unsat
(
(assume a0 (distinct y$n0s16 y$n1s16 y$n3s16))
(assume a1 (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32))
(assume a2 (= y$dve_invalid (not y$41)))
(assume a3 (= y$111 (Extract_1_15_15_16 y$v_range)))
(assume a4 (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)))
(assume a5 (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)))
(assume a6 (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)))
(assume a7 (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)))
(assume a8 (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)))
(assume a9 (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))
(assume a10 (= y$1275 (not (= y$n0s32 y$v3_1517448506_30_op))))
(assume a11 (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))
(assume a12 (= y$1278 (= y$n1s32 y$v3_1517448506_36_op)))
(assume a13 (= y$v3_1517448506_38_op (and y$1275 y$1278)))
(assume a14 (= y$id37_op (and y$41 y$v3_1517448506_38_op)))
(assume a15 (= y$id37_op (not y$1283)))
(assume a16 (= y$1284 (= y$prop y$1283)))
(assume a17 (= y$v3_1517448506_46_op (ite y$f02 y$n1s16 y$v_cage_safe)))
(assume a18 (= y$v3_1517448506_47_op (ite y$f03 y$n0s16 y$v3_1517448506_46_op)))
(assume a19 (= y$v3_1517448506_48_op (ite y$f05 y$n1s16 y$v3_1517448506_47_op)))
(assume a20 (= y$v_cage_safe$next_rhs_op (ite y$f20 y$n1s16 y$v3_1517448506_48_op)))
(assume a21 (= y$64 (= y$v_cage_safe$next y$v_cage_safe$next_rhs_op)))
(assume a22 (= y$69 (Extract_1_15_15_16 y$v_actual_resistance)))
(assume a23 (= y$w$8_op (Concat_32_16_16 y$n0s16 y$v_actual_resistance)))
(assume a24 (= y$s$203_op (BitWiseNot_32_32 y$w$8_op)))
(assume a25 (= y$v3_1517448506_57_op (ShiftR_32_32_32 y$s$203_op y$n16s32)))
(assume a26 (= y$s$202_op (BitWiseNot_32_32 y$v3_1517448506_57_op)))
(assume a27 (= y$v3_1517448506_56_op (ShiftR_32_32_32 y$w$8_op y$n16s32)))
(assume a28 (= y$v3_1517448506_55_op (ite y$69 y$s$202_op y$v3_1517448506_56_op)))
(assume a29 (= y$v3_1517448506_59_op (Sub_32_32_32 y$v3_1517448506_55_op y$n1s32)))
(assume a30 (= y$87 (Extract_16_15_0_32 y$v3_1517448506_59_op)))
(assume a31 (= y$v3_1517448506_62_op (Add_32_32_32 y$n1s32 y$v3_1517448506_55_op)))
(assume a32 (= y$92 (Extract_16_15_0_32 y$v3_1517448506_62_op)))
(assume a33 (= y$v3_1517448506_64_op (ite y$f01 y$92 y$v_actual_resistance)))
(assume a34 (= y$v3_1517448506_65_op (ite y$f02 y$87 y$v3_1517448506_64_op)))
(assume a35 (= y$v3_1517448506_66_op (ite y$f03 y$87 y$v3_1517448506_65_op)))
(assume a36 (= y$v3_1517448506_67_op (ite y$f06 y$n0s16 y$v3_1517448506_66_op)))
(assume a37 (= y$v3_1517448506_68_op (ite y$f07 y$n0s16 y$v3_1517448506_67_op)))
(assume a38 (= y$v_actual_resistance$next_rhs_op (ite y$f08 y$n0s16 y$v3_1517448506_68_op)))
(assume a39 (= y$106 (= y$v_actual_resistance$next y$v_actual_resistance$next_rhs_op)))
(assume a40 (= y$v3_1517448506_75_op (Add_32_32_32 y$n1s32 y$v3_1517448506_30_op)))
(assume a41 (= y$127 (Extract_16_15_0_32 y$v3_1517448506_75_op)))
(assume a42 (= y$v3_1517448506_78_op (Sub_32_32_32 y$v3_1517448506_30_op y$n1s32)))
(assume a43 (= y$132 (Extract_16_15_0_32 y$v3_1517448506_78_op)))
(assume a44 (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))
(assume a45 (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume a46 (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume a47 (= y$140 (= y$v_range$next y$v_range$next_rhs_op)))
(assume a48 (= y$a_measure_in_progress (not y$21)))
(assume a49 (= y$f00 (not y$143)))
(assume a50 (= y$v3_1517448506_86_op (and y$21 y$143)))
(assume a51 (= y$v3_1517448506_86_op (not y$146)))
(assume a52 (= y$146 (not y$147)))
(assume a53 (= y$f02 (not y$148)))
(assume a54 (= y$v3_1517448506_88_op (and y$147 y$148)))
(assume a55 (= y$v3_1517448506_88_op (not y$151)))
(assume a56 (= y$f03 (not y$152)))
(assume a57 (= y$v3_1517448506_89_op (and y$151 y$152)))
(assume a58 (= y$v3_1517448506_89_op (not y$155)))
(assume a59 (= y$f05 (not y$156)))
(assume a60 (= y$v3_1517448506_91_op (and y$155 y$156)))
(assume a61 (= y$v3_1517448506_91_op (not y$159)))
(assume a62 (= y$f22 (not y$161)))
(assume a63 (= y$a_measure_in_progress$next_rhs_op (and y$159 y$161)))
(assume a64 (= y$164 (= y$a_measure_in_progress$next y$a_measure_in_progress$next_rhs_op)))
(assume a65 (= y$v3_1517448506_95_op (and y$a_measure_done y$143)))
(assume a66 (= y$f04 (not y$169)))
(assume a67 (= y$v3_1517448506_97_op (and y$v3_1517448506_95_op y$169)))
(assume a68 (= y$v3_1517448506_97_op (not y$172)))
(assume a69 (= y$v3_1517448506_99_op (and y$161 y$172)))
(assume a70 (= y$v3_1517448506_99_op (not y$175)))
(assume a71 (= y$176 (= y$a_measure_done$next y$175)))
(assume a72 (= y$v3_1517448506_101_op (and y$a_cage_opened y$148)))
(assume a73 (= y$v3_1517448506_101_op (not y$180)))
(assume a74 (= y$v3_1517448506_103_op (and y$152 y$180)))
(assume a75 (= y$v3_1517448506_103_op (not y$183)))
(assume a76 (= y$183 (not y$184)))
(assume a77 (= y$v3_1517448506_105_op (and y$169 y$184)))
(assume a78 (= y$v3_1517448506_105_op (not y$187)))
(assume a79 (= y$188 (= y$a_cage_opened$next y$187)))
(assume a80 (= y$a_initial_Measuring_0 (not y$15)))
(assume a81 (= y$v3_1517448506_107_op (and y$15 y$156)))
(assume a82 (= y$v3_1517448506_107_op (not y$192)))
(assume a83 (= y$193 (= y$a_initial_Measuring_0$next y$192)))
(assume a84 (= y$f06 (not y$195)))
(assume a85 (= y$v3_1517448506_109_op (and y$a_measured y$195)))
(assume a86 (= y$f07 (not y$198)))
(assume a87 (= y$v3_1517448506_110_op (and y$v3_1517448506_109_op y$198)))
(assume a88 (= y$f08 (not y$201)))
(assume a89 (= y$v3_1517448506_111_op (and y$v3_1517448506_110_op y$201)))
(assume a90 (= y$v3_1517448506_111_op (not y$204)))
(assume a91 (= y$f24 (not y$206)))
(assume a92 (= y$v3_1517448506_114_op (and y$204 y$206)))
(assume a93 (= y$v3_1517448506_114_op (not y$209)))
(assume a94 (= y$210 (= y$a_measured$next y$209)))
(assume a95 (= y$a_new_range (not y$25)))
(assume a96 (= y$f09 (not y$212)))
(assume a97 (= y$v3_1517448506_117_op (and y$25 y$212)))
(assume a98 (= y$v3_1517448506_117_op (not y$215)))
(assume a99 (= y$215 (not y$216)))
(assume a100 (= y$f10 (not y$217)))
(assume a101 (= y$v3_1517448506_119_op (and y$216 y$217)))
(assume a102 (= y$v3_1517448506_119_op (not y$220)))
(assume a103 (= y$220 (not y$221)))
(assume a104 (= y$f11 (not y$222)))
(assume a105 (= y$v3_1517448506_121_op (and y$221 y$222)))
(assume a106 (= y$v3_1517448506_121_op (not y$225)))
(assume a107 (= y$a_new_range$next_rhs_op (and y$161 y$225)))
(assume a108 (= y$228 (= y$a_new_range$next y$a_new_range$next_rhs_op)))
(assume a109 (= y$a_small_resistance (not y$27)))
(assume a110 (= y$v3_1517448506_125_op (and y$27 y$198)))
(assume a111 (= y$v3_1517448506_125_op (not y$232)))
(assume a112 (= y$v3_1517448506_126_op (and y$212 y$232)))
(assume a113 (= y$f12 (not y$236)))
(assume a114 (= y$a_small_resistance$next_rhs_op (and y$v3_1517448506_126_op y$236)))
(assume a115 (= y$239 (= y$a_small_resistance$next y$a_small_resistance$next_rhs_op)))
(assume a116 (= y$a_correct_value (not y$7)))
(assume a117 (= y$v3_1517448506_131_op (and y$7 y$195)))
(assume a118 (= y$v3_1517448506_131_op (not y$243)))
(assume a119 (= y$f16 (not y$245)))
(assume a120 (= y$a_correct_value$next_rhs_op (and y$243 y$245)))
(assume a121 (= y$248 (= y$a_correct_value$next y$a_correct_value$next_rhs_op)))
(assume a122 (= y$a_large_resistance (not y$17)))
(assume a123 (= y$v3_1517448506_136_op (and y$17 y$201)))
(assume a124 (= y$v3_1517448506_136_op (not y$252)))
(assume a125 (= y$v3_1517448506_137_op (and y$217 y$252)))
(assume a126 (= y$f13 (not y$256)))
(assume a127 (= y$v3_1517448506_139_op (and y$v3_1517448506_137_op y$256)))
(assume a128 (= y$v3_1517448506_139_op (not y$259)))
(assume a129 (= y$f25 (not y$261)))
(assume a130 (= y$v3_1517448506_142_op (and y$259 y$261)))
(assume a131 (= y$v3_1517448506_142_op (not y$264)))
(assume a132 (= y$265 (= y$a_large_resistance$next y$264)))
(assume a133 (= y$a_too_small_resistance (not y$39)))
(assume a134 (= y$v3_1517448506_145_op (and y$39 y$236)))
(assume a135 (= y$v3_1517448506_145_op (not y$269)))
(assume a136 (= y$f14 (not y$271)))
(assume a137 (= y$a_too_small_resistance$next_rhs_op (and y$269 y$271)))
(assume a138 (= y$274 (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)))
(assume a139 (= y$a_too_large_resistance (not y$37)))
(assume a140 (= y$v3_1517448506_150_op (and y$37 y$256)))
(assume a141 (= y$v3_1517448506_150_op (not y$278)))
(assume a142 (= y$f15 (not y$280)))
(assume a143 (= y$a_too_large_resistance$next_rhs_op (and y$278 y$280)))
(assume a144 (= y$283 (= y$a_too_large_resistance$next y$a_too_large_resistance$next_rhs_op)))
(assume a145 (= y$a_start (not y$29)))
(assume a146 (= y$v3_1517448506_154_op (and y$29 y$222)))
(assume a147 (= y$v3_1517448506_154_op (not y$287)))
(assume a148 (= y$f17 (not y$289)))
(assume a149 (= y$v3_1517448506_157_op (and y$287 y$289)))
(assume a150 (= y$v3_1517448506_157_op (not y$292)))
(assume a151 (= y$292 (not y$293)))
(assume a152 (= y$294 (= y$a_start$next y$293)))
(assume a153 (= y$a_finished_measuring (not y$11)))
(assume a154 (= y$v3_1517448506_160_op (and y$11 y$271)))
(assume a155 (= y$v3_1517448506_160_op (not y$298)))
(assume a156 (= y$298 (not y$299)))
(assume a157 (= y$v3_1517448506_162_op (and y$280 y$299)))
(assume a158 (= y$v3_1517448506_162_op (not y$302)))
(assume a159 (= y$302 (not y$303)))
(assume a160 (= y$v3_1517448506_164_op (and y$245 y$303)))
(assume a161 (= y$v3_1517448506_164_op (not y$306)))
(assume a162 (= y$a_finished_measuring$next_rhs_op (and y$289 y$306)))
(assume a163 (= y$309 (= y$a_finished_measuring$next y$a_finished_measuring$next_rhs_op)))
(assume a164 (= y$a_S1 (not y$1)))
(assume a165 (= y$v3_1517448506_168_op (and y$1 y$161)))
(assume a166 (= y$v3_1517448506_168_op (not y$313)))
(assume a167 (= y$f23 (not y$315)))
(assume a168 (= y$a_S1$next_rhs_op (and y$313 y$315)))
(assume a169 (= y$318 (= y$a_S1$next y$a_S1$next_rhs_op)))
(assume a170 (= y$a_S2 (not y$3)))
(assume a171 (= y$v3_1517448506_173_op (and y$3 y$315)))
(assume a172 (= y$v3_1517448506_173_op (not y$322)))
(assume a173 (= y$v3_1517448506_174_op (and y$206 y$322)))
(assume a174 (= y$a_S2$next_rhs_op (and y$261 y$v3_1517448506_174_op)))
(assume a175 (= y$327 (= y$a_S2$next y$a_S2$next_rhs_op)))
(assume a176 (= y$f18 (not y$330)))
(assume a177 (= y$v3_1517448506_178_op (and y$a_state_request y$330)))
(assume a178 (= y$f19 (not y$334)))
(assume a179 (= y$v3_1517448506_180_op (and y$v3_1517448506_178_op y$334)))
(assume a180 (= y$v3_1517448506_180_op (not y$337)))
(assume a181 (= y$v3_1517448506_182_op (and y$315 y$337)))
(assume a182 (= y$v3_1517448506_182_op (not y$340)))
(assume a183 (= y$341 (= y$a_state_request$next y$340)))
(assume a184 (= y$a_state0 (not y$31)))
(assume a185 (= y$v3_1517448506_185_op (and y$31 y$330)))
(assume a186 (= y$v3_1517448506_185_op (not y$345)))
(assume a187 (= y$f20 (not y$346)))
(assume a188 (= y$v3_1517448506_186_op (and y$345 y$346)))
(assume a189 (= y$f21 (not y$350)))
(assume a190 (= y$a_state0$next_rhs_op (and y$v3_1517448506_186_op y$350)))
(assume a191 (= y$353 (= y$a_state0$next y$a_state0$next_rhs_op)))
(assume a192 (= y$a_err7 (not y$9)))
(assume a193 (= y$v3_1517448506_191_op (and y$9 y$334)))
(assume a194 (= y$v3_1517448506_191_op (not y$357)))
(assume a195 (= y$357 (not y$358)))
(assume a196 (= y$v3_1517448506_193_op (and y$350 y$358)))
(assume a197 (= y$v3_1517448506_193_op (not y$361)))
(assume a198 (= y$a_err7$next_rhs_op (and y$261 y$361)))
(assume a199 (= y$364 (= y$a_err7$next y$a_err7$next_rhs_op)))
(assume a200 (= y$a_initial_Device_state (not y$13)))
(assume a201 (= y$v3_1517448506_196_op (and y$13 y$315)))
(assume a202 (= y$v3_1517448506_196_op (not y$368)))
(assume a203 (= y$v3_1517448506_198_op (and y$206 y$368)))
(assume a204 (= y$v3_1517448506_198_op (not y$371)))
(assume a205 (= y$371 (not y$372)))
(assume a206 (= y$v3_1517448506_200_op (and y$261 y$372)))
(assume a207 (= y$v3_1517448506_200_op (not y$375)))
(assume a208 (= y$375 (not y$376)))
(assume a209 (= y$377 (= y$a_initial_Device_state$next y$376)))
(assume a210 (= y$a_state1 (not y$33)))
(assume a211 (= y$v3_1517448506_203_op (and y$33 y$346)))
(assume a212 (= y$v3_1517448506_203_op (not y$381)))
(assume a213 (= y$a_state1$next_rhs_op (and y$206 y$381)))
(assume a214 (= y$384 (= y$a_state1$next y$a_state1$next_rhs_op)))
(assume a215 (= y$386 (Extract_1_15_15_16 y$v_cage_safe)))
(assume a216 (= y$w$10_op (Concat_32_16_16 y$n0s16 y$v_cage_safe)))
(assume a217 (= y$s$295_op (BitWiseNot_32_32 y$w$10_op)))
(assume a218 (= y$v3_1517448506_209_op (ShiftR_32_32_32 y$s$295_op y$n16s32)))
(assume a219 (= y$s$294_op (BitWiseNot_32_32 y$v3_1517448506_209_op)))
(assume a220 (= y$v3_1517448506_208_op (ShiftR_32_32_32 y$w$10_op y$n16s32)))
(assume a221 (= y$v3_1517448506_207_op (ite y$386 y$s$294_op y$v3_1517448506_208_op)))
(assume a222 (= y$400 (= y$n1s32 y$v3_1517448506_207_op)))
(assume a223 (= y$v3_1517448506_212_op (and y$a_measure_done y$400)))
(assume a224 (= y$v3_1517448506_212_op (not y$403)))
(assume a225 (= y$v3_1517448506_214_op (and y$f00 y$403)))
(assume a226 (= y$v3_1517448506_214_op (not y$406)))
(assume a227 (= y$v3_1517448506_216_op (GrEq_1_32_32 y$v3_1517448506_55_op y$n6200s32)))
(assume a228 (= y$v3_1517448506_216_op (not y$410)))
(assume a229 (= y$v3_1517448506_217_op (and y$400 y$410)))
(assume a230 (= y$v3_1517448506_218_op (and y$a_measure_in_progress y$v3_1517448506_217_op)))
(assume a231 (= y$v3_1517448506_218_op (not y$415)))
(assume a232 (= y$v3_1517448506_220_op (and y$f01 y$415)))
(assume a233 (= y$v3_1517448506_220_op (not y$418)))
(assume a234 (= y$v3_1517448506_221_op (and y$406 y$418)))
(assume a235 (= y$a_cage_opened (not y$5)))
(assume a236 (= y$v3_1517448506_223_op (and y$5 y$f02)))
(assume a237 (= y$v3_1517448506_223_op (not y$423)))
(assume a238 (= y$v3_1517448506_224_op (and y$v3_1517448506_221_op y$423)))
(assume a239 (= y$v3_1517448506_226_op (and y$21 y$f03)))
(assume a240 (= y$v3_1517448506_226_op (not y$428)))
(assume a241 (= y$v3_1517448506_227_op (and y$v3_1517448506_224_op y$428)))
(assume a242 (= y$432 (= y$n0s32 y$v3_1517448506_207_op)))
(assume a243 (= y$v3_1517448506_229_op (and y$a_measure_done y$432)))
(assume a244 (= y$v3_1517448506_229_op (not y$435)))
(assume a245 (= y$v3_1517448506_231_op (and y$f04 y$435)))
(assume a246 (= y$v3_1517448506_231_op (not y$438)))
(assume a247 (= y$v3_1517448506_232_op (and y$v3_1517448506_227_op y$438)))
(assume a248 (= y$v3_1517448506_234_op (and y$a_initial_Measuring_0 y$f05)))
(assume a249 (= y$v3_1517448506_234_op (not y$443)))
(assume a250 (= y$v3_1517448506_235_op (and y$v3_1517448506_232_op y$443)))
(assume a251 (= y$v3_1517448506_237_op (GrEq_1_32_32 y$n999s32 y$v3_1517448506_55_op)))
(assume a252 (= y$v3_1517448506_237_op (not y$449)))
(assume a253 (= y$v3_1517448506_239_op (GrEq_1_32_32 y$v3_1517448506_55_op y$n5999s32)))
(assume a254 (= y$v3_1517448506_239_op (not y$453)))
(assume a255 (= y$v3_1517448506_240_op (and y$449 y$453)))
(assume a256 (= y$v3_1517448506_241_op (and y$a_measured y$v3_1517448506_240_op)))
(assume a257 (= y$v3_1517448506_241_op (not y$458)))
(assume a258 (= y$v3_1517448506_243_op (and y$f06 y$458)))
(assume a259 (= y$v3_1517448506_243_op (not y$461)))
(assume a260 (= y$v3_1517448506_244_op (and y$v3_1517448506_235_op y$461)))
(assume a261 (= y$v3_1517448506_246_op (GrEq_1_32_32 y$v3_1517448506_55_op y$n1000s32)))
(assume a262 (= y$v3_1517448506_246_op (not y$467)))
(assume a263 (= y$v3_1517448506_247_op (and y$a_measured y$467)))
(assume a264 (= y$v3_1517448506_247_op (not y$470)))
(assume a265 (= y$v3_1517448506_249_op (and y$f07 y$470)))
(assume a266 (= y$v3_1517448506_249_op (not y$473)))
(assume a267 (= y$v3_1517448506_250_op (and y$v3_1517448506_244_op y$473)))
(assume a268 (= y$v3_1517448506_252_op (GrEq_1_32_32 y$n5800s32 y$v3_1517448506_55_op)))
(assume a269 (= y$v3_1517448506_252_op (not y$479)))
(assume a270 (= y$v3_1517448506_253_op (and y$a_measured y$479)))
(assume a271 (= y$v3_1517448506_253_op (not y$482)))
(assume a272 (= y$v3_1517448506_255_op (and y$f08 y$482)))
(assume a273 (= y$v3_1517448506_255_op (not y$485)))
(assume a274 (= y$v3_1517448506_256_op (and y$v3_1517448506_250_op y$485)))
(assume a275 (= y$v3_1517448506_257_op (GrEq_1_32_32 y$n0s32 y$v3_1517448506_30_op)))
(assume a276 (= y$v3_1517448506_257_op (not y$490)))
(assume a277 (= y$v3_1517448506_258_op (and y$a_small_resistance y$490)))
(assume a278 (= y$v3_1517448506_258_op (not y$493)))
(assume a279 (= y$v3_1517448506_260_op (and y$f09 y$493)))
(assume a280 (= y$v3_1517448506_260_op (not y$496)))
(assume a281 (= y$v3_1517448506_261_op (and y$v3_1517448506_256_op y$496)))
(assume a282 (= y$v3_1517448506_263_op (GrEq_1_32_32 y$v3_1517448506_30_op y$n5s32)))
(assume a283 (= y$v3_1517448506_263_op (not y$502)))
(assume a284 (= y$v3_1517448506_264_op (and y$a_large_resistance y$502)))
(assume a285 (= y$v3_1517448506_264_op (not y$505)))
(assume a286 (= y$v3_1517448506_266_op (and y$f10 y$505)))
(assume a287 (= y$v3_1517448506_266_op (not y$508)))
(assume a288 (= y$v3_1517448506_267_op (and y$v3_1517448506_261_op y$508)))
(assume a289 (= y$v3_1517448506_269_op (and y$a_start y$f11)))
(assume a290 (= y$v3_1517448506_269_op (not y$513)))
(assume a291 (= y$v3_1517448506_270_op (and y$v3_1517448506_267_op y$513)))
(assume a292 (= (= y$n0s32 y$v3_1517448506_30_op) y$516))
(assume a293 (= y$v3_1517448506_271_op (and y$a_small_resistance y$516)))
(assume a294 (= y$v3_1517448506_271_op (not y$519)))
(assume a295 (= y$v3_1517448506_273_op (and y$f12 y$519)))
(assume a296 (= y$v3_1517448506_273_op (not y$522)))
(assume a297 (= y$v3_1517448506_274_op (and y$v3_1517448506_270_op y$522)))
(assume a298 (= y$525 (= y$n5s32 y$v3_1517448506_30_op)))
(assume a299 (= y$v3_1517448506_276_op (and y$a_large_resistance y$525)))
(assume a300 (= y$v3_1517448506_276_op (not y$528)))
(assume a301 (= y$v3_1517448506_278_op (and y$f13 y$528)))
(assume a302 (= y$v3_1517448506_278_op (not y$531)))
(assume a303 (= y$v3_1517448506_279_op (and y$v3_1517448506_274_op y$531)))
(assume a304 (= y$v3_1517448506_281_op (and y$39 y$f14)))
(assume a305 (= y$v3_1517448506_281_op (not y$536)))
(assume a306 (= y$v3_1517448506_282_op (and y$v3_1517448506_279_op y$536)))
(assume a307 (= y$v3_1517448506_284_op (and y$37 y$f15)))
(assume a308 (= y$v3_1517448506_284_op (not y$541)))
(assume a309 (= y$v3_1517448506_285_op (and y$v3_1517448506_282_op y$541)))
(assume a310 (= y$v3_1517448506_287_op (and y$7 y$f16)))
(assume a311 (= y$v3_1517448506_287_op (not y$546)))
(assume a312 (= y$v3_1517448506_288_op (and y$v3_1517448506_285_op y$546)))
(assume a313 (= y$v3_1517448506_290_op (and y$11 y$f17)))
(assume a314 (= y$v3_1517448506_290_op (not y$551)))
(assume a315 (= y$v3_1517448506_291_op (and y$v3_1517448506_288_op y$551)))
(assume a316 (= y$a_state_request (not y$35)))
(assume a317 (= y$v3_1517448506_293_op (and y$35 y$f18)))
(assume a318 (= y$v3_1517448506_293_op (not y$556)))
(assume a319 (= y$v3_1517448506_294_op (and y$v3_1517448506_291_op y$556)))
(assume a320 (= y$v3_1517448506_296_op (and y$35 y$f19)))
(assume a321 (= y$v3_1517448506_296_op (not y$561)))
(assume a322 (= y$v3_1517448506_297_op (and y$v3_1517448506_294_op y$561)))
(assume a323 (= y$v3_1517448506_299_op (and y$31 y$f20)))
(assume a324 (= y$v3_1517448506_299_op (not y$566)))
(assume a325 (= y$v3_1517448506_300_op (and y$v3_1517448506_297_op y$566)))
(assume a326 (= y$v3_1517448506_302_op (and y$31 y$f21)))
(assume a327 (= y$v3_1517448506_302_op (not y$571)))
(assume a328 (= y$v3_1517448506_303_op (and y$v3_1517448506_300_op y$571)))
(assume a329 (= y$v3_1517448506_304_op (and y$a_measure_in_progress y$a_new_range)))
(assume a330 (= y$v3_1517448506_305_op (and y$400 y$v3_1517448506_304_op)))
(assume a331 (= y$v3_1517448506_305_op (not y$578)))
(assume a332 (= y$v3_1517448506_307_op (and y$f22 y$578)))
(assume a333 (= y$v3_1517448506_307_op (not y$581)))
(assume a334 (= y$v3_1517448506_308_op (and y$v3_1517448506_303_op y$581)))
(assume a335 (= y$v3_1517448506_309_op (and y$a_S1 y$13)))
(assume a336 (= y$v3_1517448506_309_op (not y$586)))
(assume a337 (= y$v3_1517448506_311_op (and y$f23 y$586)))
(assume a338 (= y$v3_1517448506_311_op (not y$589)))
(assume a339 (= y$v3_1517448506_312_op (and y$v3_1517448506_308_op y$589)))
(assume a340 (= y$v3_1517448506_313_op (and y$a_S2 y$a_state1)))
(assume a341 (= y$v3_1517448506_313_op (not y$594)))
(assume a342 (= y$v3_1517448506_315_op (and y$f24 y$594)))
(assume a343 (= y$v3_1517448506_315_op (not y$597)))
(assume a344 (= y$v3_1517448506_316_op (and y$v3_1517448506_312_op y$597)))
(assume a345 (= y$v3_1517448506_317_op (and y$a_S2 y$a_err7)))
(assume a346 (= y$v3_1517448506_317_op (not y$602)))
(assume a347 (= y$v3_1517448506_319_op (and y$f25 y$602)))
(assume a348 (= y$v3_1517448506_319_op (not y$605)))
(assume a349 (= y$v3_1517448506_320_op (and y$v3_1517448506_316_op y$605)))
(assume a350 (= y$f01 (not y$608)))
(assume a351 (= y$v3_1517448506_322_op (and y$143 y$608)))
(assume a352 (= y$v3_1517448506_322_op (not y$611)))
(assume a353 (= y$611 (not y$612)))
(assume a354 (= y$v3_1517448506_324_op (and y$148 y$612)))
(assume a355 (= y$v3_1517448506_324_op (not y$615)))
(assume a356 (= y$615 (not y$616)))
(assume a357 (= y$v3_1517448506_326_op (and y$152 y$616)))
(assume a358 (= y$v3_1517448506_326_op (not y$619)))
(assume a359 (= y$619 (not y$620)))
(assume a360 (= y$v3_1517448506_328_op (and y$169 y$620)))
(assume a361 (= y$v3_1517448506_328_op (not y$623)))
(assume a362 (= y$623 (not y$624)))
(assume a363 (= y$v3_1517448506_330_op (and y$156 y$624)))
(assume a364 (= y$v3_1517448506_330_op (not y$627)))
(assume a365 (= y$627 (not y$628)))
(assume a366 (= y$v3_1517448506_332_op (and y$195 y$628)))
(assume a367 (= y$v3_1517448506_332_op (not y$631)))
(assume a368 (= y$631 (not y$632)))
(assume a369 (= y$v3_1517448506_334_op (and y$198 y$632)))
(assume a370 (= y$v3_1517448506_334_op (not y$635)))
(assume a371 (= y$635 (not y$636)))
(assume a372 (= y$v3_1517448506_336_op (and y$201 y$636)))
(assume a373 (= y$v3_1517448506_336_op (not y$639)))
(assume a374 (= y$639 (not y$640)))
(assume a375 (= y$v3_1517448506_338_op (and y$212 y$640)))
(assume a376 (= y$v3_1517448506_338_op (not y$643)))
(assume a377 (= y$643 (not y$644)))
(assume a378 (= y$v3_1517448506_340_op (and y$217 y$644)))
(assume a379 (= y$v3_1517448506_340_op (not y$647)))
(assume a380 (= y$647 (not y$648)))
(assume a381 (= y$v3_1517448506_342_op (and y$222 y$648)))
(assume a382 (= y$v3_1517448506_342_op (not y$651)))
(assume a383 (= y$651 (not y$652)))
(assume a384 (= y$v3_1517448506_344_op (and y$236 y$652)))
(assume a385 (= y$v3_1517448506_344_op (not y$655)))
(assume a386 (= y$655 (not y$656)))
(assume a387 (= y$v3_1517448506_346_op (and y$256 y$656)))
(assume a388 (= y$v3_1517448506_346_op (not y$659)))
(assume a389 (= y$659 (not y$660)))
(assume a390 (= y$v3_1517448506_348_op (and y$271 y$660)))
(assume a391 (= y$v3_1517448506_348_op (not y$663)))
(assume a392 (= y$663 (not y$664)))
(assume a393 (= y$v3_1517448506_350_op (and y$280 y$664)))
(assume a394 (= y$v3_1517448506_350_op (not y$667)))
(assume a395 (= y$667 (not y$668)))
(assume a396 (= y$v3_1517448506_352_op (and y$245 y$668)))
(assume a397 (= y$v3_1517448506_352_op (not y$671)))
(assume a398 (= y$671 (not y$672)))
(assume a399 (= y$v3_1517448506_354_op (and y$289 y$672)))
(assume a400 (= y$v3_1517448506_354_op (not y$675)))
(assume a401 (= y$675 (not y$676)))
(assume a402 (= y$v3_1517448506_356_op (and y$330 y$676)))
(assume a403 (= y$v3_1517448506_356_op (not y$679)))
(assume a404 (= y$679 (not y$680)))
(assume a405 (= y$v3_1517448506_358_op (and y$334 y$680)))
(assume a406 (= y$v3_1517448506_358_op (not y$683)))
(assume a407 (= y$683 (not y$684)))
(assume a408 (= y$v3_1517448506_360_op (and y$346 y$684)))
(assume a409 (= y$v3_1517448506_360_op (not y$687)))
(assume a410 (= y$687 (not y$688)))
(assume a411 (= y$v3_1517448506_362_op (and y$350 y$688)))
(assume a412 (= y$v3_1517448506_362_op (not y$691)))
(assume a413 (= y$691 (not y$692)))
(assume a414 (= y$v3_1517448506_364_op (and y$161 y$692)))
(assume a415 (= y$v3_1517448506_364_op (not y$695)))
(assume a416 (= y$695 (not y$696)))
(assume a417 (= y$v3_1517448506_366_op (and y$315 y$696)))
(assume a418 (= y$v3_1517448506_366_op (not y$699)))
(assume a419 (= y$699 (not y$700)))
(assume a420 (= y$v3_1517448506_368_op (and y$206 y$700)))
(assume a421 (= y$v3_1517448506_368_op (not y$703)))
(assume a422 (= y$703 (not y$704)))
(assume a423 (= y$v3_1517448506_370_op (and y$261 y$704)))
(assume a424 (= y$v3_1517448506_370_op (not y$707)))
(assume a425 (= y$v3_1517448506_371_op (and y$v3_1517448506_320_op y$707)))
(assume a426 (= y$v3_1517448506_372_op (and y$f01 y$f00)))
(assume a427 (= y$v3_1517448506_372_op (not y$712)))
(assume a428 (= y$v3_1517448506_373_op (and y$f02 y$611)))
(assume a429 (= y$v3_1517448506_373_op (not y$715)))
(assume a430 (= y$v3_1517448506_375_op (and y$712 y$715)))
(assume a431 (= y$v3_1517448506_375_op (not y$718)))
(assume a432 (= y$718 (not y$719)))
(assume a433 (= y$v3_1517448506_376_op (and y$f03 y$615)))
(assume a434 (= y$v3_1517448506_376_op (not y$722)))
(assume a435 (= y$v3_1517448506_378_op (and y$719 y$722)))
(assume a436 (= y$v3_1517448506_378_op (not y$725)))
(assume a437 (= y$725 (not y$726)))
(assume a438 (= y$v3_1517448506_379_op (and y$f04 y$619)))
(assume a439 (= y$v3_1517448506_379_op (not y$729)))
(assume a440 (= y$v3_1517448506_381_op (and y$726 y$729)))
(assume a441 (= y$v3_1517448506_381_op (not y$732)))
(assume a442 (= y$732 (not y$733)))
(assume a443 (= y$v3_1517448506_382_op (and y$f05 y$623)))
(assume a444 (= y$v3_1517448506_382_op (not y$736)))
(assume a445 (= y$v3_1517448506_384_op (and y$733 y$736)))
(assume a446 (= y$v3_1517448506_384_op (not y$739)))
(assume a447 (= y$739 (not y$740)))
(assume a448 (= y$v3_1517448506_385_op (and y$f06 y$627)))
(assume a449 (= y$v3_1517448506_385_op (not y$743)))
(assume a450 (= y$v3_1517448506_387_op (and y$740 y$743)))
(assume a451 (= y$v3_1517448506_387_op (not y$746)))
(assume a452 (= y$746 (not y$747)))
(assume a453 (= y$v3_1517448506_388_op (and y$f07 y$631)))
(assume a454 (= y$v3_1517448506_388_op (not y$750)))
(assume a455 (= y$v3_1517448506_390_op (and y$747 y$750)))
(assume a456 (= y$v3_1517448506_390_op (not y$753)))
(assume a457 (= y$753 (not y$754)))
(assume a458 (= y$v3_1517448506_391_op (and y$f08 y$635)))
(assume a459 (= y$v3_1517448506_391_op (not y$757)))
(assume a460 (= y$v3_1517448506_393_op (and y$754 y$757)))
(assume a461 (= y$v3_1517448506_393_op (not y$760)))
(assume a462 (= y$760 (not y$761)))
(assume a463 (= y$v3_1517448506_394_op (and y$f09 y$639)))
(assume a464 (= y$v3_1517448506_394_op (not y$764)))
(assume a465 (= y$v3_1517448506_396_op (and y$761 y$764)))
(assume a466 (= y$v3_1517448506_396_op (not y$767)))
(assume a467 (= y$767 (not y$768)))
(assume a468 (= y$v3_1517448506_397_op (and y$f10 y$643)))
(assume a469 (= y$v3_1517448506_397_op (not y$771)))
(assume a470 (= y$v3_1517448506_399_op (and y$768 y$771)))
(assume a471 (= y$v3_1517448506_399_op (not y$774)))
(assume a472 (= y$774 (not y$775)))
(assume a473 (= y$v3_1517448506_400_op (and y$f11 y$647)))
(assume a474 (= y$v3_1517448506_400_op (not y$778)))
(assume a475 (= y$v3_1517448506_402_op (and y$775 y$778)))
(assume a476 (= y$v3_1517448506_402_op (not y$781)))
(assume a477 (= y$781 (not y$782)))
(assume a478 (= y$v3_1517448506_403_op (and y$f12 y$651)))
(assume a479 (= y$v3_1517448506_403_op (not y$785)))
(assume a480 (= y$v3_1517448506_405_op (and y$782 y$785)))
(assume a481 (= y$v3_1517448506_405_op (not y$788)))
(assume a482 (= y$788 (not y$789)))
(assume a483 (= y$v3_1517448506_406_op (and y$f13 y$655)))
(assume a484 (= y$v3_1517448506_406_op (not y$792)))
(assume a485 (= y$v3_1517448506_408_op (and y$789 y$792)))
(assume a486 (= y$v3_1517448506_408_op (not y$795)))
(assume a487 (= y$795 (not y$796)))
(assume a488 (= y$v3_1517448506_409_op (and y$f14 y$659)))
(assume a489 (= y$v3_1517448506_409_op (not y$799)))
(assume a490 (= y$v3_1517448506_411_op (and y$796 y$799)))
(assume a491 (= y$v3_1517448506_411_op (not y$802)))
(assume a492 (= y$802 (not y$803)))
(assume a493 (= y$v3_1517448506_412_op (and y$f15 y$663)))
(assume a494 (= y$v3_1517448506_412_op (not y$806)))
(assume a495 (= y$v3_1517448506_414_op (and y$803 y$806)))
(assume a496 (= y$v3_1517448506_414_op (not y$809)))
(assume a497 (= y$809 (not y$810)))
(assume a498 (= y$v3_1517448506_415_op (and y$f16 y$667)))
(assume a499 (= y$v3_1517448506_415_op (not y$813)))
(assume a500 (= y$v3_1517448506_417_op (and y$810 y$813)))
(assume a501 (= y$v3_1517448506_417_op (not y$816)))
(assume a502 (= y$816 (not y$817)))
(assume a503 (= y$v3_1517448506_418_op (and y$f17 y$671)))
(assume a504 (= y$v3_1517448506_418_op (not y$820)))
(assume a505 (= y$v3_1517448506_420_op (and y$817 y$820)))
(assume a506 (= y$v3_1517448506_420_op (not y$823)))
(assume a507 (= y$823 (not y$824)))
(assume a508 (= y$v3_1517448506_421_op (and y$f18 y$675)))
(assume a509 (= y$v3_1517448506_421_op (not y$827)))
(assume a510 (= y$v3_1517448506_423_op (and y$824 y$827)))
(assume a511 (= y$v3_1517448506_423_op (not y$830)))
(assume a512 (= y$830 (not y$831)))
(assume a513 (= y$v3_1517448506_424_op (and y$f19 y$679)))
(assume a514 (= y$v3_1517448506_424_op (not y$834)))
(assume a515 (= y$v3_1517448506_426_op (and y$831 y$834)))
(assume a516 (= y$v3_1517448506_426_op (not y$837)))
(assume a517 (= y$837 (not y$838)))
(assume a518 (= y$v3_1517448506_427_op (and y$f20 y$683)))
(assume a519 (= y$v3_1517448506_427_op (not y$841)))
(assume a520 (= y$v3_1517448506_429_op (and y$838 y$841)))
(assume a521 (= y$v3_1517448506_429_op (not y$844)))
(assume a522 (= y$844 (not y$845)))
(assume a523 (= y$v3_1517448506_430_op (and y$f21 y$687)))
(assume a524 (= y$v3_1517448506_430_op (not y$848)))
(assume a525 (= y$v3_1517448506_432_op (and y$845 y$848)))
(assume a526 (= y$v3_1517448506_432_op (not y$851)))
(assume a527 (= y$851 (not y$852)))
(assume a528 (= y$v3_1517448506_433_op (and y$f22 y$691)))
(assume a529 (= y$v3_1517448506_433_op (not y$855)))
(assume a530 (= y$v3_1517448506_435_op (and y$852 y$855)))
(assume a531 (= y$v3_1517448506_435_op (not y$858)))
(assume a532 (= y$858 (not y$859)))
(assume a533 (= y$v3_1517448506_436_op (and y$f23 y$695)))
(assume a534 (= y$v3_1517448506_436_op (not y$862)))
(assume a535 (= y$v3_1517448506_438_op (and y$859 y$862)))
(assume a536 (= y$v3_1517448506_438_op (not y$865)))
(assume a537 (= y$865 (not y$866)))
(assume a538 (= y$v3_1517448506_439_op (and y$f24 y$699)))
(assume a539 (= y$v3_1517448506_439_op (not y$869)))
(assume a540 (= y$v3_1517448506_441_op (and y$866 y$869)))
(assume a541 (= y$v3_1517448506_441_op (not y$872)))
(assume a542 (= y$872 (not y$873)))
(assume a543 (= y$v3_1517448506_442_op (and y$f25 y$703)))
(assume a544 (= y$v3_1517448506_442_op (not y$876)))
(assume a545 (= y$v3_1517448506_444_op (and y$873 y$876)))
(assume a546 (= y$v3_1517448506_444_op (not y$879)))
(assume a547 (= y$879 (not y$880)))
(assume a548 (= y$v3_1517448506_445_op (and y$v3_1517448506_371_op y$880)))
(assume a549 (= y$v3_1517448506_446_op (and y$a_measure_done y$a_measure_in_progress)))
(assume a550 (= y$v3_1517448506_446_op (not y$885)))
(assume a551 (= y$a_measure_done (not y$19)))
(assume a552 (= y$v3_1517448506_448_op (and y$19 y$21)))
(assume a553 (= y$v3_1517448506_448_op (not y$888)))
(assume a554 (= y$v3_1517448506_449_op (and y$a_cage_opened y$888)))
(assume a555 (= y$v3_1517448506_449_op (not y$891)))
(assume a556 (= y$v3_1517448506_451_op (and y$885 y$891)))
(assume a557 (= y$v3_1517448506_451_op (not y$894)))
(assume a558 (= y$894 (not y$895)))
(assume a559 (= y$888 (not y$896)))
(assume a560 (= y$v3_1517448506_453_op (and y$5 y$896)))
(assume a561 (= y$v3_1517448506_453_op (not y$899)))
(assume a562 (= y$v3_1517448506_454_op (and y$15 y$899)))
(assume a563 (= y$v3_1517448506_454_op (not y$902)))
(assume a564 (= y$v3_1517448506_456_op (and y$895 y$902)))
(assume a565 (= y$v3_1517448506_456_op (not y$905)))
(assume a566 (= y$905 (not y$906)))
(assume a567 (= y$899 (not y$907)))
(assume a568 (= y$v3_1517448506_458_op (and y$a_initial_Measuring_0 y$907)))
(assume a569 (= y$v3_1517448506_458_op (not y$910)))
(assume a570 (= y$v3_1517448506_459_op (and y$906 y$910)))
(assume a571 (= y$v3_1517448506_460_op (and y$a_measured y$a_new_range)))
(assume a572 (= y$v3_1517448506_460_op (not y$915)))
(assume a573 (= y$a_measured (not y$23)))
(assume a574 (= y$v3_1517448506_462_op (and y$23 y$25)))
(assume a575 (= y$v3_1517448506_462_op (not y$918)))
(assume a576 (= y$v3_1517448506_463_op (and y$a_small_resistance y$918)))
(assume a577 (= y$v3_1517448506_463_op (not y$921)))
(assume a578 (= y$v3_1517448506_465_op (and y$915 y$921)))
(assume a579 (= y$v3_1517448506_465_op (not y$924)))
(assume a580 (= y$924 (not y$925)))
(assume a581 (= y$918 (not y$926)))
(assume a582 (= y$v3_1517448506_467_op (and y$27 y$926)))
(assume a583 (= y$v3_1517448506_467_op (not y$929)))
(assume a584 (= y$v3_1517448506_468_op (and y$a_correct_value y$929)))
(assume a585 (= y$v3_1517448506_468_op (not y$932)))
(assume a586 (= y$v3_1517448506_470_op (and y$925 y$932)))
(assume a587 (= y$v3_1517448506_470_op (not y$935)))
(assume a588 (= y$935 (not y$936)))
(assume a589 (= y$929 (not y$937)))
(assume a590 (= y$v3_1517448506_472_op (and y$7 y$937)))
(assume a591 (= y$v3_1517448506_472_op (not y$940)))
(assume a592 (= y$v3_1517448506_473_op (and y$a_large_resistance y$940)))
(assume a593 (= y$v3_1517448506_473_op (not y$943)))
(assume a594 (= y$v3_1517448506_475_op (and y$936 y$943)))
(assume a595 (= y$v3_1517448506_475_op (not y$946)))
(assume a596 (= y$946 (not y$947)))
(assume a597 (= y$940 (not y$948)))
(assume a598 (= y$v3_1517448506_477_op (and y$17 y$948)))
(assume a599 (= y$v3_1517448506_477_op (not y$951)))
(assume a600 (= y$v3_1517448506_478_op (and y$a_too_small_resistance y$951)))
(assume a601 (= y$v3_1517448506_478_op (not y$954)))
(assume a602 (= y$v3_1517448506_480_op (and y$947 y$954)))
(assume a603 (= y$v3_1517448506_480_op (not y$957)))
(assume a604 (= y$957 (not y$958)))
(assume a605 (= y$951 (not y$959)))
(assume a606 (= y$v3_1517448506_482_op (and y$39 y$959)))
(assume a607 (= y$v3_1517448506_482_op (not y$962)))
(assume a608 (= y$v3_1517448506_483_op (and y$a_too_large_resistance y$962)))
(assume a609 (= y$v3_1517448506_483_op (not y$965)))
(assume a610 (= y$v3_1517448506_485_op (and y$958 y$965)))
(assume a611 (= y$v3_1517448506_485_op (not y$968)))
(assume a612 (= y$968 (not y$969)))
(assume a613 (= y$962 (not y$970)))
(assume a614 (= y$v3_1517448506_487_op (and y$37 y$970)))
(assume a615 (= y$v3_1517448506_487_op (not y$973)))
(assume a616 (= y$v3_1517448506_488_op (and y$29 y$973)))
(assume a617 (= y$v3_1517448506_488_op (not y$976)))
(assume a618 (= y$v3_1517448506_490_op (and y$969 y$976)))
(assume a619 (= y$v3_1517448506_490_op (not y$979)))
(assume a620 (= y$979 (not y$980)))
(assume a621 (= y$973 (not y$981)))
(assume a622 (= y$v3_1517448506_492_op (and y$a_start y$981)))
(assume a623 (= y$v3_1517448506_492_op (not y$984)))
(assume a624 (= y$v3_1517448506_493_op (and y$a_finished_measuring y$984)))
(assume a625 (= y$v3_1517448506_493_op (not y$987)))
(assume a626 (= y$v3_1517448506_495_op (and y$980 y$987)))
(assume a627 (= y$v3_1517448506_495_op (not y$990)))
(assume a628 (= y$990 (not y$991)))
(assume a629 (= y$984 (not y$992)))
(assume a630 (= y$v3_1517448506_497_op (and y$11 y$992)))
(assume a631 (= y$v3_1517448506_497_op (not y$995)))
(assume a632 (= y$v3_1517448506_498_op (and y$a_S1 y$995)))
(assume a633 (= y$v3_1517448506_498_op (not y$998)))
(assume a634 (= y$v3_1517448506_500_op (and y$991 y$998)))
(assume a635 (= y$v3_1517448506_500_op (not y$1001)))
(assume a636 (= y$1001 (not y$1002)))
(assume a637 (= y$995 (not y$1003)))
(assume a638 (= y$v3_1517448506_502_op (and y$1 y$1003)))
(assume a639 (= y$v3_1517448506_502_op (not y$1006)))
(assume a640 (= y$v3_1517448506_503_op (and y$a_S2 y$1006)))
(assume a641 (= y$v3_1517448506_503_op (not y$1009)))
(assume a642 (= y$v3_1517448506_505_op (and y$1002 y$1009)))
(assume a643 (= y$v3_1517448506_505_op (not y$1012)))
(assume a644 (= y$1012 (not y$1013)))
(assume a645 (= y$v3_1517448506_506_op (and y$v3_1517448506_459_op y$1013)))
(assume a646 (= y$1006 (not y$1016)))
(assume a647 (= y$v3_1517448506_508_op (and y$3 y$1016)))
(assume a648 (= y$v3_1517448506_508_op (not y$1019)))
(assume a649 (= y$v3_1517448506_509_op (and y$v3_1517448506_506_op y$1019)))
(assume a650 (= y$v3_1517448506_510_op (and y$a_state0 y$a_state_request)))
(assume a651 (= y$v3_1517448506_510_op (not y$1024)))
(assume a652 (= y$v3_1517448506_512_op (and y$31 y$35)))
(assume a653 (= y$v3_1517448506_512_op (not y$1027)))
(assume a654 (= y$v3_1517448506_513_op (and y$a_err7 y$1027)))
(assume a655 (= y$v3_1517448506_513_op (not y$1030)))
(assume a656 (= y$v3_1517448506_515_op (and y$1024 y$1030)))
(assume a657 (= y$v3_1517448506_515_op (not y$1033)))
(assume a658 (= y$1033 (not y$1034)))
(assume a659 (= y$1027 (not y$1035)))
(assume a660 (= y$v3_1517448506_517_op (and y$9 y$1035)))
(assume a661 (= y$v3_1517448506_517_op (not y$1038)))
(assume a662 (= y$v3_1517448506_518_op (and y$13 y$1038)))
(assume a663 (= y$v3_1517448506_518_op (not y$1041)))
(assume a664 (= y$v3_1517448506_520_op (and y$1034 y$1041)))
(assume a665 (= y$v3_1517448506_520_op (not y$1044)))
(assume a666 (= y$1044 (not y$1045)))
(assume a667 (= y$1038 (not y$1046)))
(assume a668 (= y$v3_1517448506_522_op (and y$a_initial_Device_state y$1046)))
(assume a669 (= y$v3_1517448506_522_op (not y$1049)))
(assume a670 (= y$v3_1517448506_523_op (and y$a_state1 y$1049)))
(assume a671 (= y$v3_1517448506_523_op (not y$1052)))
(assume a672 (= y$v3_1517448506_525_op (and y$1045 y$1052)))
(assume a673 (= y$v3_1517448506_525_op (not y$1055)))
(assume a674 (= y$1055 (not y$1056)))
(assume a675 (= y$v3_1517448506_526_op (and y$v3_1517448506_509_op y$1056)))
(assume a676 (= y$1049 (not y$1059)))
(assume a677 (= y$v3_1517448506_528_op (and y$33 y$1059)))
(assume a678 (= y$v3_1517448506_528_op (not y$1062)))
(assume a679 (= y$v3_1517448506_529_op (and y$v3_1517448506_526_op y$1062)))
(assume a680 (= y$v3_1517448506_530_op (and y$v3_1517448506_445_op y$v3_1517448506_529_op)))
(assume a681 (= y$v3_1517448506_531_op (and y$a_measure_in_progress$next_rhs_op y$175)))
(assume a682 (= y$v3_1517448506_531_op (not y$1069)))
(assume a683 (= y$a_measure_in_progress$next_rhs_op (not y$1070)))
(assume a684 (= y$175 (not y$1071)))
(assume a685 (= y$v3_1517448506_533_op (and y$1070 y$1071)))
(assume a686 (= y$v3_1517448506_533_op (not y$1074)))
(assume a687 (= y$v3_1517448506_534_op (and y$187 y$1074)))
(assume a688 (= y$v3_1517448506_534_op (not y$1077)))
(assume a689 (= y$v3_1517448506_536_op (and y$1069 y$1077)))
(assume a690 (= y$v3_1517448506_536_op (not y$1080)))
(assume a691 (= y$1080 (not y$1081)))
(assume a692 (= y$187 (not y$1082)))
(assume a693 (= y$1074 (not y$1083)))
(assume a694 (= y$v3_1517448506_538_op (and y$1082 y$1083)))
(assume a695 (= y$v3_1517448506_538_op (not y$1086)))
(assume a696 (= y$v3_1517448506_539_op (and y$v3_1517448506_107_op y$1086)))
(assume a697 (= y$v3_1517448506_539_op (not y$1089)))
(assume a698 (= y$v3_1517448506_541_op (and y$1081 y$1089)))
(assume a699 (= y$v3_1517448506_541_op (not y$1092)))
(assume a700 (= y$1092 (not y$1093)))
(assume a701 (= y$1086 (not y$1094)))
(assume a702 (= y$v3_1517448506_543_op (and y$192 y$1094)))
(assume a703 (= y$v3_1517448506_543_op (not y$1097)))
(assume a704 (= y$v3_1517448506_544_op (and y$1093 y$1097)))
(assume a705 (= y$v3_1517448506_545_op (and y$209 y$a_new_range$next_rhs_op)))
(assume a706 (= y$v3_1517448506_545_op (not y$1102)))
(assume a707 (= y$a_new_range$next_rhs_op (not y$1103)))
(assume a708 (= y$209 (not y$1104)))
(assume a709 (= y$v3_1517448506_547_op (and y$1103 y$1104)))
(assume a710 (= y$v3_1517448506_547_op (not y$1107)))
(assume a711 (= y$v3_1517448506_548_op (and y$a_small_resistance$next_rhs_op y$1107)))
(assume a712 (= y$v3_1517448506_548_op (not y$1110)))
(assume a713 (= y$v3_1517448506_550_op (and y$1102 y$1110)))
(assume a714 (= y$v3_1517448506_550_op (not y$1113)))
(assume a715 (= y$1113 (not y$1114)))
(assume a716 (= y$a_small_resistance$next_rhs_op (not y$1115)))
(assume a717 (= y$1107 (not y$1116)))
(assume a718 (= y$v3_1517448506_552_op (and y$1115 y$1116)))
(assume a719 (= y$v3_1517448506_552_op (not y$1119)))
(assume a720 (= y$v3_1517448506_553_op (and y$a_correct_value$next_rhs_op y$1119)))
(assume a721 (= y$v3_1517448506_553_op (not y$1122)))
(assume a722 (= y$v3_1517448506_555_op (and y$1114 y$1122)))
(assume a723 (= y$v3_1517448506_555_op (not y$1125)))
(assume a724 (= y$1125 (not y$1126)))
(assume a725 (= y$a_correct_value$next_rhs_op (not y$1127)))
(assume a726 (= y$1119 (not y$1128)))
(assume a727 (= y$v3_1517448506_557_op (and y$1127 y$1128)))
(assume a728 (= y$v3_1517448506_557_op (not y$1131)))
(assume a729 (= y$v3_1517448506_558_op (and y$264 y$1131)))
(assume a730 (= y$v3_1517448506_558_op (not y$1134)))
(assume a731 (= y$v3_1517448506_560_op (and y$1126 y$1134)))
(assume a732 (= y$v3_1517448506_560_op (not y$1137)))
(assume a733 (= y$1137 (not y$1138)))
(assume a734 (= y$264 (not y$1139)))
(assume a735 (= y$1131 (not y$1140)))
(assume a736 (= y$v3_1517448506_562_op (and y$1139 y$1140)))
(assume a737 (= y$v3_1517448506_562_op (not y$1143)))
(assume a738 (= y$v3_1517448506_563_op (and y$a_too_small_resistance$next_rhs_op y$1143)))
(assume a739 (= y$v3_1517448506_563_op (not y$1146)))
(assume a740 (= y$v3_1517448506_565_op (and y$1138 y$1146)))
(assume a741 (= y$v3_1517448506_565_op (not y$1149)))
(assume a742 (= y$1149 (not y$1150)))
(assume a743 (= y$a_too_small_resistance$next_rhs_op (not y$1151)))
(assume a744 (= y$1143 (not y$1152)))
(assume a745 (= y$v3_1517448506_567_op (and y$1151 y$1152)))
(assume a746 (= y$v3_1517448506_567_op (not y$1155)))
(assume a747 (= y$v3_1517448506_568_op (and y$a_too_large_resistance$next_rhs_op y$1155)))
(assume a748 (= y$v3_1517448506_568_op (not y$1158)))
(assume a749 (= y$v3_1517448506_570_op (and y$1150 y$1158)))
(assume a750 (= y$v3_1517448506_570_op (not y$1161)))
(assume a751 (= y$1161 (not y$1162)))
(assume a752 (= y$a_too_large_resistance$next_rhs_op (not y$1163)))
(assume a753 (= y$1155 (not y$1164)))
(assume a754 (= y$v3_1517448506_572_op (and y$1163 y$1164)))
(assume a755 (= y$v3_1517448506_572_op (not y$1167)))
(assume a756 (= y$v3_1517448506_573_op (and y$292 y$1167)))
(assume a757 (= y$v3_1517448506_573_op (not y$1170)))
(assume a758 (= y$v3_1517448506_575_op (and y$1162 y$1170)))
(assume a759 (= y$v3_1517448506_575_op (not y$1173)))
(assume a760 (= y$1173 (not y$1174)))
(assume a761 (= y$1167 (not y$1175)))
(assume a762 (= y$v3_1517448506_577_op (and y$293 y$1175)))
(assume a763 (= y$v3_1517448506_577_op (not y$1178)))
(assume a764 (= y$v3_1517448506_578_op (and y$a_finished_measuring$next_rhs_op y$1178)))
(assume a765 (= y$v3_1517448506_578_op (not y$1181)))
(assume a766 (= y$v3_1517448506_580_op (and y$1174 y$1181)))
(assume a767 (= y$v3_1517448506_580_op (not y$1184)))
(assume a768 (= y$1184 (not y$1185)))
(assume a769 (= y$a_finished_measuring$next_rhs_op (not y$1186)))
(assume a770 (= y$1178 (not y$1187)))
(assume a771 (= y$v3_1517448506_582_op (and y$1186 y$1187)))
(assume a772 (= y$v3_1517448506_582_op (not y$1190)))
(assume a773 (= y$v3_1517448506_583_op (and y$a_S1$next_rhs_op y$1190)))
(assume a774 (= y$v3_1517448506_583_op (not y$1193)))
(assume a775 (= y$v3_1517448506_585_op (and y$1185 y$1193)))
(assume a776 (= y$v3_1517448506_585_op (not y$1196)))
(assume a777 (= y$1196 (not y$1197)))
(assume a778 (= y$a_S1$next_rhs_op (not y$1198)))
(assume a779 (= y$1190 (not y$1199)))
(assume a780 (= y$v3_1517448506_587_op (and y$1198 y$1199)))
(assume a781 (= y$v3_1517448506_587_op (not y$1202)))
(assume a782 (= y$v3_1517448506_588_op (and y$a_S2$next_rhs_op y$1202)))
(assume a783 (= y$v3_1517448506_588_op (not y$1205)))
(assume a784 (= y$v3_1517448506_590_op (and y$1197 y$1205)))
(assume a785 (= y$v3_1517448506_590_op (not y$1208)))
(assume a786 (= y$1208 (not y$1209)))
(assume a787 (= y$v3_1517448506_591_op (and y$v3_1517448506_544_op y$1209)))
(assume a788 (= y$a_S2$next_rhs_op (not y$1212)))
(assume a789 (= y$1202 (not y$1213)))
(assume a790 (= y$v3_1517448506_593_op (and y$1212 y$1213)))
(assume a791 (= y$v3_1517448506_593_op (not y$1216)))
(assume a792 (= y$v3_1517448506_594_op (and y$v3_1517448506_591_op y$1216)))
(assume a793 (= y$v3_1517448506_595_op (and y$340 y$a_state0$next_rhs_op)))
(assume a794 (= y$v3_1517448506_595_op (not y$1221)))
(assume a795 (= y$a_state0$next_rhs_op (not y$1222)))
(assume a796 (= y$340 (not y$1223)))
(assume a797 (= y$v3_1517448506_597_op (and y$1222 y$1223)))
(assume a798 (= y$v3_1517448506_597_op (not y$1226)))
(assume a799 (= y$v3_1517448506_598_op (and y$a_err7$next_rhs_op y$1226)))
(assume a800 (= y$v3_1517448506_598_op (not y$1229)))
(assume a801 (= y$v3_1517448506_600_op (and y$1221 y$1229)))
(assume a802 (= y$v3_1517448506_600_op (not y$1232)))
(assume a803 (= y$1232 (not y$1233)))
(assume a804 (= y$a_err7$next_rhs_op (not y$1234)))
(assume a805 (= y$1226 (not y$1235)))
(assume a806 (= y$v3_1517448506_602_op (and y$1234 y$1235)))
(assume a807 (= y$v3_1517448506_602_op (not y$1238)))
(assume a808 (= y$v3_1517448506_603_op (and y$375 y$1238)))
(assume a809 (= y$v3_1517448506_603_op (not y$1241)))
(assume a810 (= y$v3_1517448506_605_op (and y$1233 y$1241)))
(assume a811 (= y$v3_1517448506_605_op (not y$1244)))
(assume a812 (= y$1244 (not y$1245)))
(assume a813 (= y$1238 (not y$1246)))
(assume a814 (= y$v3_1517448506_607_op (and y$376 y$1246)))
(assume a815 (= y$v3_1517448506_607_op (not y$1249)))
(assume a816 (= y$v3_1517448506_608_op (and y$a_state1$next_rhs_op y$1249)))
(assume a817 (= y$v3_1517448506_608_op (not y$1252)))
(assume a818 (= y$v3_1517448506_610_op (and y$1245 y$1252)))
(assume a819 (= y$v3_1517448506_610_op (not y$1255)))
(assume a820 (= y$1255 (not y$1256)))
(assume a821 (= y$v3_1517448506_611_op (and y$v3_1517448506_594_op y$1256)))
(assume a822 (= y$a_state1$next_rhs_op (not y$1259)))
(assume a823 (= y$1249 (not y$1260)))
(assume a824 (= y$v3_1517448506_613_op (and y$1259 y$1260)))
(assume a825 (= y$v3_1517448506_613_op (not y$1263)))
(assume a826 (= y$v3_1517448506_614_op (and y$v3_1517448506_611_op y$1263)))
(assume a827 (= y$v3_1517448506_615_op (and y$v3_1517448506_530_op y$v3_1517448506_614_op)))
(assume a828 (= y$v3_1517448506_616_op (and y$41 y$v3_1517448506_615_op)))
(assume a829 (= y$v3_1517448506_616_op (not y$1270)))
(assume a830 (= y$1271 (= y$dve_invalid$next y$1270)))
(assume a831 (= y$1272 (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)))
(assume a832 (= y$dve_invalid$next (not y$1319)))
(assume a833 (= y$1320 (Extract_1_15_15_16 y$v_range$next)))
(assume a834 (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)))
(assume a835 (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)))
(assume a836 (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)))
(assume a837 (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)))
(assume a838 (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))
(assume a839 (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))
(assume a840 (= y$1334 (not (= y$n0s32 y$v3_1517448506_30$next_op))))
(assume a841 (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))
(assume a842 (= y$1337 (= y$n1s32 y$v3_1517448506_36$next_op)))
(assume a843 (= y$v3_1517448506_38$next_op (and y$1334 y$1337)))
(assume a844 (= y$id37$next_op (and y$1319 y$v3_1517448506_38$next_op)))
(assume a845 (= y$id37$next_op (not y$1342)))
(assume a846 (= y$1343 (= y$prop$next y$1342)))
(assume a847 (= y$prop$next (not y$1318)))
(assume a848 (= y$1346 (and y$prop y$1284 y$1272 y$1343 y$1318)))
(assume a849 y$1346)
(step t0 (cl (not (= (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op)) (ite y$1320 (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) (not (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op))) (ite y$1320 (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule equiv_pos2)
(step t1 (cl (= y$1320 y$1320)) :rule refl)
(step t2 (cl (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op))
(step t3 (cl (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op))
(step t4 (cl (= (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op)) (ite y$1320 (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule cong :premises (t1 t2 t3))
(step t5 (cl (not (= (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op)) (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op)))) (not (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op))) (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op))) :rule equiv_pos2)
(step t6 (cl (= (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) true) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule equiv_simplify)
(step t7 (cl (not (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) true)) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule equiv1 :premises (t6))
(step t8 (cl (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule refl)
(step t9 (cl (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule refl)
(step t10 (cl (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule cong :premises (t8 t9))
(step t11 (cl (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))
(step t12 (cl (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) true)) :rule trans :premises (t10 t11))
(step t13 (cl (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule resolution :premises (t7 t12))
(step t14 (cl (= y$s$198$next_op y$s$198$next_op)) :rule refl)
(step t15 (cl (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op))) :rule cong :premises (t13 t14))
(step t16 (cl (= y$v3_1517448506_31$next_op y$v3_1517448506_31$next_op)) :rule refl)
(step t17 (cl (= (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op))) :rule cong :premises (t13 t16))
(step t18 (cl (= (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op)) (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op)))) :rule cong :premises (t1 t15 t17))
(step t19 (cl (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op))) :rule hole)
(step t20 (cl (ite y$1320 (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op) (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op))) :rule resolution :premises (t5 t18 t19))
(step t21 (cl (ite y$1320 (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule resolution :premises (t0 t4 t20))
(step t22 (cl (not y$1320) (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule ite2 :premises (t21))
(step t23 (cl (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (not y$1320)) :rule reordering :premises (t22))
(step t24 (cl y$1320 (not (Extract_1_15_15_16 y$v_range$next))) :rule equiv2 :premises (a833))
(step t25 (cl (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op)) (not (Extract_1_15_15_16 y$v_range)) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op))) :rule and_neg)
(step t26 (cl (=> (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op)) (Extract_1_15_15_16 y$v_range$next)) (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) :rule implies_neg1)
(anchor :step t27)
(assume t27.a0 (Extract_1_15_15_16 y$v_range))
(assume t27.a1 (= y$v_range (ite y$f09 y$132 y$v_range)))
(assume t27.a2 (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))
(assume t27.a3 (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t27.a4 (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t27.a5 (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t27.a6 (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t27.a7 (= y$v_range$next y$v_range$next_rhs_op))
(step t27.t0 (cl (= (= (Extract_1_15_15_16 y$v_range$next) true) (Extract_1_15_15_16 y$v_range$next))) :rule equiv_simplify)
(step t27.t1 (cl (not (= (Extract_1_15_15_16 y$v_range$next) true)) (Extract_1_15_15_16 y$v_range$next)) :rule equiv1 :premises (t27.t0))
(step t27.t2 (cl (not y$140) (= y$v_range$next y$v_range$next_rhs_op)) :rule equiv1 :premises (a47))
(step t27.t3 (cl (= y$v_range$next y$v_range$next_rhs_op) (not y$140)) :rule reordering :premises (t27.t2))
(step t27.t4 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$140) :rule and_pos :args (2))
(step t27.t5 (cl y$140 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t27.t4))
(step t27.t6 (cl (not y$1272) (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) :rule equiv1 :premises (a831))
(step t27.t7 (cl (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271) (not y$1272)) :rule reordering :premises (t27.t6))
(step t27.t8 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1272) :rule and_pos :args (2))
(step t27.t9 (cl y$1272 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t27.t8))
(step t27.t10 (cl (not y$1346) (and y$prop y$1284 y$1272 y$1343 y$1318)) :rule equiv1 :premises (a848))
(step t27.t11 (cl (and y$prop y$1284 y$1272 y$1343 y$1318) (not y$1346)) :rule reordering :premises (t27.t10))
(step t27.t12 (cl (and y$prop y$1284 y$1272 y$1343 y$1318)) :rule resolution :premises (t27.t11 a849))
(step t27.t13 (cl y$1272) :rule resolution :premises (t27.t9 t27.t12))
(step t27.t14 (cl (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) :rule resolution :premises (t27.t7 t27.t13))
(step t27.t15 (cl y$140) :rule resolution :premises (t27.t5 t27.t14))
(step t27.t16 (cl (= y$v_range$next y$v_range$next_rhs_op)) :rule resolution :premises (t27.t3 t27.t15))
(step t27.t17 (cl (not (= (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule equiv_pos2)
(step t27.t18 (cl (= y$v_range$next_rhs_op y$v_range$next_rhs_op)) :rule refl)
(step t27.t19 (cl (= (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule equiv_simplify)
(step t27.t20 (cl (not (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule equiv1 :premises (t27.t19))
(step t27.t21 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule refl)
(step t27.t22 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule refl)
(step t27.t23 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule cong :premises (t27.t21 t27.t22))
(step t27.t24 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(step t27.t25 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) :rule trans :premises (t27.t23 t27.t24))
(step t27.t26 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t27.t20 t27.t25))
(step t27.t27 (cl (= (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule cong :premises (t27.t18 t27.t26))
(step t27.t28 (cl (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t27.t17 t27.t27 t27.a6))
(step t27.t29 (cl (not (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) (not (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule equiv_pos2)
(step t27.t30 (cl (= y$f11 y$f11)) :rule refl)
(step t27.t31 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16))
(step t27.t32 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))
(step t27.t33 (cl (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule cong :premises (t27.t30 t27.t31 t27.t32))
(step t27.t34 (cl (not (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)))) (not (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule equiv_pos2)
(step t27.t35 (cl (= y$n3s16 y$n3s16)) :rule refl)
(step t27.t36 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16))) :rule cong :premises (t27.t26 t27.t35))
(step t27.t37 (cl (= y$v3_1517448506_81_op y$v3_1517448506_81_op)) :rule refl)
(step t27.t38 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule cong :premises (t27.t26 t27.t37))
(step t27.t39 (cl (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)))) :rule cong :premises (t27.t30 t27.t36 t27.t38))
(step t27.t40 (cl (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule hole)
(step t27.t41 (cl (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule resolution :premises (t27.t34 t27.t39 t27.t40))
(step t27.t42 (cl (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule resolution :premises (t27.t29 t27.t33 t27.t41))
(step t27.t43 (cl y$f11 (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule ite1 :premises (t27.t42))
(step t27.t44 (cl (not y$f11) (not y$222)) :rule equiv1 :premises (a104))
(step t27.t45 (cl (not y$222) (not y$f11)) :rule reordering :premises (t27.t44))
(step t27.t46 (cl (not (and y$221 y$222)) y$222) :rule and_pos :args (1))
(step t27.t47 (cl y$222 (not (and y$221 y$222))) :rule reordering :premises (t27.t46))
(step t27.t48 (cl (not y$v3_1517448506_121_op) (and y$221 y$222)) :rule equiv1 :premises (a105))
(step t27.t49 (cl (and y$221 y$222) (not y$v3_1517448506_121_op)) :rule reordering :premises (t27.t48))
(step t27.t50 (cl (not (= (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op y$225))) (not (or y$v3_1517448506_121_op (not (not y$225)))) (or y$v3_1517448506_121_op y$225)) :rule equiv_pos2)
(step t27.t51 (cl (= y$v3_1517448506_121_op y$v3_1517448506_121_op)) :rule refl)
(step t27.t52 (cl (= (not (not y$225)) y$225)) :rule rare_rewrite :args ("bool-double-not-elim" y$225))
(step t27.t53 (cl (= (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op y$225))) :rule cong :premises (t27.t51 t27.t52))
(step t27.t54 (cl y$v3_1517448506_121_op (not (not y$225))) :rule equiv2 :premises (a106))
(step t27.t55 (cl (or y$v3_1517448506_121_op (not (not y$225))) (not y$v3_1517448506_121_op)) :rule or_neg :args (0))
(step t27.t56 (cl (or y$v3_1517448506_121_op (not (not y$225))) (not (not (not y$225)))) :rule or_neg :args (1))
(step t27.t57 (cl (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op (not (not y$225)))) :rule resolution :premises (t27.t54 t27.t55 t27.t56))
(step t27.t58 (cl (or y$v3_1517448506_121_op (not (not y$225)))) :rule contraction :premises (t27.t57))
(step t27.t59 (cl (or y$v3_1517448506_121_op y$225)) :rule resolution :premises (t27.t50 t27.t53 t27.t58))
(step t27.t60 (cl y$v3_1517448506_121_op y$225) :rule or :premises (t27.t59))
(step t27.t61 (cl y$225 y$v3_1517448506_121_op) :rule reordering :premises (t27.t60))
(step t27.t62 (cl (and y$161 y$225) (not y$161) (not y$225)) :rule and_neg)
(step t27.t63 (cl (not y$161) (not y$225) (and y$161 y$225)) :rule reordering :premises (t27.t62))
(step t27.t64 (cl (not (and y$1 y$161)) y$161) :rule and_pos :args (1))
(step t27.t65 (cl y$161 (not (and y$1 y$161))) :rule reordering :premises (t27.t64))
(step t27.t66 (cl (not y$v3_1517448506_168_op) (and y$1 y$161)) :rule equiv1 :premises (a165))
(step t27.t67 (cl (and y$1 y$161) (not y$v3_1517448506_168_op)) :rule reordering :premises (t27.t66))
(step t27.t68 (cl (not (= (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op y$313))) (not (or y$v3_1517448506_168_op (not (not y$313)))) (or y$v3_1517448506_168_op y$313)) :rule equiv_pos2)
(step t27.t69 (cl (= y$v3_1517448506_168_op y$v3_1517448506_168_op)) :rule refl)
(step t27.t70 (cl (= (not (not y$313)) y$313)) :rule rare_rewrite :args ("bool-double-not-elim" y$313))
(step t27.t71 (cl (= (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op y$313))) :rule cong :premises (t27.t69 t27.t70))
(step t27.t72 (cl y$v3_1517448506_168_op (not (not y$313))) :rule equiv2 :premises (a166))
(step t27.t73 (cl (or y$v3_1517448506_168_op (not (not y$313))) (not y$v3_1517448506_168_op)) :rule or_neg :args (0))
(step t27.t74 (cl (or y$v3_1517448506_168_op (not (not y$313))) (not (not (not y$313)))) :rule or_neg :args (1))
(step t27.t75 (cl (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op (not (not y$313)))) :rule resolution :premises (t27.t72 t27.t73 t27.t74))
(step t27.t76 (cl (or y$v3_1517448506_168_op (not (not y$313)))) :rule contraction :premises (t27.t75))
(step t27.t77 (cl (or y$v3_1517448506_168_op y$313)) :rule resolution :premises (t27.t68 t27.t71 t27.t76))
(step t27.t78 (cl y$v3_1517448506_168_op y$313) :rule or :premises (t27.t77))
(step t27.t79 (cl y$313 y$v3_1517448506_168_op) :rule reordering :premises (t27.t78))
(step t27.t80 (cl (and y$313 y$315) (not y$313) (not y$315)) :rule and_neg)
(step t27.t81 (cl (not y$313) (not y$315) (and y$313 y$315)) :rule reordering :premises (t27.t80))
(step t27.t82 (cl (not (and y$3 y$315)) y$315) :rule and_pos :args (1))
(step t27.t83 (cl y$315 (not (and y$3 y$315))) :rule reordering :premises (t27.t82))
(step t27.t84 (cl (not y$v3_1517448506_173_op) (and y$3 y$315)) :rule equiv1 :premises (a171))
(step t27.t85 (cl (and y$3 y$315) (not y$v3_1517448506_173_op)) :rule reordering :premises (t27.t84))
(step t27.t86 (cl (not (= (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op y$322))) (not (or y$v3_1517448506_173_op (not (not y$322)))) (or y$v3_1517448506_173_op y$322)) :rule equiv_pos2)
(step t27.t87 (cl (= y$v3_1517448506_173_op y$v3_1517448506_173_op)) :rule refl)
(step t27.t88 (cl (= (not (not y$322)) y$322)) :rule rare_rewrite :args ("bool-double-not-elim" y$322))
(step t27.t89 (cl (= (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op y$322))) :rule cong :premises (t27.t87 t27.t88))
(step t27.t90 (cl y$v3_1517448506_173_op (not (not y$322))) :rule equiv2 :premises (a172))
(step t27.t91 (cl (or y$v3_1517448506_173_op (not (not y$322))) (not y$v3_1517448506_173_op)) :rule or_neg :args (0))
(step t27.t92 (cl (or y$v3_1517448506_173_op (not (not y$322))) (not (not (not y$322)))) :rule or_neg :args (1))
(step t27.t93 (cl (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op (not (not y$322)))) :rule resolution :premises (t27.t90 t27.t91 t27.t92))
(step t27.t94 (cl (or y$v3_1517448506_173_op (not (not y$322)))) :rule contraction :premises (t27.t93))
(step t27.t95 (cl (or y$v3_1517448506_173_op y$322)) :rule resolution :premises (t27.t86 t27.t89 t27.t94))
(step t27.t96 (cl y$v3_1517448506_173_op y$322) :rule or :premises (t27.t95))
(step t27.t97 (cl y$322 y$v3_1517448506_173_op) :rule reordering :premises (t27.t96))
(step t27.t98 (cl (and y$206 y$322) (not y$206) (not y$322)) :rule and_neg)
(step t27.t99 (cl (not y$206) (not y$322) (and y$206 y$322)) :rule reordering :premises (t27.t98))
(step t27.t100 (cl (not (and y$204 y$206)) y$206) :rule and_pos :args (1))
(step t27.t101 (cl y$206 (not (and y$204 y$206))) :rule reordering :premises (t27.t100))
(step t27.t102 (cl (not y$v3_1517448506_114_op) (and y$204 y$206)) :rule equiv1 :premises (a92))
(step t27.t103 (cl (and y$204 y$206) (not y$v3_1517448506_114_op)) :rule reordering :premises (t27.t102))
(step t27.t104 (cl (not (= (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op y$209))) (not (or y$v3_1517448506_114_op (not (not y$209)))) (or y$v3_1517448506_114_op y$209)) :rule equiv_pos2)
(step t27.t105 (cl (= y$v3_1517448506_114_op y$v3_1517448506_114_op)) :rule refl)
(step t27.t106 (cl (= (not (not y$209)) y$209)) :rule rare_rewrite :args ("bool-double-not-elim" y$209))
(step t27.t107 (cl (= (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op y$209))) :rule cong :premises (t27.t105 t27.t106))
(step t27.t108 (cl y$v3_1517448506_114_op (not (not y$209))) :rule equiv2 :premises (a93))
(step t27.t109 (cl (or y$v3_1517448506_114_op (not (not y$209))) (not y$v3_1517448506_114_op)) :rule or_neg :args (0))
(step t27.t110 (cl (or y$v3_1517448506_114_op (not (not y$209))) (not (not (not y$209)))) :rule or_neg :args (1))
(step t27.t111 (cl (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op (not (not y$209)))) :rule resolution :premises (t27.t108 t27.t109 t27.t110))
(step t27.t112 (cl (or y$v3_1517448506_114_op (not (not y$209)))) :rule contraction :premises (t27.t111))
(step t27.t113 (cl (or y$v3_1517448506_114_op y$209)) :rule resolution :premises (t27.t104 t27.t107 t27.t112))
(step t27.t114 (cl y$v3_1517448506_114_op y$209) :rule or :premises (t27.t113))
(step t27.t115 (cl y$209 y$v3_1517448506_114_op) :rule reordering :premises (t27.t114))
(step t27.t116 (cl (not y$209) (not y$1104)) :rule equiv1 :premises (a708))
(step t27.t117 (cl (not (and y$1103 y$1104)) y$1104) :rule and_pos :args (1))
(step t27.t118 (cl y$1104 (not (and y$1103 y$1104))) :rule reordering :premises (t27.t117))
(step t27.t119 (cl (not y$v3_1517448506_547_op) (and y$1103 y$1104)) :rule equiv1 :premises (a709))
(step t27.t120 (cl (and y$1103 y$1104) (not y$v3_1517448506_547_op)) :rule reordering :premises (t27.t119))
(step t27.t121 (cl (not (= (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op y$1107))) (not (or y$v3_1517448506_547_op (not (not y$1107)))) (or y$v3_1517448506_547_op y$1107)) :rule equiv_pos2)
(step t27.t122 (cl (= y$v3_1517448506_547_op y$v3_1517448506_547_op)) :rule refl)
(step t27.t123 (cl (= (not (not y$1107)) y$1107)) :rule rare_rewrite :args ("bool-double-not-elim" y$1107))
(step t27.t124 (cl (= (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op y$1107))) :rule cong :premises (t27.t122 t27.t123))
(step t27.t125 (cl y$v3_1517448506_547_op (not (not y$1107))) :rule equiv2 :premises (a710))
(step t27.t126 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (not y$v3_1517448506_547_op)) :rule or_neg :args (0))
(step t27.t127 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (not (not (not y$1107)))) :rule or_neg :args (1))
(step t27.t128 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op (not (not y$1107)))) :rule resolution :premises (t27.t125 t27.t126 t27.t127))
(step t27.t129 (cl (or y$v3_1517448506_547_op (not (not y$1107)))) :rule contraction :premises (t27.t128))
(step t27.t130 (cl (or y$v3_1517448506_547_op y$1107)) :rule resolution :premises (t27.t121 t27.t124 t27.t129))
(step t27.t131 (cl y$v3_1517448506_547_op y$1107) :rule or :premises (t27.t130))
(step t27.t132 (cl y$1107 y$v3_1517448506_547_op) :rule reordering :premises (t27.t131))
(step t27.t133 (cl (not y$1107) (not y$1116)) :rule equiv1 :premises (a717))
(step t27.t134 (cl (not (and y$1115 y$1116)) y$1116) :rule and_pos :args (1))
(step t27.t135 (cl y$1116 (not (and y$1115 y$1116))) :rule reordering :premises (t27.t134))
(step t27.t136 (cl (not y$v3_1517448506_552_op) (and y$1115 y$1116)) :rule equiv1 :premises (a718))
(step t27.t137 (cl (and y$1115 y$1116) (not y$v3_1517448506_552_op)) :rule reordering :premises (t27.t136))
(step t27.t138 (cl (not (= (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op y$1119))) (not (or y$v3_1517448506_552_op (not (not y$1119)))) (or y$v3_1517448506_552_op y$1119)) :rule equiv_pos2)
(step t27.t139 (cl (= y$v3_1517448506_552_op y$v3_1517448506_552_op)) :rule refl)
(step t27.t140 (cl (= (not (not y$1119)) y$1119)) :rule rare_rewrite :args ("bool-double-not-elim" y$1119))
(step t27.t141 (cl (= (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op y$1119))) :rule cong :premises (t27.t139 t27.t140))
(step t27.t142 (cl y$v3_1517448506_552_op (not (not y$1119))) :rule equiv2 :premises (a719))
(step t27.t143 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (not y$v3_1517448506_552_op)) :rule or_neg :args (0))
(step t27.t144 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (not (not (not y$1119)))) :rule or_neg :args (1))
(step t27.t145 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op (not (not y$1119)))) :rule resolution :premises (t27.t142 t27.t143 t27.t144))
(step t27.t146 (cl (or y$v3_1517448506_552_op (not (not y$1119)))) :rule contraction :premises (t27.t145))
(step t27.t147 (cl (or y$v3_1517448506_552_op y$1119)) :rule resolution :premises (t27.t138 t27.t141 t27.t146))
(step t27.t148 (cl y$v3_1517448506_552_op y$1119) :rule or :premises (t27.t147))
(step t27.t149 (cl y$1119 y$v3_1517448506_552_op) :rule reordering :premises (t27.t148))
(step t27.t150 (cl (not y$1119) (not y$1128)) :rule equiv1 :premises (a726))
(step t27.t151 (cl (not (and y$1127 y$1128)) y$1128) :rule and_pos :args (1))
(step t27.t152 (cl y$1128 (not (and y$1127 y$1128))) :rule reordering :premises (t27.t151))
(step t27.t153 (cl (not y$v3_1517448506_557_op) (and y$1127 y$1128)) :rule equiv1 :premises (a727))
(step t27.t154 (cl (and y$1127 y$1128) (not y$v3_1517448506_557_op)) :rule reordering :premises (t27.t153))
(step t27.t155 (cl (not (= (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op y$1131))) (not (or y$v3_1517448506_557_op (not (not y$1131)))) (or y$v3_1517448506_557_op y$1131)) :rule equiv_pos2)
(step t27.t156 (cl (= y$v3_1517448506_557_op y$v3_1517448506_557_op)) :rule refl)
(step t27.t157 (cl (= (not (not y$1131)) y$1131)) :rule rare_rewrite :args ("bool-double-not-elim" y$1131))
(step t27.t158 (cl (= (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op y$1131))) :rule cong :premises (t27.t156 t27.t157))
(step t27.t159 (cl y$v3_1517448506_557_op (not (not y$1131))) :rule equiv2 :premises (a728))
(step t27.t160 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (not y$v3_1517448506_557_op)) :rule or_neg :args (0))
(step t27.t161 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (not (not (not y$1131)))) :rule or_neg :args (1))
(step t27.t162 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op (not (not y$1131)))) :rule resolution :premises (t27.t159 t27.t160 t27.t161))
(step t27.t163 (cl (or y$v3_1517448506_557_op (not (not y$1131)))) :rule contraction :premises (t27.t162))
(step t27.t164 (cl (or y$v3_1517448506_557_op y$1131)) :rule resolution :premises (t27.t155 t27.t158 t27.t163))
(step t27.t165 (cl y$v3_1517448506_557_op y$1131) :rule or :premises (t27.t164))
(step t27.t166 (cl y$1131 y$v3_1517448506_557_op) :rule reordering :premises (t27.t165))
(step t27.t167 (cl (not y$1131) (not y$1140)) :rule equiv1 :premises (a735))
(step t27.t168 (cl (not (and y$1139 y$1140)) y$1140) :rule and_pos :args (1))
(step t27.t169 (cl y$1140 (not (and y$1139 y$1140))) :rule reordering :premises (t27.t168))
(step t27.t170 (cl (not y$v3_1517448506_562_op) (and y$1139 y$1140)) :rule equiv1 :premises (a736))
(step t27.t171 (cl (and y$1139 y$1140) (not y$v3_1517448506_562_op)) :rule reordering :premises (t27.t170))
(step t27.t172 (cl (not (= (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op y$1143))) (not (or y$v3_1517448506_562_op (not (not y$1143)))) (or y$v3_1517448506_562_op y$1143)) :rule equiv_pos2)
(step t27.t173 (cl (= y$v3_1517448506_562_op y$v3_1517448506_562_op)) :rule refl)
(step t27.t174 (cl (= (not (not y$1143)) y$1143)) :rule rare_rewrite :args ("bool-double-not-elim" y$1143))
(step t27.t175 (cl (= (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op y$1143))) :rule cong :premises (t27.t173 t27.t174))
(step t27.t176 (cl y$v3_1517448506_562_op (not (not y$1143))) :rule equiv2 :premises (a737))
(step t27.t177 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (not y$v3_1517448506_562_op)) :rule or_neg :args (0))
(step t27.t178 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (not (not (not y$1143)))) :rule or_neg :args (1))
(step t27.t179 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op (not (not y$1143)))) :rule resolution :premises (t27.t176 t27.t177 t27.t178))
(step t27.t180 (cl (or y$v3_1517448506_562_op (not (not y$1143)))) :rule contraction :premises (t27.t179))
(step t27.t181 (cl (or y$v3_1517448506_562_op y$1143)) :rule resolution :premises (t27.t172 t27.t175 t27.t180))
(step t27.t182 (cl y$v3_1517448506_562_op y$1143) :rule or :premises (t27.t181))
(step t27.t183 (cl y$1143 y$v3_1517448506_562_op) :rule reordering :premises (t27.t182))
(step t27.t184 (cl (and y$a_too_small_resistance$next_rhs_op y$1143) (not y$a_too_small_resistance$next_rhs_op) (not y$1143)) :rule and_neg)
(step t27.t185 (cl (not y$1143) (and y$a_too_small_resistance$next_rhs_op y$1143) (not y$a_too_small_resistance$next_rhs_op)) :rule reordering :premises (t27.t184))
(step t27.t186 (cl y$v3_1517448506_563_op (not (and y$a_too_small_resistance$next_rhs_op y$1143))) :rule equiv2 :premises (a738))
(step t27.t187 (cl (not y$v3_1517448506_563_op) (not y$1146)) :rule equiv1 :premises (a739))
(step t27.t188 (cl (not y$1146) (not y$v3_1517448506_563_op)) :rule reordering :premises (t27.t187))
(step t27.t189 (cl (not (and y$1138 y$1146)) y$1146) :rule and_pos :args (1))
(step t27.t190 (cl y$1146 (not (and y$1138 y$1146))) :rule reordering :premises (t27.t189))
(step t27.t191 (cl (not y$v3_1517448506_565_op) (and y$1138 y$1146)) :rule equiv1 :premises (a740))
(step t27.t192 (cl (and y$1138 y$1146) (not y$v3_1517448506_565_op)) :rule reordering :premises (t27.t191))
(step t27.t193 (cl (not (= (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op y$1149))) (not (or y$v3_1517448506_565_op (not (not y$1149)))) (or y$v3_1517448506_565_op y$1149)) :rule equiv_pos2)
(step t27.t194 (cl (= y$v3_1517448506_565_op y$v3_1517448506_565_op)) :rule refl)
(step t27.t195 (cl (= (not (not y$1149)) y$1149)) :rule rare_rewrite :args ("bool-double-not-elim" y$1149))
(step t27.t196 (cl (= (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op y$1149))) :rule cong :premises (t27.t194 t27.t195))
(step t27.t197 (cl y$v3_1517448506_565_op (not (not y$1149))) :rule equiv2 :premises (a741))
(step t27.t198 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (not y$v3_1517448506_565_op)) :rule or_neg :args (0))
(step t27.t199 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (not (not (not y$1149)))) :rule or_neg :args (1))
(step t27.t200 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op (not (not y$1149)))) :rule resolution :premises (t27.t197 t27.t198 t27.t199))
(step t27.t201 (cl (or y$v3_1517448506_565_op (not (not y$1149)))) :rule contraction :premises (t27.t200))
(step t27.t202 (cl (or y$v3_1517448506_565_op y$1149)) :rule resolution :premises (t27.t193 t27.t196 t27.t201))
(step t27.t203 (cl y$v3_1517448506_565_op y$1149) :rule or :premises (t27.t202))
(step t27.t204 (cl y$1149 y$v3_1517448506_565_op) :rule reordering :premises (t27.t203))
(step t27.t205 (cl (not y$1149) (not y$1150)) :rule equiv1 :premises (a742))
(step t27.t206 (cl (not (and y$1150 y$1158)) y$1150) :rule and_pos :args (0))
(step t27.t207 (cl y$1150 (not (and y$1150 y$1158))) :rule reordering :premises (t27.t206))
(step t27.t208 (cl (not y$v3_1517448506_570_op) (and y$1150 y$1158)) :rule equiv1 :premises (a749))
(step t27.t209 (cl (and y$1150 y$1158) (not y$v3_1517448506_570_op)) :rule reordering :premises (t27.t208))
(step t27.t210 (cl (not (= (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op y$1161))) (not (or y$v3_1517448506_570_op (not (not y$1161)))) (or y$v3_1517448506_570_op y$1161)) :rule equiv_pos2)
(step t27.t211 (cl (= y$v3_1517448506_570_op y$v3_1517448506_570_op)) :rule refl)
(step t27.t212 (cl (= (not (not y$1161)) y$1161)) :rule rare_rewrite :args ("bool-double-not-elim" y$1161))
(step t27.t213 (cl (= (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op y$1161))) :rule cong :premises (t27.t211 t27.t212))
(step t27.t214 (cl y$v3_1517448506_570_op (not (not y$1161))) :rule equiv2 :premises (a750))
(step t27.t215 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (not y$v3_1517448506_570_op)) :rule or_neg :args (0))
(step t27.t216 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (not (not (not y$1161)))) :rule or_neg :args (1))
(step t27.t217 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op (not (not y$1161)))) :rule resolution :premises (t27.t214 t27.t215 t27.t216))
(step t27.t218 (cl (or y$v3_1517448506_570_op (not (not y$1161)))) :rule contraction :premises (t27.t217))
(step t27.t219 (cl (or y$v3_1517448506_570_op y$1161)) :rule resolution :premises (t27.t210 t27.t213 t27.t218))
(step t27.t220 (cl y$v3_1517448506_570_op y$1161) :rule or :premises (t27.t219))
(step t27.t221 (cl y$1161 y$v3_1517448506_570_op) :rule reordering :premises (t27.t220))
(step t27.t222 (cl (not y$1161) (not y$1162)) :rule equiv1 :premises (a751))
(step t27.t223 (cl (not (and y$1162 y$1170)) y$1162) :rule and_pos :args (0))
(step t27.t224 (cl y$1162 (not (and y$1162 y$1170))) :rule reordering :premises (t27.t223))
(step t27.t225 (cl (not y$v3_1517448506_575_op) (and y$1162 y$1170)) :rule equiv1 :premises (a758))
(step t27.t226 (cl (and y$1162 y$1170) (not y$v3_1517448506_575_op)) :rule reordering :premises (t27.t225))
(step t27.t227 (cl (not (= (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op y$1173))) (not (or y$v3_1517448506_575_op (not (not y$1173)))) (or y$v3_1517448506_575_op y$1173)) :rule equiv_pos2)
(step t27.t228 (cl (= y$v3_1517448506_575_op y$v3_1517448506_575_op)) :rule refl)
(step t27.t229 (cl (= (not (not y$1173)) y$1173)) :rule rare_rewrite :args ("bool-double-not-elim" y$1173))
(step t27.t230 (cl (= (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op y$1173))) :rule cong :premises (t27.t228 t27.t229))
(step t27.t231 (cl y$v3_1517448506_575_op (not (not y$1173))) :rule equiv2 :premises (a759))
(step t27.t232 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (not y$v3_1517448506_575_op)) :rule or_neg :args (0))
(step t27.t233 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (not (not (not y$1173)))) :rule or_neg :args (1))
(step t27.t234 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op (not (not y$1173)))) :rule resolution :premises (t27.t231 t27.t232 t27.t233))
(step t27.t235 (cl (or y$v3_1517448506_575_op (not (not y$1173)))) :rule contraction :premises (t27.t234))
(step t27.t236 (cl (or y$v3_1517448506_575_op y$1173)) :rule resolution :premises (t27.t227 t27.t230 t27.t235))
(step t27.t237 (cl y$v3_1517448506_575_op y$1173) :rule or :premises (t27.t236))
(step t27.t238 (cl y$1173 y$v3_1517448506_575_op) :rule reordering :premises (t27.t237))
(step t27.t239 (cl (not y$1173) (not y$1174)) :rule equiv1 :premises (a760))
(step t27.t240 (cl (not (and y$1174 y$1181)) y$1174) :rule and_pos :args (0))
(step t27.t241 (cl y$1174 (not (and y$1174 y$1181))) :rule reordering :premises (t27.t240))
(step t27.t242 (cl (not y$v3_1517448506_580_op) (and y$1174 y$1181)) :rule equiv1 :premises (a766))
(step t27.t243 (cl (and y$1174 y$1181) (not y$v3_1517448506_580_op)) :rule reordering :premises (t27.t242))
(step t27.t244 (cl (not (= (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op y$1184))) (not (or y$v3_1517448506_580_op (not (not y$1184)))) (or y$v3_1517448506_580_op y$1184)) :rule equiv_pos2)
(step t27.t245 (cl (= y$v3_1517448506_580_op y$v3_1517448506_580_op)) :rule refl)
(step t27.t246 (cl (= (not (not y$1184)) y$1184)) :rule rare_rewrite :args ("bool-double-not-elim" y$1184))
(step t27.t247 (cl (= (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op y$1184))) :rule cong :premises (t27.t245 t27.t246))
(step t27.t248 (cl y$v3_1517448506_580_op (not (not y$1184))) :rule equiv2 :premises (a767))
(step t27.t249 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (not y$v3_1517448506_580_op)) :rule or_neg :args (0))
(step t27.t250 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (not (not (not y$1184)))) :rule or_neg :args (1))
(step t27.t251 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op (not (not y$1184)))) :rule resolution :premises (t27.t248 t27.t249 t27.t250))
(step t27.t252 (cl (or y$v3_1517448506_580_op (not (not y$1184)))) :rule contraction :premises (t27.t251))
(step t27.t253 (cl (or y$v3_1517448506_580_op y$1184)) :rule resolution :premises (t27.t244 t27.t247 t27.t252))
(step t27.t254 (cl y$v3_1517448506_580_op y$1184) :rule or :premises (t27.t253))
(step t27.t255 (cl y$1184 y$v3_1517448506_580_op) :rule reordering :premises (t27.t254))
(step t27.t256 (cl (not y$1184) (not y$1185)) :rule equiv1 :premises (a768))
(step t27.t257 (cl (not (and y$1185 y$1193)) y$1185) :rule and_pos :args (0))
(step t27.t258 (cl y$1185 (not (and y$1185 y$1193))) :rule reordering :premises (t27.t257))
(step t27.t259 (cl (not y$v3_1517448506_585_op) (and y$1185 y$1193)) :rule equiv1 :premises (a775))
(step t27.t260 (cl (and y$1185 y$1193) (not y$v3_1517448506_585_op)) :rule reordering :premises (t27.t259))
(step t27.t261 (cl (not (= (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op y$1196))) (not (or y$v3_1517448506_585_op (not (not y$1196)))) (or y$v3_1517448506_585_op y$1196)) :rule equiv_pos2)
(step t27.t262 (cl (= y$v3_1517448506_585_op y$v3_1517448506_585_op)) :rule refl)
(step t27.t263 (cl (= (not (not y$1196)) y$1196)) :rule rare_rewrite :args ("bool-double-not-elim" y$1196))
(step t27.t264 (cl (= (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op y$1196))) :rule cong :premises (t27.t262 t27.t263))
(step t27.t265 (cl y$v3_1517448506_585_op (not (not y$1196))) :rule equiv2 :premises (a776))
(step t27.t266 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (not y$v3_1517448506_585_op)) :rule or_neg :args (0))
(step t27.t267 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (not (not (not y$1196)))) :rule or_neg :args (1))
(step t27.t268 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op (not (not y$1196)))) :rule resolution :premises (t27.t265 t27.t266 t27.t267))
(step t27.t269 (cl (or y$v3_1517448506_585_op (not (not y$1196)))) :rule contraction :premises (t27.t268))
(step t27.t270 (cl (or y$v3_1517448506_585_op y$1196)) :rule resolution :premises (t27.t261 t27.t264 t27.t269))
(step t27.t271 (cl y$v3_1517448506_585_op y$1196) :rule or :premises (t27.t270))
(step t27.t272 (cl y$1196 y$v3_1517448506_585_op) :rule reordering :premises (t27.t271))
(step t27.t273 (cl (not y$1196) (not y$1197)) :rule equiv1 :premises (a777))
(step t27.t274 (cl (not (and y$1197 y$1205)) y$1197) :rule and_pos :args (0))
(step t27.t275 (cl y$1197 (not (and y$1197 y$1205))) :rule reordering :premises (t27.t274))
(step t27.t276 (cl (not y$v3_1517448506_590_op) (and y$1197 y$1205)) :rule equiv1 :premises (a784))
(step t27.t277 (cl (and y$1197 y$1205) (not y$v3_1517448506_590_op)) :rule reordering :premises (t27.t276))
(step t27.t278 (cl (not (= (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op y$1208))) (not (or y$v3_1517448506_590_op (not (not y$1208)))) (or y$v3_1517448506_590_op y$1208)) :rule equiv_pos2)
(step t27.t279 (cl (= y$v3_1517448506_590_op y$v3_1517448506_590_op)) :rule refl)
(step t27.t280 (cl (= (not (not y$1208)) y$1208)) :rule rare_rewrite :args ("bool-double-not-elim" y$1208))
(step t27.t281 (cl (= (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op y$1208))) :rule cong :premises (t27.t279 t27.t280))
(step t27.t282 (cl y$v3_1517448506_590_op (not (not y$1208))) :rule equiv2 :premises (a785))
(step t27.t283 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (not y$v3_1517448506_590_op)) :rule or_neg :args (0))
(step t27.t284 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (not (not (not y$1208)))) :rule or_neg :args (1))
(step t27.t285 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op (not (not y$1208)))) :rule resolution :premises (t27.t282 t27.t283 t27.t284))
(step t27.t286 (cl (or y$v3_1517448506_590_op (not (not y$1208)))) :rule contraction :premises (t27.t285))
(step t27.t287 (cl (or y$v3_1517448506_590_op y$1208)) :rule resolution :premises (t27.t278 t27.t281 t27.t286))
(step t27.t288 (cl y$v3_1517448506_590_op y$1208) :rule or :premises (t27.t287))
(step t27.t289 (cl y$1208 y$v3_1517448506_590_op) :rule reordering :premises (t27.t288))
(step t27.t290 (cl (not y$1208) (not y$1209)) :rule equiv1 :premises (a786))
(step t27.t291 (cl (not (and y$v3_1517448506_544_op y$1209)) y$1209) :rule and_pos :args (1))
(step t27.t292 (cl y$1209 (not (and y$v3_1517448506_544_op y$1209))) :rule reordering :premises (t27.t291))
(step t27.t293 (cl (not y$v3_1517448506_591_op) (and y$v3_1517448506_544_op y$1209)) :rule equiv1 :premises (a787))
(step t27.t294 (cl (and y$v3_1517448506_544_op y$1209) (not y$v3_1517448506_591_op)) :rule reordering :premises (t27.t293))
(step t27.t295 (cl (not (and y$v3_1517448506_591_op y$1216)) y$v3_1517448506_591_op) :rule and_pos :args (0))
(step t27.t296 (cl y$v3_1517448506_591_op (not (and y$v3_1517448506_591_op y$1216))) :rule reordering :premises (t27.t295))
(step t27.t297 (cl (not y$v3_1517448506_594_op) (and y$v3_1517448506_591_op y$1216)) :rule equiv1 :premises (a792))
(step t27.t298 (cl (and y$v3_1517448506_591_op y$1216) (not y$v3_1517448506_594_op)) :rule reordering :premises (t27.t297))
(step t27.t299 (cl (not (and y$v3_1517448506_594_op y$1256)) y$v3_1517448506_594_op) :rule and_pos :args (0))
(step t27.t300 (cl y$v3_1517448506_594_op (not (and y$v3_1517448506_594_op y$1256))) :rule reordering :premises (t27.t299))
(step t27.t301 (cl (not y$v3_1517448506_611_op) (and y$v3_1517448506_594_op y$1256)) :rule equiv1 :premises (a821))
(step t27.t302 (cl (and y$v3_1517448506_594_op y$1256) (not y$v3_1517448506_611_op)) :rule reordering :premises (t27.t301))
(step t27.t303 (cl (not (and y$v3_1517448506_611_op y$1263)) y$v3_1517448506_611_op) :rule and_pos :args (0))
(step t27.t304 (cl y$v3_1517448506_611_op (not (and y$v3_1517448506_611_op y$1263))) :rule reordering :premises (t27.t303))
(step t27.t305 (cl (not y$v3_1517448506_614_op) (and y$v3_1517448506_611_op y$1263)) :rule equiv1 :premises (a826))
(step t27.t306 (cl (and y$v3_1517448506_611_op y$1263) (not y$v3_1517448506_614_op)) :rule reordering :premises (t27.t305))
(step t27.t307 (cl (not (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) y$v3_1517448506_614_op) :rule and_pos :args (1))
(step t27.t308 (cl y$v3_1517448506_614_op (not (and y$v3_1517448506_530_op y$v3_1517448506_614_op))) :rule reordering :premises (t27.t307))
(step t27.t309 (cl (not y$v3_1517448506_615_op) (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) :rule equiv1 :premises (a827))
(step t27.t310 (cl (and y$v3_1517448506_530_op y$v3_1517448506_614_op) (not y$v3_1517448506_615_op)) :rule reordering :premises (t27.t309))
(step t27.t311 (cl (not (and y$41 y$v3_1517448506_615_op)) y$v3_1517448506_615_op) :rule and_pos :args (1))
(step t27.t312 (cl y$v3_1517448506_615_op (not (and y$41 y$v3_1517448506_615_op))) :rule reordering :premises (t27.t311))
(step t27.t313 (cl (not y$v3_1517448506_616_op) (and y$41 y$v3_1517448506_615_op)) :rule equiv1 :premises (a828))
(step t27.t314 (cl (and y$41 y$v3_1517448506_615_op) (not y$v3_1517448506_616_op)) :rule reordering :premises (t27.t313))
(step t27.t315 (cl (not (= (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op y$1270))) (not (or y$v3_1517448506_616_op (not (not y$1270)))) (or y$v3_1517448506_616_op y$1270)) :rule equiv_pos2)
(step t27.t316 (cl (= y$v3_1517448506_616_op y$v3_1517448506_616_op)) :rule refl)
(step t27.t317 (cl (= (not (not y$1270)) y$1270)) :rule rare_rewrite :args ("bool-double-not-elim" y$1270))
(step t27.t318 (cl (= (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op y$1270))) :rule cong :premises (t27.t316 t27.t317))
(step t27.t319 (cl y$v3_1517448506_616_op (not (not y$1270))) :rule equiv2 :premises (a829))
(step t27.t320 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (not y$v3_1517448506_616_op)) :rule or_neg :args (0))
(step t27.t321 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (not (not (not y$1270)))) :rule or_neg :args (1))
(step t27.t322 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op (not (not y$1270)))) :rule resolution :premises (t27.t319 t27.t320 t27.t321))
(step t27.t323 (cl (or y$v3_1517448506_616_op (not (not y$1270)))) :rule contraction :premises (t27.t322))
(step t27.t324 (cl (or y$v3_1517448506_616_op y$1270)) :rule resolution :premises (t27.t315 t27.t318 t27.t323))
(step t27.t325 (cl y$v3_1517448506_616_op y$1270) :rule or :premises (t27.t324))
(step t27.t326 (cl y$1270 y$v3_1517448506_616_op) :rule reordering :premises (t27.t325))
(step t27.t327 (cl (not (= y$1270 y$dve_invalid$next)) (not y$1270) y$dve_invalid$next) :rule equiv_pos2)
(step t27.t328 (cl y$dve_invalid$next (not y$1270) (not (= y$1270 y$dve_invalid$next))) :rule reordering :premises (t27.t327))
(step t27.t329 (cl (not y$dve_invalid$next) (not y$1319)) :rule equiv1 :premises (a832))
(step t27.t330 (cl (not y$1319) (not y$dve_invalid$next)) :rule reordering :premises (t27.t329))
(step t27.t331 (cl (not (and y$1319 y$v3_1517448506_38$next_op)) y$1319) :rule and_pos :args (0))
(step t27.t332 (cl y$1319 (not (and y$1319 y$v3_1517448506_38$next_op))) :rule reordering :premises (t27.t331))
(step t27.t333 (cl (not y$id37$next_op) (and y$1319 y$v3_1517448506_38$next_op)) :rule equiv1 :premises (a844))
(step t27.t334 (cl (and y$1319 y$v3_1517448506_38$next_op) (not y$id37$next_op)) :rule reordering :premises (t27.t333))
(step t27.t335 (cl (not (= (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op y$1342))) (not (or y$id37$next_op (not (not y$1342)))) (or y$id37$next_op y$1342)) :rule equiv_pos2)
(step t27.t336 (cl (= y$id37$next_op y$id37$next_op)) :rule refl)
(step t27.t337 (cl (= (not (not y$1342)) y$1342)) :rule rare_rewrite :args ("bool-double-not-elim" y$1342))
(step t27.t338 (cl (= (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op y$1342))) :rule cong :premises (t27.t336 t27.t337))
(step t27.t339 (cl y$id37$next_op (not (not y$1342))) :rule equiv2 :premises (a845))
(step t27.t340 (cl (or y$id37$next_op (not (not y$1342))) (not y$id37$next_op)) :rule or_neg :args (0))
(step t27.t341 (cl (or y$id37$next_op (not (not y$1342))) (not (not (not y$1342)))) :rule or_neg :args (1))
(step t27.t342 (cl (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op (not (not y$1342)))) :rule resolution :premises (t27.t339 t27.t340 t27.t341))
(step t27.t343 (cl (or y$id37$next_op (not (not y$1342)))) :rule contraction :premises (t27.t342))
(step t27.t344 (cl (or y$id37$next_op y$1342)) :rule resolution :premises (t27.t335 t27.t338 t27.t343))
(step t27.t345 (cl y$id37$next_op y$1342) :rule or :premises (t27.t344))
(step t27.t346 (cl y$1342 y$id37$next_op) :rule reordering :premises (t27.t345))
(step t27.t347 (cl (not (= y$1342 y$prop$next)) (not y$1342) y$prop$next) :rule equiv_pos2)
(step t27.t348 (cl y$prop$next (not y$1342) (not (= y$1342 y$prop$next))) :rule reordering :premises (t27.t347))
(step t27.t349 (cl (not y$prop$next) (not y$1318)) :rule equiv1 :premises (a847))
(step t27.t350 (cl (not y$1318) (not y$prop$next)) :rule reordering :premises (t27.t349))
(step t27.t351 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1318) :rule and_pos :args (4))
(step t27.t352 (cl y$1318 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t27.t351))
(step t27.t353 (cl y$1318) :rule resolution :premises (t27.t352 t27.t12))
(step t27.t354 (cl (not y$prop$next)) :rule resolution :premises (t27.t350 t27.t353))
(step t27.t355 (cl (not (= (= y$1343 (= y$prop$next y$1342)) (= y$1343 (= y$1342 y$prop$next)))) (not (= y$1343 (= y$prop$next y$1342))) (= y$1343 (= y$1342 y$prop$next))) :rule equiv_pos2)
(step t27.t356 (cl (= y$1343 y$1343)) :rule refl)
(step t27.t357 (cl (= (= y$prop$next y$1342) (= y$1342 y$prop$next))) :rule rare_rewrite :args ("eq-symm" y$prop$next y$1342))
(step t27.t358 (cl (= (= y$1343 (= y$prop$next y$1342)) (= y$1343 (= y$1342 y$prop$next)))) :rule cong :premises (t27.t356 t27.t357))
(step t27.t359 (cl (= y$1343 (= y$1342 y$prop$next))) :rule resolution :premises (t27.t355 t27.t358 a846))
(step t27.t360 (cl (not y$1343) (= y$1342 y$prop$next)) :rule equiv1 :premises (t27.t359))
(step t27.t361 (cl (= y$1342 y$prop$next) (not y$1343)) :rule reordering :premises (t27.t360))
(step t27.t362 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1343) :rule and_pos :args (3))
(step t27.t363 (cl y$1343 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t27.t362))
(step t27.t364 (cl y$1343) :rule resolution :premises (t27.t363 t27.t12))
(step t27.t365 (cl (= y$1342 y$prop$next)) :rule resolution :premises (t27.t361 t27.t364))
(step t27.t366 (cl (not y$1342)) :rule resolution :premises (t27.t348 t27.t354 t27.t365))
(step t27.t367 (cl y$id37$next_op) :rule resolution :premises (t27.t346 t27.t366))
(step t27.t368 (cl (and y$1319 y$v3_1517448506_38$next_op)) :rule resolution :premises (t27.t334 t27.t367))
(step t27.t369 (cl y$1319) :rule resolution :premises (t27.t332 t27.t368))
(step t27.t370 (cl (not y$dve_invalid$next)) :rule resolution :premises (t27.t330 t27.t369))
(step t27.t371 (cl (not (= (= y$1271 (= y$dve_invalid$next y$1270)) (= y$1271 (= y$1270 y$dve_invalid$next)))) (not (= y$1271 (= y$dve_invalid$next y$1270))) (= y$1271 (= y$1270 y$dve_invalid$next))) :rule equiv_pos2)
(step t27.t372 (cl (= y$1271 y$1271)) :rule refl)
(step t27.t373 (cl (= (= y$dve_invalid$next y$1270) (= y$1270 y$dve_invalid$next))) :rule rare_rewrite :args ("eq-symm" y$dve_invalid$next y$1270))
(step t27.t374 (cl (= (= y$1271 (= y$dve_invalid$next y$1270)) (= y$1271 (= y$1270 y$dve_invalid$next)))) :rule cong :premises (t27.t372 t27.t373))
(step t27.t375 (cl (= y$1271 (= y$1270 y$dve_invalid$next))) :rule resolution :premises (t27.t371 t27.t374 a830))
(step t27.t376 (cl (not y$1271) (= y$1270 y$dve_invalid$next)) :rule equiv1 :premises (t27.t375))
(step t27.t377 (cl (= y$1270 y$dve_invalid$next) (not y$1271)) :rule reordering :premises (t27.t376))
(step t27.t378 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$1271) :rule and_pos :args (23))
(step t27.t379 (cl y$1271 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t27.t378))
(step t27.t380 (cl y$1271) :rule resolution :premises (t27.t379 t27.t14))
(step t27.t381 (cl (= y$1270 y$dve_invalid$next)) :rule resolution :premises (t27.t377 t27.t380))
(step t27.t382 (cl (not y$1270)) :rule resolution :premises (t27.t328 t27.t370 t27.t381))
(step t27.t383 (cl y$v3_1517448506_616_op) :rule resolution :premises (t27.t326 t27.t382))
(step t27.t384 (cl (and y$41 y$v3_1517448506_615_op)) :rule resolution :premises (t27.t314 t27.t383))
(step t27.t385 (cl y$v3_1517448506_615_op) :rule resolution :premises (t27.t312 t27.t384))
(step t27.t386 (cl (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) :rule resolution :premises (t27.t310 t27.t385))
(step t27.t387 (cl y$v3_1517448506_614_op) :rule resolution :premises (t27.t308 t27.t386))
(step t27.t388 (cl (and y$v3_1517448506_611_op y$1263)) :rule resolution :premises (t27.t306 t27.t387))
(step t27.t389 (cl y$v3_1517448506_611_op) :rule resolution :premises (t27.t304 t27.t388))
(step t27.t390 (cl (and y$v3_1517448506_594_op y$1256)) :rule resolution :premises (t27.t302 t27.t389))
(step t27.t391 (cl y$v3_1517448506_594_op) :rule resolution :premises (t27.t300 t27.t390))
(step t27.t392 (cl (and y$v3_1517448506_591_op y$1216)) :rule resolution :premises (t27.t298 t27.t391))
(step t27.t393 (cl y$v3_1517448506_591_op) :rule resolution :premises (t27.t296 t27.t392))
(step t27.t394 (cl (and y$v3_1517448506_544_op y$1209)) :rule resolution :premises (t27.t294 t27.t393))
(step t27.t395 (cl y$1209) :rule resolution :premises (t27.t292 t27.t394))
(step t27.t396 (cl (not y$1208)) :rule resolution :premises (t27.t290 t27.t395))
(step t27.t397 (cl y$v3_1517448506_590_op) :rule resolution :premises (t27.t289 t27.t396))
(step t27.t398 (cl (and y$1197 y$1205)) :rule resolution :premises (t27.t277 t27.t397))
(step t27.t399 (cl y$1197) :rule resolution :premises (t27.t275 t27.t398))
(step t27.t400 (cl (not y$1196)) :rule resolution :premises (t27.t273 t27.t399))
(step t27.t401 (cl y$v3_1517448506_585_op) :rule resolution :premises (t27.t272 t27.t400))
(step t27.t402 (cl (and y$1185 y$1193)) :rule resolution :premises (t27.t260 t27.t401))
(step t27.t403 (cl y$1185) :rule resolution :premises (t27.t258 t27.t402))
(step t27.t404 (cl (not y$1184)) :rule resolution :premises (t27.t256 t27.t403))
(step t27.t405 (cl y$v3_1517448506_580_op) :rule resolution :premises (t27.t255 t27.t404))
(step t27.t406 (cl (and y$1174 y$1181)) :rule resolution :premises (t27.t243 t27.t405))
(step t27.t407 (cl y$1174) :rule resolution :premises (t27.t241 t27.t406))
(step t27.t408 (cl (not y$1173)) :rule resolution :premises (t27.t239 t27.t407))
(step t27.t409 (cl y$v3_1517448506_575_op) :rule resolution :premises (t27.t238 t27.t408))
(step t27.t410 (cl (and y$1162 y$1170)) :rule resolution :premises (t27.t226 t27.t409))
(step t27.t411 (cl y$1162) :rule resolution :premises (t27.t224 t27.t410))
(step t27.t412 (cl (not y$1161)) :rule resolution :premises (t27.t222 t27.t411))
(step t27.t413 (cl y$v3_1517448506_570_op) :rule resolution :premises (t27.t221 t27.t412))
(step t27.t414 (cl (and y$1150 y$1158)) :rule resolution :premises (t27.t209 t27.t413))
(step t27.t415 (cl y$1150) :rule resolution :premises (t27.t207 t27.t414))
(step t27.t416 (cl (not y$1149)) :rule resolution :premises (t27.t205 t27.t415))
(step t27.t417 (cl y$v3_1517448506_565_op) :rule resolution :premises (t27.t204 t27.t416))
(step t27.t418 (cl (and y$1138 y$1146)) :rule resolution :premises (t27.t192 t27.t417))
(step t27.t419 (cl y$1146) :rule resolution :premises (t27.t190 t27.t418))
(step t27.t420 (cl (not y$v3_1517448506_563_op)) :rule resolution :premises (t27.t188 t27.t419))
(step t27.t421 (cl (not (and y$a_too_small_resistance$next_rhs_op y$1143))) :rule resolution :premises (t27.t186 t27.t420))
(step t27.t422 (cl (not (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) (not y$a_too_small_resistance$next) y$a_too_small_resistance$next_rhs_op) :rule equiv_pos2)
(step t27.t423 (cl y$a_too_small_resistance$next_rhs_op (not y$a_too_small_resistance$next) (not (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op))) :rule reordering :premises (t27.t422))
(step t27.t424 (cl (not (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) (not (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv_pos2)
(step t27.t425 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next)) :rule refl)
(step t27.t426 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule rare_rewrite :args ("eq-symm" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32))
(step t27.t427 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule rare_rewrite :args ("eq-symm" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))
(step t27.t428 (cl (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) :rule cong :premises (t27.t425 t27.t426 t27.t427))
(step t27.t429 (cl (not (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)))) (not (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule equiv_pos2)
(step t27.t430 (cl (= (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv_simplify)
(step t27.t431 (cl (not (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule equiv1 :premises (t27.t430))
(step t27.t432 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule refl)
(step t27.t433 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule refl)
(step t27.t434 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule cong :premises (t27.t432 t27.t433))
(step t27.t435 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))
(step t27.t436 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) :rule trans :premises (t27.t434 t27.t435))
(step t27.t437 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule resolution :premises (t27.t431 t27.t436))
(step t27.t438 (cl (= y$n1s32 y$n1s32)) :rule refl)
(step t27.t439 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32))) :rule cong :premises (t27.t437 t27.t438))
(step t27.t440 (cl (= y$n0s32 y$n0s32)) :rule refl)
(step t27.t441 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule cong :premises (t27.t437 t27.t440))
(step t27.t442 (cl (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)))) :rule cong :premises (t27.t425 t27.t439 t27.t441))
(step t27.t443 (cl (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule hole)
(step t27.t444 (cl (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule resolution :premises (t27.t429 t27.t442 t27.t443))
(step t27.t445 (cl (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule resolution :premises (t27.t424 t27.t428 t27.t444))
(step t27.t446 (cl y$a_too_small_resistance$next (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule ite1 :premises (t27.t445))
(step t27.t447 (cl (= (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false) (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) :rule equiv_simplify)
(step t27.t448 (cl (not (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false)) (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv1 :premises (t27.t447))
(step t27.t449 (cl (not (= (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) (not (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule equiv_pos2)
(step t27.t450 (cl (= y$v3_1517448506_36$next_op y$v3_1517448506_36$next_op)) :rule refl)
(step t27.t451 (cl (= (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule cong :premises (t27.t450 t27.t437))
(step t27.t452 (cl (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule resolution :premises (t27.t449 t27.t451 a841))
(step t27.t453 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$v3_1517448506_36$next_op)) :rule symm :premises (t27.t452))
(step t27.t454 (cl (not y$1337) (= y$n1s32 y$v3_1517448506_36$next_op)) :rule equiv1 :premises (a842))
(step t27.t455 (cl (= y$n1s32 y$v3_1517448506_36$next_op) (not y$1337)) :rule reordering :premises (t27.t454))
(step t27.t456 (cl (not (and y$1334 y$1337)) y$1337) :rule and_pos :args (1))
(step t27.t457 (cl y$1337 (not (and y$1334 y$1337))) :rule reordering :premises (t27.t456))
(step t27.t458 (cl (not y$v3_1517448506_38$next_op) (and y$1334 y$1337)) :rule equiv1 :premises (a843))
(step t27.t459 (cl (and y$1334 y$1337) (not y$v3_1517448506_38$next_op)) :rule reordering :premises (t27.t458))
(step t27.t460 (cl (not (and y$1319 y$v3_1517448506_38$next_op)) y$v3_1517448506_38$next_op) :rule and_pos :args (1))
(step t27.t461 (cl y$v3_1517448506_38$next_op (not (and y$1319 y$v3_1517448506_38$next_op))) :rule reordering :premises (t27.t460))
(step t27.t462 (cl y$v3_1517448506_38$next_op) :rule resolution :premises (t27.t461 t27.t368))
(step t27.t463 (cl (and y$1334 y$1337)) :rule resolution :premises (t27.t459 t27.t462))
(step t27.t464 (cl y$1337) :rule resolution :premises (t27.t457 t27.t463))
(step t27.t465 (cl (= y$n1s32 y$v3_1517448506_36$next_op)) :rule resolution :premises (t27.t455 t27.t464))
(step t27.t466 (cl (= y$v3_1517448506_36$next_op y$n1s32)) :rule symm :premises (t27.t465))
(step t27.t467 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32)) :rule trans :premises (t27.t453 t27.t466))
(step t27.t468 (cl (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 y$n1s32))) :rule cong :premises (t27.t440 t27.t467))
(step t27.t469 (cl (= (= (= y$n0s32 y$n1s32) false) (not (= y$n0s32 y$n1s32)))) :rule equiv_simplify)
(step t27.t470 (cl (= (= y$n0s32 y$n1s32) false) (not (not (= y$n0s32 y$n1s32)))) :rule equiv2 :premises (t27.t469))
(step t27.t471 (cl (not (not (not (= y$n0s32 y$n1s32)))) (= y$n0s32 y$n1s32)) :rule not_not)
(step t27.t472 (cl (= (= y$n0s32 y$n1s32) false) (= y$n0s32 y$n1s32)) :rule resolution :premises (t27.t470 t27.t471))
(step t27.t473 (cl (not (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) (not (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32)) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule equiv_pos2)
(step t27.t474 (cl (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n0s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n1000s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n0s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n1000s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n6200s32 y$n0s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n5999s32)) (not (= y$n6200s32 y$n1000s32)) (not (= y$n6200s32 y$n5800s32)) (not (= y$n6200s32 y$n5s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n999s32 y$n5999s32)) (not (= y$n999s32 y$n1000s32)) (not (= y$n999s32 y$n5800s32)) (not (= y$n999s32 y$n5s32)) (not (= y$n5999s32 y$n1000s32)) (not (= y$n5999s32 y$n5800s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule distinct_elim)
(step t27.t475 (cl (= (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n1s32)))) :rule refl)
(step t27.t476 (cl (= (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n6200s32)))) :rule refl)
(step t27.t477 (cl (= (= y$n16s32 y$n0s32) (= y$n0s32 y$n16s32))) :rule rare_rewrite :args ("eq-symm" y$n16s32 y$n0s32))
(step t27.t478 (cl (= (not (= y$n16s32 y$n0s32)) (not (= y$n0s32 y$n16s32)))) :rule cong :premises (t27.t477))
(step t27.t479 (cl (= (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n999s32)))) :rule refl)
(step t27.t480 (cl (= (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n5999s32)))) :rule refl)
(step t27.t481 (cl (= (= y$n16s32 y$n1000s32) (= y$n1000s32 y$n16s32))) :rule rare_rewrite :args ("eq-symm" y$n16s32 y$n1000s32))
(step t27.t482 (cl (= (not (= y$n16s32 y$n1000s32)) (not (= y$n1000s32 y$n16s32)))) :rule cong :premises (t27.t481))
(step t27.t483 (cl (= (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5800s32)))) :rule refl)
(step t27.t484 (cl (= (not (= y$n16s32 y$n5s32)) (not (= y$n16s32 y$n5s32)))) :rule refl)
(step t27.t485 (cl (= (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n6200s32)))) :rule refl)
(step t27.t486 (cl (= (= y$n1s32 y$n0s32) (= y$n0s32 y$n1s32))) :rule rare_rewrite :args ("eq-symm" y$n1s32 y$n0s32))
(step t27.t487 (cl (= (not (= y$n1s32 y$n0s32)) (not (= y$n0s32 y$n1s32)))) :rule cong :premises (t27.t486))
(step t27.t488 (cl (= (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n999s32)))) :rule refl)
(step t27.t489 (cl (= (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n5999s32)))) :rule refl)
(step t27.t490 (cl (= (= y$n1s32 y$n1000s32) (= y$n1000s32 y$n1s32))) :rule rare_rewrite :args ("eq-symm" y$n1s32 y$n1000s32))
(step t27.t491 (cl (= (not (= y$n1s32 y$n1000s32)) (not (= y$n1000s32 y$n1s32)))) :rule cong :premises (t27.t490))
(step t27.t492 (cl (= (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5800s32)))) :rule refl)
(step t27.t493 (cl (= (not (= y$n1s32 y$n5s32)) (not (= y$n1s32 y$n5s32)))) :rule refl)
(step t27.t494 (cl (= (= y$n6200s32 y$n0s32) (= y$n0s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n0s32))
(step t27.t495 (cl (= (not (= y$n6200s32 y$n0s32)) (not (= y$n0s32 y$n6200s32)))) :rule cong :premises (t27.t494))
(step t27.t496 (cl (= (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n999s32)))) :rule refl)
(step t27.t497 (cl (= (= y$n6200s32 y$n5999s32) (= y$n5999s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5999s32))
(step t27.t498 (cl (= (not (= y$n6200s32 y$n5999s32)) (not (= y$n5999s32 y$n6200s32)))) :rule cong :premises (t27.t497))
(step t27.t499 (cl (= (= y$n6200s32 y$n1000s32) (= y$n1000s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n1000s32))
(step t27.t500 (cl (= (not (= y$n6200s32 y$n1000s32)) (not (= y$n1000s32 y$n6200s32)))) :rule cong :premises (t27.t499))
(step t27.t501 (cl (= (= y$n6200s32 y$n5800s32) (= y$n5800s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5800s32))
(step t27.t502 (cl (= (not (= y$n6200s32 y$n5800s32)) (not (= y$n5800s32 y$n6200s32)))) :rule cong :premises (t27.t501))
(step t27.t503 (cl (= (= y$n6200s32 y$n5s32) (= y$n5s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5s32))
(step t27.t504 (cl (= (not (= y$n6200s32 y$n5s32)) (not (= y$n5s32 y$n6200s32)))) :rule cong :premises (t27.t503))
(step t27.t505 (cl (= (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n999s32)))) :rule refl)
(step t27.t506 (cl (= (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n5999s32)))) :rule refl)
(step t27.t507 (cl (= (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n1000s32)))) :rule refl)
(step t27.t508 (cl (= (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5800s32)))) :rule refl)
(step t27.t509 (cl (= (not (= y$n0s32 y$n5s32)) (not (= y$n0s32 y$n5s32)))) :rule refl)
(step t27.t510 (cl (= (= y$n999s32 y$n5999s32) (= y$n5999s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5999s32))
(step t27.t511 (cl (= (not (= y$n999s32 y$n5999s32)) (not (= y$n5999s32 y$n999s32)))) :rule cong :premises (t27.t510))
(step t27.t512 (cl (= (= y$n999s32 y$n1000s32) (= y$n1000s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n1000s32))
(step t27.t513 (cl (= (not (= y$n999s32 y$n1000s32)) (not (= y$n1000s32 y$n999s32)))) :rule cong :premises (t27.t512))
(step t27.t514 (cl (= (= y$n999s32 y$n5800s32) (= y$n5800s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5800s32))
(step t27.t515 (cl (= (not (= y$n999s32 y$n5800s32)) (not (= y$n5800s32 y$n999s32)))) :rule cong :premises (t27.t514))
(step t27.t516 (cl (= (= y$n999s32 y$n5s32) (= y$n5s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5s32))
(step t27.t517 (cl (= (not (= y$n999s32 y$n5s32)) (not (= y$n5s32 y$n999s32)))) :rule cong :premises (t27.t516))
(step t27.t518 (cl (= (= y$n5999s32 y$n1000s32) (= y$n1000s32 y$n5999s32))) :rule rare_rewrite :args ("eq-symm" y$n5999s32 y$n1000s32))
(step t27.t519 (cl (= (not (= y$n5999s32 y$n1000s32)) (not (= y$n1000s32 y$n5999s32)))) :rule cong :premises (t27.t518))
(step t27.t520 (cl (= (= y$n5999s32 y$n5800s32) (= y$n5800s32 y$n5999s32))) :rule rare_rewrite :args ("eq-symm" y$n5999s32 y$n5800s32))
(step t27.t521 (cl (= (not (= y$n5999s32 y$n5800s32)) (not (= y$n5800s32 y$n5999s32)))) :rule cong :premises (t27.t520))
(step t27.t522 (cl (= (not (= y$n5999s32 y$n5s32)) (not (= y$n5999s32 y$n5s32)))) :rule refl)
(step t27.t523 (cl (= (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5800s32)))) :rule refl)
(step t27.t524 (cl (= (not (= y$n1000s32 y$n5s32)) (not (= y$n1000s32 y$n5s32)))) :rule refl)
(step t27.t525 (cl (= (not (= y$n5800s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule refl)
(step t27.t526 (cl (= (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n0s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n1000s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n0s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n1000s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n6200s32 y$n0s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n5999s32)) (not (= y$n6200s32 y$n1000s32)) (not (= y$n6200s32 y$n5800s32)) (not (= y$n6200s32 y$n5s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n999s32 y$n5999s32)) (not (= y$n999s32 y$n1000s32)) (not (= y$n999s32 y$n5800s32)) (not (= y$n999s32 y$n5s32)) (not (= y$n5999s32 y$n1000s32)) (not (= y$n5999s32 y$n5800s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule cong :premises (t27.t475 t27.t476 t27.t478 t27.t479 t27.t480 t27.t482 t27.t483 t27.t484 t27.t485 t27.t487 t27.t488 t27.t489 t27.t491 t27.t492 t27.t493 t27.t495 t27.t496 t27.t498 t27.t500 t27.t502 t27.t504 t27.t505 t27.t506 t27.t507 t27.t508 t27.t509 t27.t511 t27.t513 t27.t515 t27.t517 t27.t519 t27.t521 t27.t522 t27.t523 t27.t524 t27.t525))
(step t27.t527 (cl (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule trans :premises (t27.t474 t27.t526))
(step t27.t528 (cl (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule resolution :premises (t27.t473 t27.t527 a1))
(step t27.t529 (cl (not (= y$n0s32 y$n1s32))) :rule and :premises (t27.t528) :args (9))
(step t27.t530 (cl (= (= y$n0s32 y$n1s32) false)) :rule resolution :premises (t27.t472 t27.t529))
(step t27.t531 (cl (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false)) :rule trans :premises (t27.t468 t27.t530))
(step t27.t532 (cl (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule resolution :premises (t27.t448 t27.t531))
(step t27.t533 (cl y$a_too_small_resistance$next) :rule resolution :premises (t27.t446 t27.t532))
(step t27.t534 (cl (not y$274) (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) :rule equiv1 :premises (a138))
(step t27.t535 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op) (not y$274)) :rule reordering :premises (t27.t534))
(step t27.t536 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$274) :rule and_pos :args (12))
(step t27.t537 (cl y$274 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t27.t536))
(step t27.t538 (cl y$274) :rule resolution :premises (t27.t537 t27.t14))
(step t27.t539 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) :rule resolution :premises (t27.t535 t27.t538))
(step t27.t540 (cl y$a_too_small_resistance$next_rhs_op) :rule resolution :premises (t27.t423 t27.t533 t27.t539))
(step t27.t541 (cl (not y$1143)) :rule resolution :premises (t27.t185 t27.t421 t27.t540))
(step t27.t542 (cl y$v3_1517448506_562_op) :rule resolution :premises (t27.t183 t27.t541))
(step t27.t543 (cl (and y$1139 y$1140)) :rule resolution :premises (t27.t171 t27.t542))
(step t27.t544 (cl y$1140) :rule resolution :premises (t27.t169 t27.t543))
(step t27.t545 (cl (not y$1131)) :rule resolution :premises (t27.t167 t27.t544))
(step t27.t546 (cl y$v3_1517448506_557_op) :rule resolution :premises (t27.t166 t27.t545))
(step t27.t547 (cl (and y$1127 y$1128)) :rule resolution :premises (t27.t154 t27.t546))
(step t27.t548 (cl y$1128) :rule resolution :premises (t27.t152 t27.t547))
(step t27.t549 (cl (not y$1119)) :rule resolution :premises (t27.t150 t27.t548))
(step t27.t550 (cl y$v3_1517448506_552_op) :rule resolution :premises (t27.t149 t27.t549))
(step t27.t551 (cl (and y$1115 y$1116)) :rule resolution :premises (t27.t137 t27.t550))
(step t27.t552 (cl y$1116) :rule resolution :premises (t27.t135 t27.t551))
(step t27.t553 (cl (not y$1107)) :rule resolution :premises (t27.t133 t27.t552))
(step t27.t554 (cl y$v3_1517448506_547_op) :rule resolution :premises (t27.t132 t27.t553))
(step t27.t555 (cl (and y$1103 y$1104)) :rule resolution :premises (t27.t120 t27.t554))
(step t27.t556 (cl y$1104) :rule resolution :premises (t27.t118 t27.t555))
(step t27.t557 (cl (not y$209)) :rule resolution :premises (t27.t116 t27.t556))
(step t27.t558 (cl y$v3_1517448506_114_op) :rule resolution :premises (t27.t115 t27.t557))
(step t27.t559 (cl (and y$204 y$206)) :rule resolution :premises (t27.t103 t27.t558))
(step t27.t560 (cl y$206) :rule resolution :premises (t27.t101 t27.t559))
(step t27.t561 (cl y$v3_1517448506_174_op (not (and y$206 y$322))) :rule equiv2 :premises (a173))
(step t27.t562 (cl (and y$261 y$v3_1517448506_174_op) (not y$261) (not y$v3_1517448506_174_op)) :rule and_neg)
(step t27.t563 (cl (not y$261) (and y$261 y$v3_1517448506_174_op) (not y$v3_1517448506_174_op)) :rule reordering :premises (t27.t562))
(step t27.t564 (cl (not (and y$259 y$261)) y$261) :rule and_pos :args (1))
(step t27.t565 (cl y$261 (not (and y$259 y$261))) :rule reordering :premises (t27.t564))
(step t27.t566 (cl (not y$v3_1517448506_142_op) (and y$259 y$261)) :rule equiv1 :premises (a130))
(step t27.t567 (cl (and y$259 y$261) (not y$v3_1517448506_142_op)) :rule reordering :premises (t27.t566))
(step t27.t568 (cl (not (= (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op y$264))) (not (or y$v3_1517448506_142_op (not (not y$264)))) (or y$v3_1517448506_142_op y$264)) :rule equiv_pos2)
(step t27.t569 (cl (= y$v3_1517448506_142_op y$v3_1517448506_142_op)) :rule refl)
(step t27.t570 (cl (= (not (not y$264)) y$264)) :rule rare_rewrite :args ("bool-double-not-elim" y$264))
(step t27.t571 (cl (= (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op y$264))) :rule cong :premises (t27.t569 t27.t570))
(step t27.t572 (cl y$v3_1517448506_142_op (not (not y$264))) :rule equiv2 :premises (a131))
(step t27.t573 (cl (or y$v3_1517448506_142_op (not (not y$264))) (not y$v3_1517448506_142_op)) :rule or_neg :args (0))
(step t27.t574 (cl (or y$v3_1517448506_142_op (not (not y$264))) (not (not (not y$264)))) :rule or_neg :args (1))
(step t27.t575 (cl (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op (not (not y$264)))) :rule resolution :premises (t27.t572 t27.t573 t27.t574))
(step t27.t576 (cl (or y$v3_1517448506_142_op (not (not y$264)))) :rule contraction :premises (t27.t575))
(step t27.t577 (cl (or y$v3_1517448506_142_op y$264)) :rule resolution :premises (t27.t568 t27.t571 t27.t576))
(step t27.t578 (cl y$v3_1517448506_142_op y$264) :rule or :premises (t27.t577))
(step t27.t579 (cl y$264 y$v3_1517448506_142_op) :rule reordering :premises (t27.t578))
(step t27.t580 (cl (not y$264) (not y$1139)) :rule equiv1 :premises (a734))
(step t27.t581 (cl (not (and y$1139 y$1140)) y$1139) :rule and_pos :args (0))
(step t27.t582 (cl y$1139 (not (and y$1139 y$1140))) :rule reordering :premises (t27.t581))
(step t27.t583 (cl y$1139) :rule resolution :premises (t27.t582 t27.t543))
(step t27.t584 (cl (not y$264)) :rule resolution :premises (t27.t580 t27.t583))
(step t27.t585 (cl y$v3_1517448506_142_op) :rule resolution :premises (t27.t579 t27.t584))
(step t27.t586 (cl (and y$259 y$261)) :rule resolution :premises (t27.t567 t27.t585))
(step t27.t587 (cl y$261) :rule resolution :premises (t27.t565 t27.t586))
(step t27.t588 (cl y$a_S2$next_rhs_op (not (and y$261 y$v3_1517448506_174_op))) :rule equiv2 :premises (a174))
(step t27.t589 (cl (and y$a_S2$next_rhs_op y$1202) (not y$a_S2$next_rhs_op) (not y$1202)) :rule and_neg)
(step t27.t590 (cl (not y$1202) (and y$a_S2$next_rhs_op y$1202) (not y$a_S2$next_rhs_op)) :rule reordering :premises (t27.t589))
(step t27.t591 (cl (not (= (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op y$1202))) (not (or y$v3_1517448506_587_op (not (not y$1202)))) (or y$v3_1517448506_587_op y$1202)) :rule equiv_pos2)
(step t27.t592 (cl (= y$v3_1517448506_587_op y$v3_1517448506_587_op)) :rule refl)
(step t27.t593 (cl (= (not (not y$1202)) y$1202)) :rule rare_rewrite :args ("bool-double-not-elim" y$1202))
(step t27.t594 (cl (= (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op y$1202))) :rule cong :premises (t27.t592 t27.t593))
(step t27.t595 (cl y$v3_1517448506_587_op (not (not y$1202))) :rule equiv2 :premises (a781))
(step t27.t596 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (not y$v3_1517448506_587_op)) :rule or_neg :args (0))
(step t27.t597 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (not (not (not y$1202)))) :rule or_neg :args (1))
(step t27.t598 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op (not (not y$1202)))) :rule resolution :premises (t27.t595 t27.t596 t27.t597))
(step t27.t599 (cl (or y$v3_1517448506_587_op (not (not y$1202)))) :rule contraction :premises (t27.t598))
(step t27.t600 (cl (or y$v3_1517448506_587_op y$1202)) :rule resolution :premises (t27.t591 t27.t594 t27.t599))
(step t27.t601 (cl y$v3_1517448506_587_op y$1202) :rule or :premises (t27.t600))
(step t27.t602 (cl y$1202 y$v3_1517448506_587_op) :rule reordering :premises (t27.t601))
(step t27.t603 (cl (not y$v3_1517448506_587_op) (and y$1198 y$1199)) :rule equiv1 :premises (a780))
(step t27.t604 (cl (and y$1198 y$1199) (not y$v3_1517448506_587_op)) :rule reordering :premises (t27.t603))
(step t27.t605 (cl (not (and y$1198 y$1199)) y$1199) :rule and_pos :args (1))
(step t27.t606 (cl y$1199 (not (and y$1198 y$1199))) :rule reordering :premises (t27.t605))
(step t27.t607 (cl (not y$1190) (not y$1199)) :rule equiv1 :premises (a779))
(step t27.t608 (cl (not (= (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op y$1190))) (not (or y$v3_1517448506_582_op (not (not y$1190)))) (or y$v3_1517448506_582_op y$1190)) :rule equiv_pos2)
(step t27.t609 (cl (= y$v3_1517448506_582_op y$v3_1517448506_582_op)) :rule refl)
(step t27.t610 (cl (= (not (not y$1190)) y$1190)) :rule rare_rewrite :args ("bool-double-not-elim" y$1190))
(step t27.t611 (cl (= (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op y$1190))) :rule cong :premises (t27.t609 t27.t610))
(step t27.t612 (cl y$v3_1517448506_582_op (not (not y$1190))) :rule equiv2 :premises (a772))
(step t27.t613 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (not y$v3_1517448506_582_op)) :rule or_neg :args (0))
(step t27.t614 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (not (not (not y$1190)))) :rule or_neg :args (1))
(step t27.t615 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op (not (not y$1190)))) :rule resolution :premises (t27.t612 t27.t613 t27.t614))
(step t27.t616 (cl (or y$v3_1517448506_582_op (not (not y$1190)))) :rule contraction :premises (t27.t615))
(step t27.t617 (cl (or y$v3_1517448506_582_op y$1190)) :rule resolution :premises (t27.t608 t27.t611 t27.t616))
(step t27.t618 (cl y$v3_1517448506_582_op y$1190) :rule or :premises (t27.t617))
(step t27.t619 (cl y$1190 y$v3_1517448506_582_op) :rule reordering :premises (t27.t618))
(step t27.t620 (cl (not y$v3_1517448506_582_op) (and y$1186 y$1187)) :rule equiv1 :premises (a771))
(step t27.t621 (cl (and y$1186 y$1187) (not y$v3_1517448506_582_op)) :rule reordering :premises (t27.t620))
(step t27.t622 (cl (not (and y$1186 y$1187)) y$1187) :rule and_pos :args (1))
(step t27.t623 (cl y$1187 (not (and y$1186 y$1187))) :rule reordering :premises (t27.t622))
(step t27.t624 (cl (not y$1178) (not y$1187)) :rule equiv1 :premises (a770))
(step t27.t625 (cl (not (= (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op y$1178))) (not (or y$v3_1517448506_577_op (not (not y$1178)))) (or y$v3_1517448506_577_op y$1178)) :rule equiv_pos2)
(step t27.t626 (cl (= y$v3_1517448506_577_op y$v3_1517448506_577_op)) :rule refl)
(step t27.t627 (cl (= (not (not y$1178)) y$1178)) :rule rare_rewrite :args ("bool-double-not-elim" y$1178))
(step t27.t628 (cl (= (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op y$1178))) :rule cong :premises (t27.t626 t27.t627))
(step t27.t629 (cl y$v3_1517448506_577_op (not (not y$1178))) :rule equiv2 :premises (a763))
(step t27.t630 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (not y$v3_1517448506_577_op)) :rule or_neg :args (0))
(step t27.t631 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (not (not (not y$1178)))) :rule or_neg :args (1))
(step t27.t632 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op (not (not y$1178)))) :rule resolution :premises (t27.t629 t27.t630 t27.t631))
(step t27.t633 (cl (or y$v3_1517448506_577_op (not (not y$1178)))) :rule contraction :premises (t27.t632))
(step t27.t634 (cl (or y$v3_1517448506_577_op y$1178)) :rule resolution :premises (t27.t625 t27.t628 t27.t633))
(step t27.t635 (cl y$v3_1517448506_577_op y$1178) :rule or :premises (t27.t634))
(step t27.t636 (cl y$1178 y$v3_1517448506_577_op) :rule reordering :premises (t27.t635))
(step t27.t637 (cl (not y$v3_1517448506_577_op) (and y$293 y$1175)) :rule equiv1 :premises (a762))
(step t27.t638 (cl (and y$293 y$1175) (not y$v3_1517448506_577_op)) :rule reordering :premises (t27.t637))
(step t27.t639 (cl (not (and y$293 y$1175)) y$1175) :rule and_pos :args (1))
(step t27.t640 (cl y$1175 (not (and y$293 y$1175))) :rule reordering :premises (t27.t639))
(step t27.t641 (cl (not y$1167) (not y$1175)) :rule equiv1 :premises (a761))
(step t27.t642 (cl (not (= (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op y$1167))) (not (or y$v3_1517448506_572_op (not (not y$1167)))) (or y$v3_1517448506_572_op y$1167)) :rule equiv_pos2)
(step t27.t643 (cl (= y$v3_1517448506_572_op y$v3_1517448506_572_op)) :rule refl)
(step t27.t644 (cl (= (not (not y$1167)) y$1167)) :rule rare_rewrite :args ("bool-double-not-elim" y$1167))
(step t27.t645 (cl (= (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op y$1167))) :rule cong :premises (t27.t643 t27.t644))
(step t27.t646 (cl y$v3_1517448506_572_op (not (not y$1167))) :rule equiv2 :premises (a755))
(step t27.t647 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (not y$v3_1517448506_572_op)) :rule or_neg :args (0))
(step t27.t648 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (not (not (not y$1167)))) :rule or_neg :args (1))
(step t27.t649 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op (not (not y$1167)))) :rule resolution :premises (t27.t646 t27.t647 t27.t648))
(step t27.t650 (cl (or y$v3_1517448506_572_op (not (not y$1167)))) :rule contraction :premises (t27.t649))
(step t27.t651 (cl (or y$v3_1517448506_572_op y$1167)) :rule resolution :premises (t27.t642 t27.t645 t27.t650))
(step t27.t652 (cl y$v3_1517448506_572_op y$1167) :rule or :premises (t27.t651))
(step t27.t653 (cl y$1167 y$v3_1517448506_572_op) :rule reordering :premises (t27.t652))
(step t27.t654 (cl (not y$v3_1517448506_572_op) (and y$1163 y$1164)) :rule equiv1 :premises (a754))
(step t27.t655 (cl (and y$1163 y$1164) (not y$v3_1517448506_572_op)) :rule reordering :premises (t27.t654))
(step t27.t656 (cl (not (and y$1163 y$1164)) y$1164) :rule and_pos :args (1))
(step t27.t657 (cl y$1164 (not (and y$1163 y$1164))) :rule reordering :premises (t27.t656))
(step t27.t658 (cl (not y$1155) (not y$1164)) :rule equiv1 :premises (a753))
(step t27.t659 (cl (not (= (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op y$1155))) (not (or y$v3_1517448506_567_op (not (not y$1155)))) (or y$v3_1517448506_567_op y$1155)) :rule equiv_pos2)
(step t27.t660 (cl (= y$v3_1517448506_567_op y$v3_1517448506_567_op)) :rule refl)
(step t27.t661 (cl (= (not (not y$1155)) y$1155)) :rule rare_rewrite :args ("bool-double-not-elim" y$1155))
(step t27.t662 (cl (= (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op y$1155))) :rule cong :premises (t27.t660 t27.t661))
(step t27.t663 (cl y$v3_1517448506_567_op (not (not y$1155))) :rule equiv2 :premises (a746))
(step t27.t664 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (not y$v3_1517448506_567_op)) :rule or_neg :args (0))
(step t27.t665 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (not (not (not y$1155)))) :rule or_neg :args (1))
(step t27.t666 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op (not (not y$1155)))) :rule resolution :premises (t27.t663 t27.t664 t27.t665))
(step t27.t667 (cl (or y$v3_1517448506_567_op (not (not y$1155)))) :rule contraction :premises (t27.t666))
(step t27.t668 (cl (or y$v3_1517448506_567_op y$1155)) :rule resolution :premises (t27.t659 t27.t662 t27.t667))
(step t27.t669 (cl y$v3_1517448506_567_op y$1155) :rule or :premises (t27.t668))
(step t27.t670 (cl y$1155 y$v3_1517448506_567_op) :rule reordering :premises (t27.t669))
(step t27.t671 (cl (not y$v3_1517448506_567_op) (and y$1151 y$1152)) :rule equiv1 :premises (a745))
(step t27.t672 (cl (and y$1151 y$1152) (not y$v3_1517448506_567_op)) :rule reordering :premises (t27.t671))
(step t27.t673 (cl (not (and y$1151 y$1152)) y$1151) :rule and_pos :args (0))
(step t27.t674 (cl y$1151 (not (and y$1151 y$1152))) :rule reordering :premises (t27.t673))
(step t27.t675 (cl (not y$a_too_small_resistance$next_rhs_op) (not y$1151)) :rule equiv1 :premises (a743))
(step t27.t676 (cl (not y$1151) (not y$a_too_small_resistance$next_rhs_op)) :rule reordering :premises (t27.t675))
(step t27.t677 (cl (not y$1151)) :rule resolution :premises (t27.t676 t27.t540))
(step t27.t678 (cl (not (and y$1151 y$1152))) :rule resolution :premises (t27.t674 t27.t677))
(step t27.t679 (cl (not y$v3_1517448506_567_op)) :rule resolution :premises (t27.t672 t27.t678))
(step t27.t680 (cl y$1155) :rule resolution :premises (t27.t670 t27.t679))
(step t27.t681 (cl (not y$1164)) :rule resolution :premises (t27.t658 t27.t680))
(step t27.t682 (cl (not (and y$1163 y$1164))) :rule resolution :premises (t27.t657 t27.t681))
(step t27.t683 (cl (not y$v3_1517448506_572_op)) :rule resolution :premises (t27.t655 t27.t682))
(step t27.t684 (cl y$1167) :rule resolution :premises (t27.t653 t27.t683))
(step t27.t685 (cl (not y$1175)) :rule resolution :premises (t27.t641 t27.t684))
(step t27.t686 (cl (not (and y$293 y$1175))) :rule resolution :premises (t27.t640 t27.t685))
(step t27.t687 (cl (not y$v3_1517448506_577_op)) :rule resolution :premises (t27.t638 t27.t686))
(step t27.t688 (cl y$1178) :rule resolution :premises (t27.t636 t27.t687))
(step t27.t689 (cl (not y$1187)) :rule resolution :premises (t27.t624 t27.t688))
(step t27.t690 (cl (not (and y$1186 y$1187))) :rule resolution :premises (t27.t623 t27.t689))
(step t27.t691 (cl (not y$v3_1517448506_582_op)) :rule resolution :premises (t27.t621 t27.t690))
(step t27.t692 (cl y$1190) :rule resolution :premises (t27.t619 t27.t691))
(step t27.t693 (cl (not y$1199)) :rule resolution :premises (t27.t607 t27.t692))
(step t27.t694 (cl (not (and y$1198 y$1199))) :rule resolution :premises (t27.t606 t27.t693))
(step t27.t695 (cl (not y$v3_1517448506_587_op)) :rule resolution :premises (t27.t604 t27.t694))
(step t27.t696 (cl y$1202) :rule resolution :premises (t27.t602 t27.t695))
(step t27.t697 (cl y$v3_1517448506_588_op (not (and y$a_S2$next_rhs_op y$1202))) :rule equiv2 :premises (a782))
(step t27.t698 (cl (not y$v3_1517448506_588_op) (not y$1205)) :rule equiv1 :premises (a783))
(step t27.t699 (cl (not y$1205) (not y$v3_1517448506_588_op)) :rule reordering :premises (t27.t698))
(step t27.t700 (cl (not (and y$1197 y$1205)) y$1205) :rule and_pos :args (1))
(step t27.t701 (cl y$1205 (not (and y$1197 y$1205))) :rule reordering :premises (t27.t700))
(step t27.t702 (cl y$1205) :rule resolution :premises (t27.t701 t27.t398))
(step t27.t703 (cl (not y$v3_1517448506_588_op)) :rule resolution :premises (t27.t699 t27.t702))
(step t27.t704 (cl (not (and y$a_S2$next_rhs_op y$1202))) :rule resolution :premises (t27.t697 t27.t703))
(step t27.t705 (cl (not y$a_S2$next_rhs_op)) :rule resolution :premises (t27.t590 t27.t696 t27.t704))
(step t27.t706 (cl (not (and y$261 y$v3_1517448506_174_op))) :rule resolution :premises (t27.t588 t27.t705))
(step t27.t707 (cl (not y$v3_1517448506_174_op)) :rule resolution :premises (t27.t563 t27.t587 t27.t706))
(step t27.t708 (cl (not (and y$206 y$322))) :rule resolution :premises (t27.t561 t27.t707))
(step t27.t709 (cl (not y$322)) :rule resolution :premises (t27.t99 t27.t560 t27.t708))
(step t27.t710 (cl y$v3_1517448506_173_op) :rule resolution :premises (t27.t97 t27.t709))
(step t27.t711 (cl (and y$3 y$315)) :rule resolution :premises (t27.t85 t27.t710))
(step t27.t712 (cl y$315) :rule resolution :premises (t27.t83 t27.t711))
(step t27.t713 (cl y$a_S1$next_rhs_op (not (and y$313 y$315))) :rule equiv2 :premises (a168))
(step t27.t714 (cl (and y$a_S1$next_rhs_op y$1190) (not y$a_S1$next_rhs_op) (not y$1190)) :rule and_neg)
(step t27.t715 (cl (not y$1190) (and y$a_S1$next_rhs_op y$1190) (not y$a_S1$next_rhs_op)) :rule reordering :premises (t27.t714))
(step t27.t716 (cl y$v3_1517448506_583_op (not (and y$a_S1$next_rhs_op y$1190))) :rule equiv2 :premises (a773))
(step t27.t717 (cl (not y$v3_1517448506_583_op) (not y$1193)) :rule equiv1 :premises (a774))
(step t27.t718 (cl (not y$1193) (not y$v3_1517448506_583_op)) :rule reordering :premises (t27.t717))
(step t27.t719 (cl (not (and y$1185 y$1193)) y$1193) :rule and_pos :args (1))
(step t27.t720 (cl y$1193 (not (and y$1185 y$1193))) :rule reordering :premises (t27.t719))
(step t27.t721 (cl y$1193) :rule resolution :premises (t27.t720 t27.t402))
(step t27.t722 (cl (not y$v3_1517448506_583_op)) :rule resolution :premises (t27.t718 t27.t721))
(step t27.t723 (cl (not (and y$a_S1$next_rhs_op y$1190))) :rule resolution :premises (t27.t716 t27.t722))
(step t27.t724 (cl (not y$a_S1$next_rhs_op)) :rule resolution :premises (t27.t715 t27.t692 t27.t723))
(step t27.t725 (cl (not (and y$313 y$315))) :rule resolution :premises (t27.t713 t27.t724))
(step t27.t726 (cl (not y$313)) :rule resolution :premises (t27.t81 t27.t712 t27.t725))
(step t27.t727 (cl y$v3_1517448506_168_op) :rule resolution :premises (t27.t79 t27.t726))
(step t27.t728 (cl (and y$1 y$161)) :rule resolution :premises (t27.t67 t27.t727))
(step t27.t729 (cl y$161) :rule resolution :premises (t27.t65 t27.t728))
(step t27.t730 (cl y$a_new_range$next_rhs_op (not (and y$161 y$225))) :rule equiv2 :premises (a107))
(step t27.t731 (cl (not y$a_new_range$next_rhs_op) (not y$1103)) :rule equiv1 :premises (a707))
(step t27.t732 (cl (not y$1103) (not y$a_new_range$next_rhs_op)) :rule reordering :premises (t27.t731))
(step t27.t733 (cl (not (and y$1103 y$1104)) y$1103) :rule and_pos :args (0))
(step t27.t734 (cl y$1103 (not (and y$1103 y$1104))) :rule reordering :premises (t27.t733))
(step t27.t735 (cl y$1103) :rule resolution :premises (t27.t734 t27.t555))
(step t27.t736 (cl (not y$a_new_range$next_rhs_op)) :rule resolution :premises (t27.t732 t27.t735))
(step t27.t737 (cl (not (and y$161 y$225))) :rule resolution :premises (t27.t730 t27.t736))
(step t27.t738 (cl (not y$225)) :rule resolution :premises (t27.t63 t27.t729 t27.t737))
(step t27.t739 (cl y$v3_1517448506_121_op) :rule resolution :premises (t27.t61 t27.t738))
(step t27.t740 (cl (and y$221 y$222)) :rule resolution :premises (t27.t49 t27.t739))
(step t27.t741 (cl y$222) :rule resolution :premises (t27.t47 t27.t740))
(step t27.t742 (cl (not y$f11)) :rule resolution :premises (t27.t45 t27.t741))
(step t27.t743 (cl (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t27.t43 t27.t742))
(step t27.t744 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) :rule symm :premises (t27.t743))
(step t27.t745 (cl (not (= (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule equiv_pos2)
(step t27.t746 (cl (= (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule equiv_simplify)
(step t27.t747 (cl (not (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule equiv1 :premises (t27.t746))
(step t27.t748 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule refl)
(step t27.t749 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule refl)
(step t27.t750 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule cong :premises (t27.t748 t27.t749))
(step t27.t751 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f10 y$127 y$v3_1517448506_80_op)))
(step t27.t752 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) :rule trans :premises (t27.t750 t27.t751))
(step t27.t753 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t27.t747 t27.t752))
(step t27.t754 (cl (= (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule cong :premises (t27.t37 t27.t753))
(step t27.t755 (cl (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t27.t745 t27.t754 t27.a4))
(step t27.t756 (cl (not (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) (not (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule equiv_pos2)
(step t27.t757 (cl (= y$f10 y$f10)) :rule refl)
(step t27.t758 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f10 y$127 y$v3_1517448506_80_op) y$127))
(step t27.t759 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))
(step t27.t760 (cl (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule cong :premises (t27.t757 t27.t758 t27.t759))
(step t27.t761 (cl (not (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)))) (not (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule equiv_pos2)
(step t27.t762 (cl (= y$127 y$127)) :rule refl)
(step t27.t763 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127))) :rule cong :premises (t27.t753 t27.t762))
(step t27.t764 (cl (= y$v3_1517448506_80_op y$v3_1517448506_80_op)) :rule refl)
(step t27.t765 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule cong :premises (t27.t753 t27.t764))
(step t27.t766 (cl (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)))) :rule cong :premises (t27.t757 t27.t763 t27.t765))
(step t27.t767 (cl (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule hole)
(step t27.t768 (cl (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule resolution :premises (t27.t761 t27.t766 t27.t767))
(step t27.t769 (cl (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule resolution :premises (t27.t756 t27.t760 t27.t768))
(step t27.t770 (cl y$f10 (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule ite1 :premises (t27.t769))
(step t27.t771 (cl (not y$f10) (not y$217)) :rule equiv1 :premises (a100))
(step t27.t772 (cl (not y$217) (not y$f10)) :rule reordering :premises (t27.t771))
(step t27.t773 (cl (not (and y$216 y$217)) y$217) :rule and_pos :args (1))
(step t27.t774 (cl y$217 (not (and y$216 y$217))) :rule reordering :premises (t27.t773))
(step t27.t775 (cl (not y$v3_1517448506_119_op) (and y$216 y$217)) :rule equiv1 :premises (a101))
(step t27.t776 (cl (and y$216 y$217) (not y$v3_1517448506_119_op)) :rule reordering :premises (t27.t775))
(step t27.t777 (cl (not (= (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op y$220))) (not (or y$v3_1517448506_119_op (not (not y$220)))) (or y$v3_1517448506_119_op y$220)) :rule equiv_pos2)
(step t27.t778 (cl (= y$v3_1517448506_119_op y$v3_1517448506_119_op)) :rule refl)
(step t27.t779 (cl (= (not (not y$220)) y$220)) :rule rare_rewrite :args ("bool-double-not-elim" y$220))
(step t27.t780 (cl (= (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op y$220))) :rule cong :premises (t27.t778 t27.t779))
(step t27.t781 (cl y$v3_1517448506_119_op (not (not y$220))) :rule equiv2 :premises (a102))
(step t27.t782 (cl (or y$v3_1517448506_119_op (not (not y$220))) (not y$v3_1517448506_119_op)) :rule or_neg :args (0))
(step t27.t783 (cl (or y$v3_1517448506_119_op (not (not y$220))) (not (not (not y$220)))) :rule or_neg :args (1))
(step t27.t784 (cl (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op (not (not y$220)))) :rule resolution :premises (t27.t781 t27.t782 t27.t783))
(step t27.t785 (cl (or y$v3_1517448506_119_op (not (not y$220)))) :rule contraction :premises (t27.t784))
(step t27.t786 (cl (or y$v3_1517448506_119_op y$220)) :rule resolution :premises (t27.t777 t27.t780 t27.t785))
(step t27.t787 (cl y$v3_1517448506_119_op y$220) :rule or :premises (t27.t786))
(step t27.t788 (cl y$220 y$v3_1517448506_119_op) :rule reordering :premises (t27.t787))
(step t27.t789 (cl (not y$220) (not y$221)) :rule equiv1 :premises (a103))
(step t27.t790 (cl (not (and y$221 y$222)) y$221) :rule and_pos :args (0))
(step t27.t791 (cl y$221 (not (and y$221 y$222))) :rule reordering :premises (t27.t790))
(step t27.t792 (cl y$221) :rule resolution :premises (t27.t791 t27.t740))
(step t27.t793 (cl (not y$220)) :rule resolution :premises (t27.t789 t27.t792))
(step t27.t794 (cl y$v3_1517448506_119_op) :rule resolution :premises (t27.t788 t27.t793))
(step t27.t795 (cl (and y$216 y$217)) :rule resolution :premises (t27.t776 t27.t794))
(step t27.t796 (cl y$217) :rule resolution :premises (t27.t774 t27.t795))
(step t27.t797 (cl (not y$f10)) :rule resolution :premises (t27.t772 t27.t796))
(step t27.t798 (cl (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t27.t770 t27.t797))
(step t27.t799 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) :rule symm :premises (t27.t798))
(step t27.t800 (cl (not (= (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule equiv_pos2)
(step t27.t801 (cl (= (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)))) :rule equiv_simplify)
(step t27.t802 (cl (not (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule equiv1 :premises (t27.t801))
(step t27.t803 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule refl)
(step t27.t804 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule refl)
(step t27.t805 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)))) :rule cong :premises (t27.t803 t27.t804))
(step t27.t806 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f09 y$132 y$v_range)))
(step t27.t807 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) :rule trans :premises (t27.t805 t27.t806))
(step t27.t808 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t27.t802 t27.t807))
(step t27.t809 (cl (= (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))) :rule cong :premises (t27.t764 t27.t808))
(step t27.t810 (cl (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t27.t800 t27.t809 t27.a2))
(step t27.t811 (cl (not (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))))) (not (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule equiv_pos2)
(step t27.t812 (cl (= y$f09 y$f09)) :rule refl)
(step t27.t813 (cl (= (= (ite y$f09 y$132 y$v_range) y$132) (= y$132 (ite y$f09 y$132 y$v_range)))) :rule rare_rewrite :args ("eq-symm" (ite y$f09 y$132 y$v_range) y$132))
(step t27.t814 (cl (= (= (ite y$f09 y$132 y$v_range) y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule rare_rewrite :args ("eq-symm" (ite y$f09 y$132 y$v_range) y$v_range))
(step t27.t815 (cl (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))))) :rule cong :premises (t27.t812 t27.t813 t27.t814))
(step t27.t816 (cl (not (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)))) (not (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule equiv_pos2)
(step t27.t817 (cl (= y$132 y$132)) :rule refl)
(step t27.t818 (cl (= (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$132))) :rule cong :premises (t27.t808 t27.t817))
(step t27.t819 (cl (= y$v_range y$v_range)) :rule refl)
(step t27.t820 (cl (= (= (ite y$f09 y$132 y$v_range) y$v_range) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule cong :premises (t27.t808 t27.t819))
(step t27.t821 (cl (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)))) :rule cong :premises (t27.t812 t27.t818 t27.t820))
(step t27.t822 (cl (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule hole)
(step t27.t823 (cl (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule resolution :premises (t27.t816 t27.t821 t27.t822))
(step t27.t824 (cl (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule resolution :premises (t27.t811 t27.t815 t27.t823))
(step t27.t825 (cl y$f09 (= y$v_range (ite y$f09 y$132 y$v_range))) :rule ite1 :premises (t27.t824))
(step t27.t826 (cl (not y$f09) (not y$212)) :rule equiv1 :premises (a96))
(step t27.t827 (cl (not y$212) (not y$f09)) :rule reordering :premises (t27.t826))
(step t27.t828 (cl (not (and y$25 y$212)) y$212) :rule and_pos :args (1))
(step t27.t829 (cl y$212 (not (and y$25 y$212))) :rule reordering :premises (t27.t828))
(step t27.t830 (cl (not y$v3_1517448506_117_op) (and y$25 y$212)) :rule equiv1 :premises (a97))
(step t27.t831 (cl (and y$25 y$212) (not y$v3_1517448506_117_op)) :rule reordering :premises (t27.t830))
(step t27.t832 (cl (not (= (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op y$215))) (not (or y$v3_1517448506_117_op (not (not y$215)))) (or y$v3_1517448506_117_op y$215)) :rule equiv_pos2)
(step t27.t833 (cl (= y$v3_1517448506_117_op y$v3_1517448506_117_op)) :rule refl)
(step t27.t834 (cl (= (not (not y$215)) y$215)) :rule rare_rewrite :args ("bool-double-not-elim" y$215))
(step t27.t835 (cl (= (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op y$215))) :rule cong :premises (t27.t833 t27.t834))
(step t27.t836 (cl y$v3_1517448506_117_op (not (not y$215))) :rule equiv2 :premises (a98))
(step t27.t837 (cl (or y$v3_1517448506_117_op (not (not y$215))) (not y$v3_1517448506_117_op)) :rule or_neg :args (0))
(step t27.t838 (cl (or y$v3_1517448506_117_op (not (not y$215))) (not (not (not y$215)))) :rule or_neg :args (1))
(step t27.t839 (cl (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op (not (not y$215)))) :rule resolution :premises (t27.t836 t27.t837 t27.t838))
(step t27.t840 (cl (or y$v3_1517448506_117_op (not (not y$215)))) :rule contraction :premises (t27.t839))
(step t27.t841 (cl (or y$v3_1517448506_117_op y$215)) :rule resolution :premises (t27.t832 t27.t835 t27.t840))
(step t27.t842 (cl y$v3_1517448506_117_op y$215) :rule or :premises (t27.t841))
(step t27.t843 (cl y$215 y$v3_1517448506_117_op) :rule reordering :premises (t27.t842))
(step t27.t844 (cl (not y$215) (not y$216)) :rule equiv1 :premises (a99))
(step t27.t845 (cl (not (and y$216 y$217)) y$216) :rule and_pos :args (0))
(step t27.t846 (cl y$216 (not (and y$216 y$217))) :rule reordering :premises (t27.t845))
(step t27.t847 (cl y$216) :rule resolution :premises (t27.t846 t27.t795))
(step t27.t848 (cl (not y$215)) :rule resolution :premises (t27.t844 t27.t847))
(step t27.t849 (cl y$v3_1517448506_117_op) :rule resolution :premises (t27.t843 t27.t848))
(step t27.t850 (cl (and y$25 y$212)) :rule resolution :premises (t27.t831 t27.t849))
(step t27.t851 (cl y$212) :rule resolution :premises (t27.t829 t27.t850))
(step t27.t852 (cl (not y$f09)) :rule resolution :premises (t27.t827 t27.t851))
(step t27.t853 (cl (= y$v_range (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t27.t825 t27.t852))
(step t27.t854 (cl (= (ite y$f09 y$132 y$v_range) y$v_range)) :rule symm :premises (t27.t853))
(step t27.t855 (cl (= y$v_range$next y$v_range)) :rule trans :premises (t27.t16 t27.t28 t27.t744 t27.t755 t27.t799 t27.t810 t27.t854))
(step t27.t856 (cl (= (Extract_1_15_15_16 y$v_range$next) (Extract_1_15_15_16 y$v_range))) :rule cong :premises (t27.t855))
(step t27.t857 (cl (= (= (Extract_1_15_15_16 y$v_range) true) (Extract_1_15_15_16 y$v_range))) :rule equiv_simplify)
(step t27.t858 (cl (= (Extract_1_15_15_16 y$v_range) true) (not (Extract_1_15_15_16 y$v_range))) :rule equiv2 :premises (t27.t857))
(step t27.t859 (cl (= (Extract_1_15_15_16 y$v_range) true)) :rule resolution :premises (t27.t858 t27.a0))
(step t27.t860 (cl (= (Extract_1_15_15_16 y$v_range$next) true)) :rule trans :premises (t27.t856 t27.t859))
(step t27.t861 (cl (Extract_1_15_15_16 y$v_range$next)) :rule resolution :premises (t27.t1 t27.t860))
(step t27 (cl (not (Extract_1_15_15_16 y$v_range)) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (Extract_1_15_15_16 y$v_range$next)) :rule subproof :discharge (t27.a0 t27.a1 t27.a2 t27.a3 t27.a4 t27.a5 t27.a6 t27.a7))
(step t28 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (Extract_1_15_15_16 y$v_range)) :rule and_pos :args (0))
(step t29 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (= y$v_range (ite y$f09 y$132 y$v_range))) :rule and_pos :args (1))
(step t30 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule and_pos :args (2))
(step t31 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (3))
(step t32 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (4))
(step t33 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (5))
(step t34 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (6))
(step t35 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (= y$v_range$next y$v_range$next_rhs_op)) :rule and_pos :args (7))
(step t36 (cl (Extract_1_15_15_16 y$v_range$next) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op)))) :rule resolution :premises (t27 t28 t29 t30 t31 t32 t33 t34 t35))
(step t37 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (Extract_1_15_15_16 y$v_range$next)) :rule reordering :premises (t36))
(step t38 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (Extract_1_15_15_16 y$v_range$next)) :rule contraction :premises (t37))
(step t39 (cl (=> (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op)) (Extract_1_15_15_16 y$v_range$next)) (Extract_1_15_15_16 y$v_range$next)) :rule resolution :premises (t26 t38))
(step t40 (cl (=> (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op)) (Extract_1_15_15_16 y$v_range$next)) (not (Extract_1_15_15_16 y$v_range$next))) :rule implies_neg2)
(step t41 (cl (=> (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op)) (Extract_1_15_15_16 y$v_range$next)) (=> (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op)) (Extract_1_15_15_16 y$v_range$next))) :rule resolution :premises (t39 t40))
(step t42 (cl (=> (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op)) (Extract_1_15_15_16 y$v_range$next))) :rule contraction :premises (t41))
(step t43 (cl (not (and (Extract_1_15_15_16 y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op))) (Extract_1_15_15_16 y$v_range$next)) :rule implies :premises (t42))
(step t44 (cl (not (Extract_1_15_15_16 y$v_range)) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (Extract_1_15_15_16 y$v_range$next)) :rule resolution :premises (t25 t43))
(step t45 (cl (Extract_1_15_15_16 y$v_range$next) (not (Extract_1_15_15_16 y$v_range)) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule reordering :premises (t44))
(step t46 (cl (not y$111) (Extract_1_15_15_16 y$v_range)) :rule equiv1 :premises (a3))
(step t47 (cl (Extract_1_15_15_16 y$v_range) (not y$111)) :rule reordering :premises (t46))
(step t48 (cl (not (= (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op)) (ite y$111 (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) (not (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op))) (ite y$111 (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule equiv_pos2)
(step t49 (cl (= y$111 y$111)) :rule refl)
(step t50 (cl (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op))
(step t51 (cl (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op))
(step t52 (cl (= (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op)) (ite y$111 (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule cong :premises (t49 t50 t51))
(step t53 (cl (not (= (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op)) (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op)))) (not (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op))) (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op))) :rule equiv_pos2)
(step t54 (cl (= (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op)) true) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule equiv_simplify)
(step t55 (cl (not (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op)) true)) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule equiv1 :premises (t54))
(step t56 (cl (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule refl)
(step t57 (cl (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule refl)
(step t58 (cl (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule cong :premises (t56 t57))
(step t59 (cl (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$111 y$s$198_op y$v3_1517448506_31_op)))
(step t60 (cl (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op)) true)) :rule trans :premises (t58 t59))
(step t61 (cl (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule resolution :premises (t55 t60))
(step t62 (cl (= y$s$198_op y$s$198_op)) :rule refl)
(step t63 (cl (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op))) :rule cong :premises (t61 t62))
(step t64 (cl (= y$v3_1517448506_31_op y$v3_1517448506_31_op)) :rule refl)
(step t65 (cl (= (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op))) :rule cong :premises (t61 t64))
(step t66 (cl (= (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op)) (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op)))) :rule cong :premises (t49 t63 t65))
(step t67 (cl (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op))) :rule hole)
(step t68 (cl (ite y$111 (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op) (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op))) :rule resolution :premises (t53 t66 t67))
(step t69 (cl (ite y$111 (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule resolution :premises (t48 t52 t68))
(step t70 (cl y$111 (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule ite1 :premises (t69))
(step t71 (cl (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) :rule and_neg)
(step t72 (cl (=> (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) :rule implies_neg1)
(anchor :step t73)
(assume t73.a0 (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))
(assume t73.a1 (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)))
(assume t73.a2 (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)))
(assume t73.a3 (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)))
(assume t73.a4 (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)))
(assume t73.a5 (= y$v_range$next y$v_range$next_rhs_op))
(assume t73.a6 (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t73.a7 (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t73.a8 (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t73.a9 (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t73.a10 (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))
(assume t73.a11 (= y$v_range (ite y$f09 y$132 y$v_range)))
(assume t73.a12 (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)))
(assume t73.a13 (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)))
(assume t73.a14 (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)))
(assume t73.a15 (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)))
(assume t73.a16 (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))
(assume t73.a17 (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))
(assume t73.a18 (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)))
(assume t73.a19 (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))
(step t73.t0 (cl (= y$n0s16 y$n0s16)) :rule refl)
(step t73.t1 (cl (not (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))))) (not (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule equiv_pos2)
(step t73.t2 (cl (= y$f09 y$f09)) :rule refl)
(step t73.t3 (cl (= (= (ite y$f09 y$132 y$v_range) y$132) (= y$132 (ite y$f09 y$132 y$v_range)))) :rule rare_rewrite :args ("eq-symm" (ite y$f09 y$132 y$v_range) y$132))
(step t73.t4 (cl (= (= (ite y$f09 y$132 y$v_range) y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule rare_rewrite :args ("eq-symm" (ite y$f09 y$132 y$v_range) y$v_range))
(step t73.t5 (cl (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))))) :rule cong :premises (t73.t2 t73.t3 t73.t4))
(step t73.t6 (cl (not (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)))) (not (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule equiv_pos2)
(step t73.t7 (cl (= (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)))) :rule equiv_simplify)
(step t73.t8 (cl (not (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule equiv1 :premises (t73.t7))
(step t73.t9 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule refl)
(step t73.t10 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule refl)
(step t73.t11 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)))) :rule cong :premises (t73.t9 t73.t10))
(step t73.t12 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f09 y$132 y$v_range)))
(step t73.t13 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) :rule trans :premises (t73.t11 t73.t12))
(step t73.t14 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t73.t8 t73.t13))
(step t73.t15 (cl (= y$132 y$132)) :rule refl)
(step t73.t16 (cl (= (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$132))) :rule cong :premises (t73.t14 t73.t15))
(step t73.t17 (cl (= y$v_range y$v_range)) :rule refl)
(step t73.t18 (cl (= (= (ite y$f09 y$132 y$v_range) y$v_range) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule cong :premises (t73.t14 t73.t17))
(step t73.t19 (cl (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)))) :rule cong :premises (t73.t2 t73.t16 t73.t18))
(step t73.t20 (cl (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule hole)
(step t73.t21 (cl (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule resolution :premises (t73.t6 t73.t19 t73.t20))
(step t73.t22 (cl (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule resolution :premises (t73.t1 t73.t5 t73.t21))
(step t73.t23 (cl y$f09 (= y$v_range (ite y$f09 y$132 y$v_range))) :rule ite1 :premises (t73.t22))
(step t73.t24 (cl (not y$f09) (not y$212)) :rule equiv1 :premises (a96))
(step t73.t25 (cl (not y$212) (not y$f09)) :rule reordering :premises (t73.t24))
(step t73.t26 (cl (not (and y$25 y$212)) y$212) :rule and_pos :args (1))
(step t73.t27 (cl y$212 (not (and y$25 y$212))) :rule reordering :premises (t73.t26))
(step t73.t28 (cl (not y$v3_1517448506_117_op) (and y$25 y$212)) :rule equiv1 :premises (a97))
(step t73.t29 (cl (and y$25 y$212) (not y$v3_1517448506_117_op)) :rule reordering :premises (t73.t28))
(step t73.t30 (cl (not (= (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op y$215))) (not (or y$v3_1517448506_117_op (not (not y$215)))) (or y$v3_1517448506_117_op y$215)) :rule equiv_pos2)
(step t73.t31 (cl (= y$v3_1517448506_117_op y$v3_1517448506_117_op)) :rule refl)
(step t73.t32 (cl (= (not (not y$215)) y$215)) :rule rare_rewrite :args ("bool-double-not-elim" y$215))
(step t73.t33 (cl (= (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op y$215))) :rule cong :premises (t73.t31 t73.t32))
(step t73.t34 (cl y$v3_1517448506_117_op (not (not y$215))) :rule equiv2 :premises (a98))
(step t73.t35 (cl (or y$v3_1517448506_117_op (not (not y$215))) (not y$v3_1517448506_117_op)) :rule or_neg :args (0))
(step t73.t36 (cl (or y$v3_1517448506_117_op (not (not y$215))) (not (not (not y$215)))) :rule or_neg :args (1))
(step t73.t37 (cl (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op (not (not y$215)))) :rule resolution :premises (t73.t34 t73.t35 t73.t36))
(step t73.t38 (cl (or y$v3_1517448506_117_op (not (not y$215)))) :rule contraction :premises (t73.t37))
(step t73.t39 (cl (or y$v3_1517448506_117_op y$215)) :rule resolution :premises (t73.t30 t73.t33 t73.t38))
(step t73.t40 (cl y$v3_1517448506_117_op y$215) :rule or :premises (t73.t39))
(step t73.t41 (cl y$215 y$v3_1517448506_117_op) :rule reordering :premises (t73.t40))
(step t73.t42 (cl (not y$215) (not y$216)) :rule equiv1 :premises (a99))
(step t73.t43 (cl (not (and y$216 y$217)) y$216) :rule and_pos :args (0))
(step t73.t44 (cl y$216 (not (and y$216 y$217))) :rule reordering :premises (t73.t43))
(step t73.t45 (cl (not y$v3_1517448506_119_op) (and y$216 y$217)) :rule equiv1 :premises (a101))
(step t73.t46 (cl (and y$216 y$217) (not y$v3_1517448506_119_op)) :rule reordering :premises (t73.t45))
(step t73.t47 (cl (not (= (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op y$220))) (not (or y$v3_1517448506_119_op (not (not y$220)))) (or y$v3_1517448506_119_op y$220)) :rule equiv_pos2)
(step t73.t48 (cl (= y$v3_1517448506_119_op y$v3_1517448506_119_op)) :rule refl)
(step t73.t49 (cl (= (not (not y$220)) y$220)) :rule rare_rewrite :args ("bool-double-not-elim" y$220))
(step t73.t50 (cl (= (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op y$220))) :rule cong :premises (t73.t48 t73.t49))
(step t73.t51 (cl y$v3_1517448506_119_op (not (not y$220))) :rule equiv2 :premises (a102))
(step t73.t52 (cl (or y$v3_1517448506_119_op (not (not y$220))) (not y$v3_1517448506_119_op)) :rule or_neg :args (0))
(step t73.t53 (cl (or y$v3_1517448506_119_op (not (not y$220))) (not (not (not y$220)))) :rule or_neg :args (1))
(step t73.t54 (cl (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op (not (not y$220)))) :rule resolution :premises (t73.t51 t73.t52 t73.t53))
(step t73.t55 (cl (or y$v3_1517448506_119_op (not (not y$220)))) :rule contraction :premises (t73.t54))
(step t73.t56 (cl (or y$v3_1517448506_119_op y$220)) :rule resolution :premises (t73.t47 t73.t50 t73.t55))
(step t73.t57 (cl y$v3_1517448506_119_op y$220) :rule or :premises (t73.t56))
(step t73.t58 (cl y$220 y$v3_1517448506_119_op) :rule reordering :premises (t73.t57))
(step t73.t59 (cl (not y$220) (not y$221)) :rule equiv1 :premises (a103))
(step t73.t60 (cl (not (and y$221 y$222)) y$221) :rule and_pos :args (0))
(step t73.t61 (cl y$221 (not (and y$221 y$222))) :rule reordering :premises (t73.t60))
(step t73.t62 (cl (not y$v3_1517448506_121_op) (and y$221 y$222)) :rule equiv1 :premises (a105))
(step t73.t63 (cl (and y$221 y$222) (not y$v3_1517448506_121_op)) :rule reordering :premises (t73.t62))
(step t73.t64 (cl (not (= (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op y$225))) (not (or y$v3_1517448506_121_op (not (not y$225)))) (or y$v3_1517448506_121_op y$225)) :rule equiv_pos2)
(step t73.t65 (cl (= y$v3_1517448506_121_op y$v3_1517448506_121_op)) :rule refl)
(step t73.t66 (cl (= (not (not y$225)) y$225)) :rule rare_rewrite :args ("bool-double-not-elim" y$225))
(step t73.t67 (cl (= (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op y$225))) :rule cong :premises (t73.t65 t73.t66))
(step t73.t68 (cl y$v3_1517448506_121_op (not (not y$225))) :rule equiv2 :premises (a106))
(step t73.t69 (cl (or y$v3_1517448506_121_op (not (not y$225))) (not y$v3_1517448506_121_op)) :rule or_neg :args (0))
(step t73.t70 (cl (or y$v3_1517448506_121_op (not (not y$225))) (not (not (not y$225)))) :rule or_neg :args (1))
(step t73.t71 (cl (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op (not (not y$225)))) :rule resolution :premises (t73.t68 t73.t69 t73.t70))
(step t73.t72 (cl (or y$v3_1517448506_121_op (not (not y$225)))) :rule contraction :premises (t73.t71))
(step t73.t73 (cl (or y$v3_1517448506_121_op y$225)) :rule resolution :premises (t73.t64 t73.t67 t73.t72))
(step t73.t74 (cl y$v3_1517448506_121_op y$225) :rule or :premises (t73.t73))
(step t73.t75 (cl y$225 y$v3_1517448506_121_op) :rule reordering :premises (t73.t74))
(step t73.t76 (cl (and y$161 y$225) (not y$161) (not y$225)) :rule and_neg)
(step t73.t77 (cl (not y$161) (not y$225) (and y$161 y$225)) :rule reordering :premises (t73.t76))
(step t73.t78 (cl (not (and y$1 y$161)) y$161) :rule and_pos :args (1))
(step t73.t79 (cl y$161 (not (and y$1 y$161))) :rule reordering :premises (t73.t78))
(step t73.t80 (cl (not y$v3_1517448506_168_op) (and y$1 y$161)) :rule equiv1 :premises (a165))
(step t73.t81 (cl (and y$1 y$161) (not y$v3_1517448506_168_op)) :rule reordering :premises (t73.t80))
(step t73.t82 (cl (not (= (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op y$313))) (not (or y$v3_1517448506_168_op (not (not y$313)))) (or y$v3_1517448506_168_op y$313)) :rule equiv_pos2)
(step t73.t83 (cl (= y$v3_1517448506_168_op y$v3_1517448506_168_op)) :rule refl)
(step t73.t84 (cl (= (not (not y$313)) y$313)) :rule rare_rewrite :args ("bool-double-not-elim" y$313))
(step t73.t85 (cl (= (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op y$313))) :rule cong :premises (t73.t83 t73.t84))
(step t73.t86 (cl y$v3_1517448506_168_op (not (not y$313))) :rule equiv2 :premises (a166))
(step t73.t87 (cl (or y$v3_1517448506_168_op (not (not y$313))) (not y$v3_1517448506_168_op)) :rule or_neg :args (0))
(step t73.t88 (cl (or y$v3_1517448506_168_op (not (not y$313))) (not (not (not y$313)))) :rule or_neg :args (1))
(step t73.t89 (cl (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op (not (not y$313)))) :rule resolution :premises (t73.t86 t73.t87 t73.t88))
(step t73.t90 (cl (or y$v3_1517448506_168_op (not (not y$313)))) :rule contraction :premises (t73.t89))
(step t73.t91 (cl (or y$v3_1517448506_168_op y$313)) :rule resolution :premises (t73.t82 t73.t85 t73.t90))
(step t73.t92 (cl y$v3_1517448506_168_op y$313) :rule or :premises (t73.t91))
(step t73.t93 (cl y$313 y$v3_1517448506_168_op) :rule reordering :premises (t73.t92))
(step t73.t94 (cl (and y$313 y$315) (not y$313) (not y$315)) :rule and_neg)
(step t73.t95 (cl (not y$313) (not y$315) (and y$313 y$315)) :rule reordering :premises (t73.t94))
(step t73.t96 (cl (not (and y$3 y$315)) y$315) :rule and_pos :args (1))
(step t73.t97 (cl y$315 (not (and y$3 y$315))) :rule reordering :premises (t73.t96))
(step t73.t98 (cl (not y$v3_1517448506_173_op) (and y$3 y$315)) :rule equiv1 :premises (a171))
(step t73.t99 (cl (and y$3 y$315) (not y$v3_1517448506_173_op)) :rule reordering :premises (t73.t98))
(step t73.t100 (cl (not (= (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op y$322))) (not (or y$v3_1517448506_173_op (not (not y$322)))) (or y$v3_1517448506_173_op y$322)) :rule equiv_pos2)
(step t73.t101 (cl (= y$v3_1517448506_173_op y$v3_1517448506_173_op)) :rule refl)
(step t73.t102 (cl (= (not (not y$322)) y$322)) :rule rare_rewrite :args ("bool-double-not-elim" y$322))
(step t73.t103 (cl (= (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op y$322))) :rule cong :premises (t73.t101 t73.t102))
(step t73.t104 (cl y$v3_1517448506_173_op (not (not y$322))) :rule equiv2 :premises (a172))
(step t73.t105 (cl (or y$v3_1517448506_173_op (not (not y$322))) (not y$v3_1517448506_173_op)) :rule or_neg :args (0))
(step t73.t106 (cl (or y$v3_1517448506_173_op (not (not y$322))) (not (not (not y$322)))) :rule or_neg :args (1))
(step t73.t107 (cl (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op (not (not y$322)))) :rule resolution :premises (t73.t104 t73.t105 t73.t106))
(step t73.t108 (cl (or y$v3_1517448506_173_op (not (not y$322)))) :rule contraction :premises (t73.t107))
(step t73.t109 (cl (or y$v3_1517448506_173_op y$322)) :rule resolution :premises (t73.t100 t73.t103 t73.t108))
(step t73.t110 (cl y$v3_1517448506_173_op y$322) :rule or :premises (t73.t109))
(step t73.t111 (cl y$322 y$v3_1517448506_173_op) :rule reordering :premises (t73.t110))
(step t73.t112 (cl (and y$206 y$322) (not y$206) (not y$322)) :rule and_neg)
(step t73.t113 (cl (not y$206) (not y$322) (and y$206 y$322)) :rule reordering :premises (t73.t112))
(step t73.t114 (cl (not (and y$204 y$206)) y$206) :rule and_pos :args (1))
(step t73.t115 (cl y$206 (not (and y$204 y$206))) :rule reordering :premises (t73.t114))
(step t73.t116 (cl (not y$v3_1517448506_114_op) (and y$204 y$206)) :rule equiv1 :premises (a92))
(step t73.t117 (cl (and y$204 y$206) (not y$v3_1517448506_114_op)) :rule reordering :premises (t73.t116))
(step t73.t118 (cl (not (= (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op y$209))) (not (or y$v3_1517448506_114_op (not (not y$209)))) (or y$v3_1517448506_114_op y$209)) :rule equiv_pos2)
(step t73.t119 (cl (= y$v3_1517448506_114_op y$v3_1517448506_114_op)) :rule refl)
(step t73.t120 (cl (= (not (not y$209)) y$209)) :rule rare_rewrite :args ("bool-double-not-elim" y$209))
(step t73.t121 (cl (= (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op y$209))) :rule cong :premises (t73.t119 t73.t120))
(step t73.t122 (cl y$v3_1517448506_114_op (not (not y$209))) :rule equiv2 :premises (a93))
(step t73.t123 (cl (or y$v3_1517448506_114_op (not (not y$209))) (not y$v3_1517448506_114_op)) :rule or_neg :args (0))
(step t73.t124 (cl (or y$v3_1517448506_114_op (not (not y$209))) (not (not (not y$209)))) :rule or_neg :args (1))
(step t73.t125 (cl (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op (not (not y$209)))) :rule resolution :premises (t73.t122 t73.t123 t73.t124))
(step t73.t126 (cl (or y$v3_1517448506_114_op (not (not y$209)))) :rule contraction :premises (t73.t125))
(step t73.t127 (cl (or y$v3_1517448506_114_op y$209)) :rule resolution :premises (t73.t118 t73.t121 t73.t126))
(step t73.t128 (cl y$v3_1517448506_114_op y$209) :rule or :premises (t73.t127))
(step t73.t129 (cl y$209 y$v3_1517448506_114_op) :rule reordering :premises (t73.t128))
(step t73.t130 (cl (not y$209) (not y$1104)) :rule equiv1 :premises (a708))
(step t73.t131 (cl (not (and y$1103 y$1104)) y$1104) :rule and_pos :args (1))
(step t73.t132 (cl y$1104 (not (and y$1103 y$1104))) :rule reordering :premises (t73.t131))
(step t73.t133 (cl (not y$v3_1517448506_547_op) (and y$1103 y$1104)) :rule equiv1 :premises (a709))
(step t73.t134 (cl (and y$1103 y$1104) (not y$v3_1517448506_547_op)) :rule reordering :premises (t73.t133))
(step t73.t135 (cl (not (= (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op y$1107))) (not (or y$v3_1517448506_547_op (not (not y$1107)))) (or y$v3_1517448506_547_op y$1107)) :rule equiv_pos2)
(step t73.t136 (cl (= y$v3_1517448506_547_op y$v3_1517448506_547_op)) :rule refl)
(step t73.t137 (cl (= (not (not y$1107)) y$1107)) :rule rare_rewrite :args ("bool-double-not-elim" y$1107))
(step t73.t138 (cl (= (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op y$1107))) :rule cong :premises (t73.t136 t73.t137))
(step t73.t139 (cl y$v3_1517448506_547_op (not (not y$1107))) :rule equiv2 :premises (a710))
(step t73.t140 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (not y$v3_1517448506_547_op)) :rule or_neg :args (0))
(step t73.t141 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (not (not (not y$1107)))) :rule or_neg :args (1))
(step t73.t142 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op (not (not y$1107)))) :rule resolution :premises (t73.t139 t73.t140 t73.t141))
(step t73.t143 (cl (or y$v3_1517448506_547_op (not (not y$1107)))) :rule contraction :premises (t73.t142))
(step t73.t144 (cl (or y$v3_1517448506_547_op y$1107)) :rule resolution :premises (t73.t135 t73.t138 t73.t143))
(step t73.t145 (cl y$v3_1517448506_547_op y$1107) :rule or :premises (t73.t144))
(step t73.t146 (cl y$1107 y$v3_1517448506_547_op) :rule reordering :premises (t73.t145))
(step t73.t147 (cl (not y$1107) (not y$1116)) :rule equiv1 :premises (a717))
(step t73.t148 (cl (not (and y$1115 y$1116)) y$1116) :rule and_pos :args (1))
(step t73.t149 (cl y$1116 (not (and y$1115 y$1116))) :rule reordering :premises (t73.t148))
(step t73.t150 (cl (not y$v3_1517448506_552_op) (and y$1115 y$1116)) :rule equiv1 :premises (a718))
(step t73.t151 (cl (and y$1115 y$1116) (not y$v3_1517448506_552_op)) :rule reordering :premises (t73.t150))
(step t73.t152 (cl (not (= (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op y$1119))) (not (or y$v3_1517448506_552_op (not (not y$1119)))) (or y$v3_1517448506_552_op y$1119)) :rule equiv_pos2)
(step t73.t153 (cl (= y$v3_1517448506_552_op y$v3_1517448506_552_op)) :rule refl)
(step t73.t154 (cl (= (not (not y$1119)) y$1119)) :rule rare_rewrite :args ("bool-double-not-elim" y$1119))
(step t73.t155 (cl (= (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op y$1119))) :rule cong :premises (t73.t153 t73.t154))
(step t73.t156 (cl y$v3_1517448506_552_op (not (not y$1119))) :rule equiv2 :premises (a719))
(step t73.t157 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (not y$v3_1517448506_552_op)) :rule or_neg :args (0))
(step t73.t158 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (not (not (not y$1119)))) :rule or_neg :args (1))
(step t73.t159 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op (not (not y$1119)))) :rule resolution :premises (t73.t156 t73.t157 t73.t158))
(step t73.t160 (cl (or y$v3_1517448506_552_op (not (not y$1119)))) :rule contraction :premises (t73.t159))
(step t73.t161 (cl (or y$v3_1517448506_552_op y$1119)) :rule resolution :premises (t73.t152 t73.t155 t73.t160))
(step t73.t162 (cl y$v3_1517448506_552_op y$1119) :rule or :premises (t73.t161))
(step t73.t163 (cl y$1119 y$v3_1517448506_552_op) :rule reordering :premises (t73.t162))
(step t73.t164 (cl (not y$1119) (not y$1128)) :rule equiv1 :premises (a726))
(step t73.t165 (cl (not (and y$1127 y$1128)) y$1128) :rule and_pos :args (1))
(step t73.t166 (cl y$1128 (not (and y$1127 y$1128))) :rule reordering :premises (t73.t165))
(step t73.t167 (cl (not y$v3_1517448506_557_op) (and y$1127 y$1128)) :rule equiv1 :premises (a727))
(step t73.t168 (cl (and y$1127 y$1128) (not y$v3_1517448506_557_op)) :rule reordering :premises (t73.t167))
(step t73.t169 (cl (not (= (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op y$1131))) (not (or y$v3_1517448506_557_op (not (not y$1131)))) (or y$v3_1517448506_557_op y$1131)) :rule equiv_pos2)
(step t73.t170 (cl (= y$v3_1517448506_557_op y$v3_1517448506_557_op)) :rule refl)
(step t73.t171 (cl (= (not (not y$1131)) y$1131)) :rule rare_rewrite :args ("bool-double-not-elim" y$1131))
(step t73.t172 (cl (= (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op y$1131))) :rule cong :premises (t73.t170 t73.t171))
(step t73.t173 (cl y$v3_1517448506_557_op (not (not y$1131))) :rule equiv2 :premises (a728))
(step t73.t174 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (not y$v3_1517448506_557_op)) :rule or_neg :args (0))
(step t73.t175 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (not (not (not y$1131)))) :rule or_neg :args (1))
(step t73.t176 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op (not (not y$1131)))) :rule resolution :premises (t73.t173 t73.t174 t73.t175))
(step t73.t177 (cl (or y$v3_1517448506_557_op (not (not y$1131)))) :rule contraction :premises (t73.t176))
(step t73.t178 (cl (or y$v3_1517448506_557_op y$1131)) :rule resolution :premises (t73.t169 t73.t172 t73.t177))
(step t73.t179 (cl y$v3_1517448506_557_op y$1131) :rule or :premises (t73.t178))
(step t73.t180 (cl y$1131 y$v3_1517448506_557_op) :rule reordering :premises (t73.t179))
(step t73.t181 (cl (not y$1131) (not y$1140)) :rule equiv1 :premises (a735))
(step t73.t182 (cl (not (and y$1139 y$1140)) y$1140) :rule and_pos :args (1))
(step t73.t183 (cl y$1140 (not (and y$1139 y$1140))) :rule reordering :premises (t73.t182))
(step t73.t184 (cl (not y$v3_1517448506_562_op) (and y$1139 y$1140)) :rule equiv1 :premises (a736))
(step t73.t185 (cl (and y$1139 y$1140) (not y$v3_1517448506_562_op)) :rule reordering :premises (t73.t184))
(step t73.t186 (cl (not (= (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op y$1143))) (not (or y$v3_1517448506_562_op (not (not y$1143)))) (or y$v3_1517448506_562_op y$1143)) :rule equiv_pos2)
(step t73.t187 (cl (= y$v3_1517448506_562_op y$v3_1517448506_562_op)) :rule refl)
(step t73.t188 (cl (= (not (not y$1143)) y$1143)) :rule rare_rewrite :args ("bool-double-not-elim" y$1143))
(step t73.t189 (cl (= (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op y$1143))) :rule cong :premises (t73.t187 t73.t188))
(step t73.t190 (cl y$v3_1517448506_562_op (not (not y$1143))) :rule equiv2 :premises (a737))
(step t73.t191 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (not y$v3_1517448506_562_op)) :rule or_neg :args (0))
(step t73.t192 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (not (not (not y$1143)))) :rule or_neg :args (1))
(step t73.t193 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op (not (not y$1143)))) :rule resolution :premises (t73.t190 t73.t191 t73.t192))
(step t73.t194 (cl (or y$v3_1517448506_562_op (not (not y$1143)))) :rule contraction :premises (t73.t193))
(step t73.t195 (cl (or y$v3_1517448506_562_op y$1143)) :rule resolution :premises (t73.t186 t73.t189 t73.t194))
(step t73.t196 (cl y$v3_1517448506_562_op y$1143) :rule or :premises (t73.t195))
(step t73.t197 (cl y$1143 y$v3_1517448506_562_op) :rule reordering :premises (t73.t196))
(step t73.t198 (cl (and y$a_too_small_resistance$next_rhs_op y$1143) (not y$a_too_small_resistance$next_rhs_op) (not y$1143)) :rule and_neg)
(step t73.t199 (cl (not y$1143) (and y$a_too_small_resistance$next_rhs_op y$1143) (not y$a_too_small_resistance$next_rhs_op)) :rule reordering :premises (t73.t198))
(step t73.t200 (cl y$v3_1517448506_563_op (not (and y$a_too_small_resistance$next_rhs_op y$1143))) :rule equiv2 :premises (a738))
(step t73.t201 (cl (not y$v3_1517448506_563_op) (not y$1146)) :rule equiv1 :premises (a739))
(step t73.t202 (cl (not y$1146) (not y$v3_1517448506_563_op)) :rule reordering :premises (t73.t201))
(step t73.t203 (cl (not (and y$1138 y$1146)) y$1146) :rule and_pos :args (1))
(step t73.t204 (cl y$1146 (not (and y$1138 y$1146))) :rule reordering :premises (t73.t203))
(step t73.t205 (cl (not y$v3_1517448506_565_op) (and y$1138 y$1146)) :rule equiv1 :premises (a740))
(step t73.t206 (cl (and y$1138 y$1146) (not y$v3_1517448506_565_op)) :rule reordering :premises (t73.t205))
(step t73.t207 (cl (not (= (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op y$1149))) (not (or y$v3_1517448506_565_op (not (not y$1149)))) (or y$v3_1517448506_565_op y$1149)) :rule equiv_pos2)
(step t73.t208 (cl (= y$v3_1517448506_565_op y$v3_1517448506_565_op)) :rule refl)
(step t73.t209 (cl (= (not (not y$1149)) y$1149)) :rule rare_rewrite :args ("bool-double-not-elim" y$1149))
(step t73.t210 (cl (= (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op y$1149))) :rule cong :premises (t73.t208 t73.t209))
(step t73.t211 (cl y$v3_1517448506_565_op (not (not y$1149))) :rule equiv2 :premises (a741))
(step t73.t212 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (not y$v3_1517448506_565_op)) :rule or_neg :args (0))
(step t73.t213 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (not (not (not y$1149)))) :rule or_neg :args (1))
(step t73.t214 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op (not (not y$1149)))) :rule resolution :premises (t73.t211 t73.t212 t73.t213))
(step t73.t215 (cl (or y$v3_1517448506_565_op (not (not y$1149)))) :rule contraction :premises (t73.t214))
(step t73.t216 (cl (or y$v3_1517448506_565_op y$1149)) :rule resolution :premises (t73.t207 t73.t210 t73.t215))
(step t73.t217 (cl y$v3_1517448506_565_op y$1149) :rule or :premises (t73.t216))
(step t73.t218 (cl y$1149 y$v3_1517448506_565_op) :rule reordering :premises (t73.t217))
(step t73.t219 (cl (not y$1149) (not y$1150)) :rule equiv1 :premises (a742))
(step t73.t220 (cl (not (and y$1150 y$1158)) y$1150) :rule and_pos :args (0))
(step t73.t221 (cl y$1150 (not (and y$1150 y$1158))) :rule reordering :premises (t73.t220))
(step t73.t222 (cl (not y$v3_1517448506_570_op) (and y$1150 y$1158)) :rule equiv1 :premises (a749))
(step t73.t223 (cl (and y$1150 y$1158) (not y$v3_1517448506_570_op)) :rule reordering :premises (t73.t222))
(step t73.t224 (cl (not (= (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op y$1161))) (not (or y$v3_1517448506_570_op (not (not y$1161)))) (or y$v3_1517448506_570_op y$1161)) :rule equiv_pos2)
(step t73.t225 (cl (= y$v3_1517448506_570_op y$v3_1517448506_570_op)) :rule refl)
(step t73.t226 (cl (= (not (not y$1161)) y$1161)) :rule rare_rewrite :args ("bool-double-not-elim" y$1161))
(step t73.t227 (cl (= (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op y$1161))) :rule cong :premises (t73.t225 t73.t226))
(step t73.t228 (cl y$v3_1517448506_570_op (not (not y$1161))) :rule equiv2 :premises (a750))
(step t73.t229 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (not y$v3_1517448506_570_op)) :rule or_neg :args (0))
(step t73.t230 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (not (not (not y$1161)))) :rule or_neg :args (1))
(step t73.t231 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op (not (not y$1161)))) :rule resolution :premises (t73.t228 t73.t229 t73.t230))
(step t73.t232 (cl (or y$v3_1517448506_570_op (not (not y$1161)))) :rule contraction :premises (t73.t231))
(step t73.t233 (cl (or y$v3_1517448506_570_op y$1161)) :rule resolution :premises (t73.t224 t73.t227 t73.t232))
(step t73.t234 (cl y$v3_1517448506_570_op y$1161) :rule or :premises (t73.t233))
(step t73.t235 (cl y$1161 y$v3_1517448506_570_op) :rule reordering :premises (t73.t234))
(step t73.t236 (cl (not y$1161) (not y$1162)) :rule equiv1 :premises (a751))
(step t73.t237 (cl (not (and y$1162 y$1170)) y$1162) :rule and_pos :args (0))
(step t73.t238 (cl y$1162 (not (and y$1162 y$1170))) :rule reordering :premises (t73.t237))
(step t73.t239 (cl (not y$v3_1517448506_575_op) (and y$1162 y$1170)) :rule equiv1 :premises (a758))
(step t73.t240 (cl (and y$1162 y$1170) (not y$v3_1517448506_575_op)) :rule reordering :premises (t73.t239))
(step t73.t241 (cl (not (= (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op y$1173))) (not (or y$v3_1517448506_575_op (not (not y$1173)))) (or y$v3_1517448506_575_op y$1173)) :rule equiv_pos2)
(step t73.t242 (cl (= y$v3_1517448506_575_op y$v3_1517448506_575_op)) :rule refl)
(step t73.t243 (cl (= (not (not y$1173)) y$1173)) :rule rare_rewrite :args ("bool-double-not-elim" y$1173))
(step t73.t244 (cl (= (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op y$1173))) :rule cong :premises (t73.t242 t73.t243))
(step t73.t245 (cl y$v3_1517448506_575_op (not (not y$1173))) :rule equiv2 :premises (a759))
(step t73.t246 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (not y$v3_1517448506_575_op)) :rule or_neg :args (0))
(step t73.t247 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (not (not (not y$1173)))) :rule or_neg :args (1))
(step t73.t248 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op (not (not y$1173)))) :rule resolution :premises (t73.t245 t73.t246 t73.t247))
(step t73.t249 (cl (or y$v3_1517448506_575_op (not (not y$1173)))) :rule contraction :premises (t73.t248))
(step t73.t250 (cl (or y$v3_1517448506_575_op y$1173)) :rule resolution :premises (t73.t241 t73.t244 t73.t249))
(step t73.t251 (cl y$v3_1517448506_575_op y$1173) :rule or :premises (t73.t250))
(step t73.t252 (cl y$1173 y$v3_1517448506_575_op) :rule reordering :premises (t73.t251))
(step t73.t253 (cl (not y$1173) (not y$1174)) :rule equiv1 :premises (a760))
(step t73.t254 (cl (not (and y$1174 y$1181)) y$1174) :rule and_pos :args (0))
(step t73.t255 (cl y$1174 (not (and y$1174 y$1181))) :rule reordering :premises (t73.t254))
(step t73.t256 (cl (not y$v3_1517448506_580_op) (and y$1174 y$1181)) :rule equiv1 :premises (a766))
(step t73.t257 (cl (and y$1174 y$1181) (not y$v3_1517448506_580_op)) :rule reordering :premises (t73.t256))
(step t73.t258 (cl (not (= (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op y$1184))) (not (or y$v3_1517448506_580_op (not (not y$1184)))) (or y$v3_1517448506_580_op y$1184)) :rule equiv_pos2)
(step t73.t259 (cl (= y$v3_1517448506_580_op y$v3_1517448506_580_op)) :rule refl)
(step t73.t260 (cl (= (not (not y$1184)) y$1184)) :rule rare_rewrite :args ("bool-double-not-elim" y$1184))
(step t73.t261 (cl (= (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op y$1184))) :rule cong :premises (t73.t259 t73.t260))
(step t73.t262 (cl y$v3_1517448506_580_op (not (not y$1184))) :rule equiv2 :premises (a767))
(step t73.t263 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (not y$v3_1517448506_580_op)) :rule or_neg :args (0))
(step t73.t264 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (not (not (not y$1184)))) :rule or_neg :args (1))
(step t73.t265 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op (not (not y$1184)))) :rule resolution :premises (t73.t262 t73.t263 t73.t264))
(step t73.t266 (cl (or y$v3_1517448506_580_op (not (not y$1184)))) :rule contraction :premises (t73.t265))
(step t73.t267 (cl (or y$v3_1517448506_580_op y$1184)) :rule resolution :premises (t73.t258 t73.t261 t73.t266))
(step t73.t268 (cl y$v3_1517448506_580_op y$1184) :rule or :premises (t73.t267))
(step t73.t269 (cl y$1184 y$v3_1517448506_580_op) :rule reordering :premises (t73.t268))
(step t73.t270 (cl (not y$1184) (not y$1185)) :rule equiv1 :premises (a768))
(step t73.t271 (cl (not (and y$1185 y$1193)) y$1185) :rule and_pos :args (0))
(step t73.t272 (cl y$1185 (not (and y$1185 y$1193))) :rule reordering :premises (t73.t271))
(step t73.t273 (cl (not y$v3_1517448506_585_op) (and y$1185 y$1193)) :rule equiv1 :premises (a775))
(step t73.t274 (cl (and y$1185 y$1193) (not y$v3_1517448506_585_op)) :rule reordering :premises (t73.t273))
(step t73.t275 (cl (not (= (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op y$1196))) (not (or y$v3_1517448506_585_op (not (not y$1196)))) (or y$v3_1517448506_585_op y$1196)) :rule equiv_pos2)
(step t73.t276 (cl (= y$v3_1517448506_585_op y$v3_1517448506_585_op)) :rule refl)
(step t73.t277 (cl (= (not (not y$1196)) y$1196)) :rule rare_rewrite :args ("bool-double-not-elim" y$1196))
(step t73.t278 (cl (= (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op y$1196))) :rule cong :premises (t73.t276 t73.t277))
(step t73.t279 (cl y$v3_1517448506_585_op (not (not y$1196))) :rule equiv2 :premises (a776))
(step t73.t280 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (not y$v3_1517448506_585_op)) :rule or_neg :args (0))
(step t73.t281 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (not (not (not y$1196)))) :rule or_neg :args (1))
(step t73.t282 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op (not (not y$1196)))) :rule resolution :premises (t73.t279 t73.t280 t73.t281))
(step t73.t283 (cl (or y$v3_1517448506_585_op (not (not y$1196)))) :rule contraction :premises (t73.t282))
(step t73.t284 (cl (or y$v3_1517448506_585_op y$1196)) :rule resolution :premises (t73.t275 t73.t278 t73.t283))
(step t73.t285 (cl y$v3_1517448506_585_op y$1196) :rule or :premises (t73.t284))
(step t73.t286 (cl y$1196 y$v3_1517448506_585_op) :rule reordering :premises (t73.t285))
(step t73.t287 (cl (not y$1196) (not y$1197)) :rule equiv1 :premises (a777))
(step t73.t288 (cl (not (and y$1197 y$1205)) y$1197) :rule and_pos :args (0))
(step t73.t289 (cl y$1197 (not (and y$1197 y$1205))) :rule reordering :premises (t73.t288))
(step t73.t290 (cl (not y$v3_1517448506_590_op) (and y$1197 y$1205)) :rule equiv1 :premises (a784))
(step t73.t291 (cl (and y$1197 y$1205) (not y$v3_1517448506_590_op)) :rule reordering :premises (t73.t290))
(step t73.t292 (cl (not (= (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op y$1208))) (not (or y$v3_1517448506_590_op (not (not y$1208)))) (or y$v3_1517448506_590_op y$1208)) :rule equiv_pos2)
(step t73.t293 (cl (= y$v3_1517448506_590_op y$v3_1517448506_590_op)) :rule refl)
(step t73.t294 (cl (= (not (not y$1208)) y$1208)) :rule rare_rewrite :args ("bool-double-not-elim" y$1208))
(step t73.t295 (cl (= (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op y$1208))) :rule cong :premises (t73.t293 t73.t294))
(step t73.t296 (cl y$v3_1517448506_590_op (not (not y$1208))) :rule equiv2 :premises (a785))
(step t73.t297 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (not y$v3_1517448506_590_op)) :rule or_neg :args (0))
(step t73.t298 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (not (not (not y$1208)))) :rule or_neg :args (1))
(step t73.t299 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op (not (not y$1208)))) :rule resolution :premises (t73.t296 t73.t297 t73.t298))
(step t73.t300 (cl (or y$v3_1517448506_590_op (not (not y$1208)))) :rule contraction :premises (t73.t299))
(step t73.t301 (cl (or y$v3_1517448506_590_op y$1208)) :rule resolution :premises (t73.t292 t73.t295 t73.t300))
(step t73.t302 (cl y$v3_1517448506_590_op y$1208) :rule or :premises (t73.t301))
(step t73.t303 (cl y$1208 y$v3_1517448506_590_op) :rule reordering :premises (t73.t302))
(step t73.t304 (cl (not y$1208) (not y$1209)) :rule equiv1 :premises (a786))
(step t73.t305 (cl (not (and y$v3_1517448506_544_op y$1209)) y$1209) :rule and_pos :args (1))
(step t73.t306 (cl y$1209 (not (and y$v3_1517448506_544_op y$1209))) :rule reordering :premises (t73.t305))
(step t73.t307 (cl (not y$v3_1517448506_591_op) (and y$v3_1517448506_544_op y$1209)) :rule equiv1 :premises (a787))
(step t73.t308 (cl (and y$v3_1517448506_544_op y$1209) (not y$v3_1517448506_591_op)) :rule reordering :premises (t73.t307))
(step t73.t309 (cl (not (and y$v3_1517448506_591_op y$1216)) y$v3_1517448506_591_op) :rule and_pos :args (0))
(step t73.t310 (cl y$v3_1517448506_591_op (not (and y$v3_1517448506_591_op y$1216))) :rule reordering :premises (t73.t309))
(step t73.t311 (cl (not y$v3_1517448506_594_op) (and y$v3_1517448506_591_op y$1216)) :rule equiv1 :premises (a792))
(step t73.t312 (cl (and y$v3_1517448506_591_op y$1216) (not y$v3_1517448506_594_op)) :rule reordering :premises (t73.t311))
(step t73.t313 (cl (not (and y$v3_1517448506_594_op y$1256)) y$v3_1517448506_594_op) :rule and_pos :args (0))
(step t73.t314 (cl y$v3_1517448506_594_op (not (and y$v3_1517448506_594_op y$1256))) :rule reordering :premises (t73.t313))
(step t73.t315 (cl (not y$v3_1517448506_611_op) (and y$v3_1517448506_594_op y$1256)) :rule equiv1 :premises (a821))
(step t73.t316 (cl (and y$v3_1517448506_594_op y$1256) (not y$v3_1517448506_611_op)) :rule reordering :premises (t73.t315))
(step t73.t317 (cl (not (and y$v3_1517448506_611_op y$1263)) y$v3_1517448506_611_op) :rule and_pos :args (0))
(step t73.t318 (cl y$v3_1517448506_611_op (not (and y$v3_1517448506_611_op y$1263))) :rule reordering :premises (t73.t317))
(step t73.t319 (cl (not y$v3_1517448506_614_op) (and y$v3_1517448506_611_op y$1263)) :rule equiv1 :premises (a826))
(step t73.t320 (cl (and y$v3_1517448506_611_op y$1263) (not y$v3_1517448506_614_op)) :rule reordering :premises (t73.t319))
(step t73.t321 (cl (not (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) y$v3_1517448506_614_op) :rule and_pos :args (1))
(step t73.t322 (cl y$v3_1517448506_614_op (not (and y$v3_1517448506_530_op y$v3_1517448506_614_op))) :rule reordering :premises (t73.t321))
(step t73.t323 (cl (not y$v3_1517448506_615_op) (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) :rule equiv1 :premises (a827))
(step t73.t324 (cl (and y$v3_1517448506_530_op y$v3_1517448506_614_op) (not y$v3_1517448506_615_op)) :rule reordering :premises (t73.t323))
(step t73.t325 (cl (not (and y$41 y$v3_1517448506_615_op)) y$v3_1517448506_615_op) :rule and_pos :args (1))
(step t73.t326 (cl y$v3_1517448506_615_op (not (and y$41 y$v3_1517448506_615_op))) :rule reordering :premises (t73.t325))
(step t73.t327 (cl (not y$v3_1517448506_616_op) (and y$41 y$v3_1517448506_615_op)) :rule equiv1 :premises (a828))
(step t73.t328 (cl (and y$41 y$v3_1517448506_615_op) (not y$v3_1517448506_616_op)) :rule reordering :premises (t73.t327))
(step t73.t329 (cl (not (= (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op y$1270))) (not (or y$v3_1517448506_616_op (not (not y$1270)))) (or y$v3_1517448506_616_op y$1270)) :rule equiv_pos2)
(step t73.t330 (cl (= y$v3_1517448506_616_op y$v3_1517448506_616_op)) :rule refl)
(step t73.t331 (cl (= (not (not y$1270)) y$1270)) :rule rare_rewrite :args ("bool-double-not-elim" y$1270))
(step t73.t332 (cl (= (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op y$1270))) :rule cong :premises (t73.t330 t73.t331))
(step t73.t333 (cl y$v3_1517448506_616_op (not (not y$1270))) :rule equiv2 :premises (a829))
(step t73.t334 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (not y$v3_1517448506_616_op)) :rule or_neg :args (0))
(step t73.t335 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (not (not (not y$1270)))) :rule or_neg :args (1))
(step t73.t336 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op (not (not y$1270)))) :rule resolution :premises (t73.t333 t73.t334 t73.t335))
(step t73.t337 (cl (or y$v3_1517448506_616_op (not (not y$1270)))) :rule contraction :premises (t73.t336))
(step t73.t338 (cl (or y$v3_1517448506_616_op y$1270)) :rule resolution :premises (t73.t329 t73.t332 t73.t337))
(step t73.t339 (cl y$v3_1517448506_616_op y$1270) :rule or :premises (t73.t338))
(step t73.t340 (cl y$1270 y$v3_1517448506_616_op) :rule reordering :premises (t73.t339))
(step t73.t341 (cl (not (= y$1270 y$dve_invalid$next)) (not y$1270) y$dve_invalid$next) :rule equiv_pos2)
(step t73.t342 (cl y$dve_invalid$next (not y$1270) (not (= y$1270 y$dve_invalid$next))) :rule reordering :premises (t73.t341))
(step t73.t343 (cl (not y$dve_invalid$next) (not y$1319)) :rule equiv1 :premises (a832))
(step t73.t344 (cl (not y$1319) (not y$dve_invalid$next)) :rule reordering :premises (t73.t343))
(step t73.t345 (cl (not (and y$1319 y$v3_1517448506_38$next_op)) y$1319) :rule and_pos :args (0))
(step t73.t346 (cl y$1319 (not (and y$1319 y$v3_1517448506_38$next_op))) :rule reordering :premises (t73.t345))
(step t73.t347 (cl (not y$id37$next_op) (and y$1319 y$v3_1517448506_38$next_op)) :rule equiv1 :premises (a844))
(step t73.t348 (cl (and y$1319 y$v3_1517448506_38$next_op) (not y$id37$next_op)) :rule reordering :premises (t73.t347))
(step t73.t349 (cl (not (= (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op y$1342))) (not (or y$id37$next_op (not (not y$1342)))) (or y$id37$next_op y$1342)) :rule equiv_pos2)
(step t73.t350 (cl (= y$id37$next_op y$id37$next_op)) :rule refl)
(step t73.t351 (cl (= (not (not y$1342)) y$1342)) :rule rare_rewrite :args ("bool-double-not-elim" y$1342))
(step t73.t352 (cl (= (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op y$1342))) :rule cong :premises (t73.t350 t73.t351))
(step t73.t353 (cl y$id37$next_op (not (not y$1342))) :rule equiv2 :premises (a845))
(step t73.t354 (cl (or y$id37$next_op (not (not y$1342))) (not y$id37$next_op)) :rule or_neg :args (0))
(step t73.t355 (cl (or y$id37$next_op (not (not y$1342))) (not (not (not y$1342)))) :rule or_neg :args (1))
(step t73.t356 (cl (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op (not (not y$1342)))) :rule resolution :premises (t73.t353 t73.t354 t73.t355))
(step t73.t357 (cl (or y$id37$next_op (not (not y$1342)))) :rule contraction :premises (t73.t356))
(step t73.t358 (cl (or y$id37$next_op y$1342)) :rule resolution :premises (t73.t349 t73.t352 t73.t357))
(step t73.t359 (cl y$id37$next_op y$1342) :rule or :premises (t73.t358))
(step t73.t360 (cl y$1342 y$id37$next_op) :rule reordering :premises (t73.t359))
(step t73.t361 (cl (not (= y$1342 y$prop$next)) (not y$1342) y$prop$next) :rule equiv_pos2)
(step t73.t362 (cl y$prop$next (not y$1342) (not (= y$1342 y$prop$next))) :rule reordering :premises (t73.t361))
(step t73.t363 (cl (not y$prop$next) (not y$1318)) :rule equiv1 :premises (a847))
(step t73.t364 (cl (not y$1318) (not y$prop$next)) :rule reordering :premises (t73.t363))
(step t73.t365 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1318) :rule and_pos :args (4))
(step t73.t366 (cl y$1318 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t73.t365))
(step t73.t367 (cl (not y$1346) (and y$prop y$1284 y$1272 y$1343 y$1318)) :rule equiv1 :premises (a848))
(step t73.t368 (cl (and y$prop y$1284 y$1272 y$1343 y$1318) (not y$1346)) :rule reordering :premises (t73.t367))
(step t73.t369 (cl (and y$prop y$1284 y$1272 y$1343 y$1318)) :rule resolution :premises (t73.t368 a849))
(step t73.t370 (cl y$1318) :rule resolution :premises (t73.t366 t73.t369))
(step t73.t371 (cl (not y$prop$next)) :rule resolution :premises (t73.t364 t73.t370))
(step t73.t372 (cl (not (= (= y$1343 (= y$prop$next y$1342)) (= y$1343 (= y$1342 y$prop$next)))) (not (= y$1343 (= y$prop$next y$1342))) (= y$1343 (= y$1342 y$prop$next))) :rule equiv_pos2)
(step t73.t373 (cl (= y$1343 y$1343)) :rule refl)
(step t73.t374 (cl (= (= y$prop$next y$1342) (= y$1342 y$prop$next))) :rule rare_rewrite :args ("eq-symm" y$prop$next y$1342))
(step t73.t375 (cl (= (= y$1343 (= y$prop$next y$1342)) (= y$1343 (= y$1342 y$prop$next)))) :rule cong :premises (t73.t373 t73.t374))
(step t73.t376 (cl (= y$1343 (= y$1342 y$prop$next))) :rule resolution :premises (t73.t372 t73.t375 a846))
(step t73.t377 (cl (not y$1343) (= y$1342 y$prop$next)) :rule equiv1 :premises (t73.t376))
(step t73.t378 (cl (= y$1342 y$prop$next) (not y$1343)) :rule reordering :premises (t73.t377))
(step t73.t379 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1343) :rule and_pos :args (3))
(step t73.t380 (cl y$1343 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t73.t379))
(step t73.t381 (cl y$1343) :rule resolution :premises (t73.t380 t73.t369))
(step t73.t382 (cl (= y$1342 y$prop$next)) :rule resolution :premises (t73.t378 t73.t381))
(step t73.t383 (cl (not y$1342)) :rule resolution :premises (t73.t362 t73.t371 t73.t382))
(step t73.t384 (cl y$id37$next_op) :rule resolution :premises (t73.t360 t73.t383))
(step t73.t385 (cl (and y$1319 y$v3_1517448506_38$next_op)) :rule resolution :premises (t73.t348 t73.t384))
(step t73.t386 (cl y$1319) :rule resolution :premises (t73.t346 t73.t385))
(step t73.t387 (cl (not y$dve_invalid$next)) :rule resolution :premises (t73.t344 t73.t386))
(step t73.t388 (cl (not (= (= y$1271 (= y$dve_invalid$next y$1270)) (= y$1271 (= y$1270 y$dve_invalid$next)))) (not (= y$1271 (= y$dve_invalid$next y$1270))) (= y$1271 (= y$1270 y$dve_invalid$next))) :rule equiv_pos2)
(step t73.t389 (cl (= y$1271 y$1271)) :rule refl)
(step t73.t390 (cl (= (= y$dve_invalid$next y$1270) (= y$1270 y$dve_invalid$next))) :rule rare_rewrite :args ("eq-symm" y$dve_invalid$next y$1270))
(step t73.t391 (cl (= (= y$1271 (= y$dve_invalid$next y$1270)) (= y$1271 (= y$1270 y$dve_invalid$next)))) :rule cong :premises (t73.t389 t73.t390))
(step t73.t392 (cl (= y$1271 (= y$1270 y$dve_invalid$next))) :rule resolution :premises (t73.t388 t73.t391 a830))
(step t73.t393 (cl (not y$1271) (= y$1270 y$dve_invalid$next)) :rule equiv1 :premises (t73.t392))
(step t73.t394 (cl (= y$1270 y$dve_invalid$next) (not y$1271)) :rule reordering :premises (t73.t393))
(step t73.t395 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$1271) :rule and_pos :args (23))
(step t73.t396 (cl y$1271 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t73.t395))
(step t73.t397 (cl (not y$1272) (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) :rule equiv1 :premises (a831))
(step t73.t398 (cl (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271) (not y$1272)) :rule reordering :premises (t73.t397))
(step t73.t399 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1272) :rule and_pos :args (2))
(step t73.t400 (cl y$1272 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t73.t399))
(step t73.t401 (cl y$1272) :rule resolution :premises (t73.t400 t73.t369))
(step t73.t402 (cl (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) :rule resolution :premises (t73.t398 t73.t401))
(step t73.t403 (cl y$1271) :rule resolution :premises (t73.t396 t73.t402))
(step t73.t404 (cl (= y$1270 y$dve_invalid$next)) :rule resolution :premises (t73.t394 t73.t403))
(step t73.t405 (cl (not y$1270)) :rule resolution :premises (t73.t342 t73.t387 t73.t404))
(step t73.t406 (cl y$v3_1517448506_616_op) :rule resolution :premises (t73.t340 t73.t405))
(step t73.t407 (cl (and y$41 y$v3_1517448506_615_op)) :rule resolution :premises (t73.t328 t73.t406))
(step t73.t408 (cl y$v3_1517448506_615_op) :rule resolution :premises (t73.t326 t73.t407))
(step t73.t409 (cl (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) :rule resolution :premises (t73.t324 t73.t408))
(step t73.t410 (cl y$v3_1517448506_614_op) :rule resolution :premises (t73.t322 t73.t409))
(step t73.t411 (cl (and y$v3_1517448506_611_op y$1263)) :rule resolution :premises (t73.t320 t73.t410))
(step t73.t412 (cl y$v3_1517448506_611_op) :rule resolution :premises (t73.t318 t73.t411))
(step t73.t413 (cl (and y$v3_1517448506_594_op y$1256)) :rule resolution :premises (t73.t316 t73.t412))
(step t73.t414 (cl y$v3_1517448506_594_op) :rule resolution :premises (t73.t314 t73.t413))
(step t73.t415 (cl (and y$v3_1517448506_591_op y$1216)) :rule resolution :premises (t73.t312 t73.t414))
(step t73.t416 (cl y$v3_1517448506_591_op) :rule resolution :premises (t73.t310 t73.t415))
(step t73.t417 (cl (and y$v3_1517448506_544_op y$1209)) :rule resolution :premises (t73.t308 t73.t416))
(step t73.t418 (cl y$1209) :rule resolution :premises (t73.t306 t73.t417))
(step t73.t419 (cl (not y$1208)) :rule resolution :premises (t73.t304 t73.t418))
(step t73.t420 (cl y$v3_1517448506_590_op) :rule resolution :premises (t73.t303 t73.t419))
(step t73.t421 (cl (and y$1197 y$1205)) :rule resolution :premises (t73.t291 t73.t420))
(step t73.t422 (cl y$1197) :rule resolution :premises (t73.t289 t73.t421))
(step t73.t423 (cl (not y$1196)) :rule resolution :premises (t73.t287 t73.t422))
(step t73.t424 (cl y$v3_1517448506_585_op) :rule resolution :premises (t73.t286 t73.t423))
(step t73.t425 (cl (and y$1185 y$1193)) :rule resolution :premises (t73.t274 t73.t424))
(step t73.t426 (cl y$1185) :rule resolution :premises (t73.t272 t73.t425))
(step t73.t427 (cl (not y$1184)) :rule resolution :premises (t73.t270 t73.t426))
(step t73.t428 (cl y$v3_1517448506_580_op) :rule resolution :premises (t73.t269 t73.t427))
(step t73.t429 (cl (and y$1174 y$1181)) :rule resolution :premises (t73.t257 t73.t428))
(step t73.t430 (cl y$1174) :rule resolution :premises (t73.t255 t73.t429))
(step t73.t431 (cl (not y$1173)) :rule resolution :premises (t73.t253 t73.t430))
(step t73.t432 (cl y$v3_1517448506_575_op) :rule resolution :premises (t73.t252 t73.t431))
(step t73.t433 (cl (and y$1162 y$1170)) :rule resolution :premises (t73.t240 t73.t432))
(step t73.t434 (cl y$1162) :rule resolution :premises (t73.t238 t73.t433))
(step t73.t435 (cl (not y$1161)) :rule resolution :premises (t73.t236 t73.t434))
(step t73.t436 (cl y$v3_1517448506_570_op) :rule resolution :premises (t73.t235 t73.t435))
(step t73.t437 (cl (and y$1150 y$1158)) :rule resolution :premises (t73.t223 t73.t436))
(step t73.t438 (cl y$1150) :rule resolution :premises (t73.t221 t73.t437))
(step t73.t439 (cl (not y$1149)) :rule resolution :premises (t73.t219 t73.t438))
(step t73.t440 (cl y$v3_1517448506_565_op) :rule resolution :premises (t73.t218 t73.t439))
(step t73.t441 (cl (and y$1138 y$1146)) :rule resolution :premises (t73.t206 t73.t440))
(step t73.t442 (cl y$1146) :rule resolution :premises (t73.t204 t73.t441))
(step t73.t443 (cl (not y$v3_1517448506_563_op)) :rule resolution :premises (t73.t202 t73.t442))
(step t73.t444 (cl (not (and y$a_too_small_resistance$next_rhs_op y$1143))) :rule resolution :premises (t73.t200 t73.t443))
(step t73.t445 (cl (not (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) (not y$a_too_small_resistance$next) y$a_too_small_resistance$next_rhs_op) :rule equiv_pos2)
(step t73.t446 (cl y$a_too_small_resistance$next_rhs_op (not y$a_too_small_resistance$next) (not (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op))) :rule reordering :premises (t73.t445))
(step t73.t447 (cl (not (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) (not (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv_pos2)
(step t73.t448 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next)) :rule refl)
(step t73.t449 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule rare_rewrite :args ("eq-symm" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32))
(step t73.t450 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule rare_rewrite :args ("eq-symm" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))
(step t73.t451 (cl (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) :rule cong :premises (t73.t448 t73.t449 t73.t450))
(step t73.t452 (cl (not (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)))) (not (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule equiv_pos2)
(step t73.t453 (cl (= (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv_simplify)
(step t73.t454 (cl (not (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule equiv1 :premises (t73.t453))
(step t73.t455 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule refl)
(step t73.t456 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule refl)
(step t73.t457 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule cong :premises (t73.t455 t73.t456))
(step t73.t458 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))
(step t73.t459 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) :rule trans :premises (t73.t457 t73.t458))
(step t73.t460 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule resolution :premises (t73.t454 t73.t459))
(step t73.t461 (cl (= y$n1s32 y$n1s32)) :rule refl)
(step t73.t462 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32))) :rule cong :premises (t73.t460 t73.t461))
(step t73.t463 (cl (= y$n0s32 y$n0s32)) :rule refl)
(step t73.t464 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule cong :premises (t73.t460 t73.t463))
(step t73.t465 (cl (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)))) :rule cong :premises (t73.t448 t73.t462 t73.t464))
(step t73.t466 (cl (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule hole)
(step t73.t467 (cl (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule resolution :premises (t73.t452 t73.t465 t73.t466))
(step t73.t468 (cl (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule resolution :premises (t73.t447 t73.t451 t73.t467))
(step t73.t469 (cl y$a_too_small_resistance$next (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule ite1 :premises (t73.t468))
(step t73.t470 (cl (= (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false) (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) :rule equiv_simplify)
(step t73.t471 (cl (not (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false)) (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv1 :premises (t73.t470))
(step t73.t472 (cl (not (= (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) (not (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule equiv_pos2)
(step t73.t473 (cl (= y$v3_1517448506_36$next_op y$v3_1517448506_36$next_op)) :rule refl)
(step t73.t474 (cl (= (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule cong :premises (t73.t473 t73.t460))
(step t73.t475 (cl (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule resolution :premises (t73.t472 t73.t474 a841))
(step t73.t476 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$v3_1517448506_36$next_op)) :rule symm :premises (t73.t475))
(step t73.t477 (cl (not y$1337) (= y$n1s32 y$v3_1517448506_36$next_op)) :rule equiv1 :premises (a842))
(step t73.t478 (cl (= y$n1s32 y$v3_1517448506_36$next_op) (not y$1337)) :rule reordering :premises (t73.t477))
(step t73.t479 (cl (not (and y$1334 y$1337)) y$1337) :rule and_pos :args (1))
(step t73.t480 (cl y$1337 (not (and y$1334 y$1337))) :rule reordering :premises (t73.t479))
(step t73.t481 (cl (not y$v3_1517448506_38$next_op) (and y$1334 y$1337)) :rule equiv1 :premises (a843))
(step t73.t482 (cl (and y$1334 y$1337) (not y$v3_1517448506_38$next_op)) :rule reordering :premises (t73.t481))
(step t73.t483 (cl (not (and y$1319 y$v3_1517448506_38$next_op)) y$v3_1517448506_38$next_op) :rule and_pos :args (1))
(step t73.t484 (cl y$v3_1517448506_38$next_op (not (and y$1319 y$v3_1517448506_38$next_op))) :rule reordering :premises (t73.t483))
(step t73.t485 (cl y$v3_1517448506_38$next_op) :rule resolution :premises (t73.t484 t73.t385))
(step t73.t486 (cl (and y$1334 y$1337)) :rule resolution :premises (t73.t482 t73.t485))
(step t73.t487 (cl y$1337) :rule resolution :premises (t73.t480 t73.t486))
(step t73.t488 (cl (= y$n1s32 y$v3_1517448506_36$next_op)) :rule resolution :premises (t73.t478 t73.t487))
(step t73.t489 (cl (= y$v3_1517448506_36$next_op y$n1s32)) :rule symm :premises (t73.t488))
(step t73.t490 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32)) :rule trans :premises (t73.t476 t73.t489))
(step t73.t491 (cl (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 y$n1s32))) :rule cong :premises (t73.t463 t73.t490))
(step t73.t492 (cl (= (= (= y$n0s32 y$n1s32) false) (not (= y$n0s32 y$n1s32)))) :rule equiv_simplify)
(step t73.t493 (cl (= (= y$n0s32 y$n1s32) false) (not (not (= y$n0s32 y$n1s32)))) :rule equiv2 :premises (t73.t492))
(step t73.t494 (cl (not (not (not (= y$n0s32 y$n1s32)))) (= y$n0s32 y$n1s32)) :rule not_not)
(step t73.t495 (cl (= (= y$n0s32 y$n1s32) false) (= y$n0s32 y$n1s32)) :rule resolution :premises (t73.t493 t73.t494))
(step t73.t496 (cl (not (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) (not (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32)) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule equiv_pos2)
(step t73.t497 (cl (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n0s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n1000s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n0s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n1000s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n6200s32 y$n0s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n5999s32)) (not (= y$n6200s32 y$n1000s32)) (not (= y$n6200s32 y$n5800s32)) (not (= y$n6200s32 y$n5s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n999s32 y$n5999s32)) (not (= y$n999s32 y$n1000s32)) (not (= y$n999s32 y$n5800s32)) (not (= y$n999s32 y$n5s32)) (not (= y$n5999s32 y$n1000s32)) (not (= y$n5999s32 y$n5800s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule distinct_elim)
(step t73.t498 (cl (= (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n1s32)))) :rule refl)
(step t73.t499 (cl (= (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n6200s32)))) :rule refl)
(step t73.t500 (cl (= (= y$n16s32 y$n0s32) (= y$n0s32 y$n16s32))) :rule rare_rewrite :args ("eq-symm" y$n16s32 y$n0s32))
(step t73.t501 (cl (= (not (= y$n16s32 y$n0s32)) (not (= y$n0s32 y$n16s32)))) :rule cong :premises (t73.t500))
(step t73.t502 (cl (= (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n999s32)))) :rule refl)
(step t73.t503 (cl (= (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n5999s32)))) :rule refl)
(step t73.t504 (cl (= (= y$n16s32 y$n1000s32) (= y$n1000s32 y$n16s32))) :rule rare_rewrite :args ("eq-symm" y$n16s32 y$n1000s32))
(step t73.t505 (cl (= (not (= y$n16s32 y$n1000s32)) (not (= y$n1000s32 y$n16s32)))) :rule cong :premises (t73.t504))
(step t73.t506 (cl (= (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5800s32)))) :rule refl)
(step t73.t507 (cl (= (not (= y$n16s32 y$n5s32)) (not (= y$n16s32 y$n5s32)))) :rule refl)
(step t73.t508 (cl (= (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n6200s32)))) :rule refl)
(step t73.t509 (cl (= (= y$n1s32 y$n0s32) (= y$n0s32 y$n1s32))) :rule rare_rewrite :args ("eq-symm" y$n1s32 y$n0s32))
(step t73.t510 (cl (= (not (= y$n1s32 y$n0s32)) (not (= y$n0s32 y$n1s32)))) :rule cong :premises (t73.t509))
(step t73.t511 (cl (= (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n999s32)))) :rule refl)
(step t73.t512 (cl (= (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n5999s32)))) :rule refl)
(step t73.t513 (cl (= (= y$n1s32 y$n1000s32) (= y$n1000s32 y$n1s32))) :rule rare_rewrite :args ("eq-symm" y$n1s32 y$n1000s32))
(step t73.t514 (cl (= (not (= y$n1s32 y$n1000s32)) (not (= y$n1000s32 y$n1s32)))) :rule cong :premises (t73.t513))
(step t73.t515 (cl (= (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5800s32)))) :rule refl)
(step t73.t516 (cl (= (not (= y$n1s32 y$n5s32)) (not (= y$n1s32 y$n5s32)))) :rule refl)
(step t73.t517 (cl (= (= y$n6200s32 y$n0s32) (= y$n0s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n0s32))
(step t73.t518 (cl (= (not (= y$n6200s32 y$n0s32)) (not (= y$n0s32 y$n6200s32)))) :rule cong :premises (t73.t517))
(step t73.t519 (cl (= (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n999s32)))) :rule refl)
(step t73.t520 (cl (= (= y$n6200s32 y$n5999s32) (= y$n5999s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5999s32))
(step t73.t521 (cl (= (not (= y$n6200s32 y$n5999s32)) (not (= y$n5999s32 y$n6200s32)))) :rule cong :premises (t73.t520))
(step t73.t522 (cl (= (= y$n6200s32 y$n1000s32) (= y$n1000s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n1000s32))
(step t73.t523 (cl (= (not (= y$n6200s32 y$n1000s32)) (not (= y$n1000s32 y$n6200s32)))) :rule cong :premises (t73.t522))
(step t73.t524 (cl (= (= y$n6200s32 y$n5800s32) (= y$n5800s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5800s32))
(step t73.t525 (cl (= (not (= y$n6200s32 y$n5800s32)) (not (= y$n5800s32 y$n6200s32)))) :rule cong :premises (t73.t524))
(step t73.t526 (cl (= (= y$n6200s32 y$n5s32) (= y$n5s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5s32))
(step t73.t527 (cl (= (not (= y$n6200s32 y$n5s32)) (not (= y$n5s32 y$n6200s32)))) :rule cong :premises (t73.t526))
(step t73.t528 (cl (= (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n999s32)))) :rule refl)
(step t73.t529 (cl (= (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n5999s32)))) :rule refl)
(step t73.t530 (cl (= (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n1000s32)))) :rule refl)
(step t73.t531 (cl (= (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5800s32)))) :rule refl)
(step t73.t532 (cl (= (not (= y$n0s32 y$n5s32)) (not (= y$n0s32 y$n5s32)))) :rule refl)
(step t73.t533 (cl (= (= y$n999s32 y$n5999s32) (= y$n5999s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5999s32))
(step t73.t534 (cl (= (not (= y$n999s32 y$n5999s32)) (not (= y$n5999s32 y$n999s32)))) :rule cong :premises (t73.t533))
(step t73.t535 (cl (= (= y$n999s32 y$n1000s32) (= y$n1000s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n1000s32))
(step t73.t536 (cl (= (not (= y$n999s32 y$n1000s32)) (not (= y$n1000s32 y$n999s32)))) :rule cong :premises (t73.t535))
(step t73.t537 (cl (= (= y$n999s32 y$n5800s32) (= y$n5800s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5800s32))
(step t73.t538 (cl (= (not (= y$n999s32 y$n5800s32)) (not (= y$n5800s32 y$n999s32)))) :rule cong :premises (t73.t537))
(step t73.t539 (cl (= (= y$n999s32 y$n5s32) (= y$n5s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5s32))
(step t73.t540 (cl (= (not (= y$n999s32 y$n5s32)) (not (= y$n5s32 y$n999s32)))) :rule cong :premises (t73.t539))
(step t73.t541 (cl (= (= y$n5999s32 y$n1000s32) (= y$n1000s32 y$n5999s32))) :rule rare_rewrite :args ("eq-symm" y$n5999s32 y$n1000s32))
(step t73.t542 (cl (= (not (= y$n5999s32 y$n1000s32)) (not (= y$n1000s32 y$n5999s32)))) :rule cong :premises (t73.t541))
(step t73.t543 (cl (= (= y$n5999s32 y$n5800s32) (= y$n5800s32 y$n5999s32))) :rule rare_rewrite :args ("eq-symm" y$n5999s32 y$n5800s32))
(step t73.t544 (cl (= (not (= y$n5999s32 y$n5800s32)) (not (= y$n5800s32 y$n5999s32)))) :rule cong :premises (t73.t543))
(step t73.t545 (cl (= (not (= y$n5999s32 y$n5s32)) (not (= y$n5999s32 y$n5s32)))) :rule refl)
(step t73.t546 (cl (= (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5800s32)))) :rule refl)
(step t73.t547 (cl (= (not (= y$n1000s32 y$n5s32)) (not (= y$n1000s32 y$n5s32)))) :rule refl)
(step t73.t548 (cl (= (not (= y$n5800s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule refl)
(step t73.t549 (cl (= (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n0s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n1000s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n0s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n1000s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n6200s32 y$n0s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n5999s32)) (not (= y$n6200s32 y$n1000s32)) (not (= y$n6200s32 y$n5800s32)) (not (= y$n6200s32 y$n5s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n999s32 y$n5999s32)) (not (= y$n999s32 y$n1000s32)) (not (= y$n999s32 y$n5800s32)) (not (= y$n999s32 y$n5s32)) (not (= y$n5999s32 y$n1000s32)) (not (= y$n5999s32 y$n5800s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule cong :premises (t73.t498 t73.t499 t73.t501 t73.t502 t73.t503 t73.t505 t73.t506 t73.t507 t73.t508 t73.t510 t73.t511 t73.t512 t73.t514 t73.t515 t73.t516 t73.t518 t73.t519 t73.t521 t73.t523 t73.t525 t73.t527 t73.t528 t73.t529 t73.t530 t73.t531 t73.t532 t73.t534 t73.t536 t73.t538 t73.t540 t73.t542 t73.t544 t73.t545 t73.t546 t73.t547 t73.t548))
(step t73.t550 (cl (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule trans :premises (t73.t497 t73.t549))
(step t73.t551 (cl (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule resolution :premises (t73.t496 t73.t550 a1))
(step t73.t552 (cl (not (= y$n0s32 y$n1s32))) :rule and :premises (t73.t551) :args (9))
(step t73.t553 (cl (= (= y$n0s32 y$n1s32) false)) :rule resolution :premises (t73.t495 t73.t552))
(step t73.t554 (cl (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false)) :rule trans :premises (t73.t491 t73.t553))
(step t73.t555 (cl (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule resolution :premises (t73.t471 t73.t554))
(step t73.t556 (cl y$a_too_small_resistance$next) :rule resolution :premises (t73.t469 t73.t555))
(step t73.t557 (cl (not y$274) (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) :rule equiv1 :premises (a138))
(step t73.t558 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op) (not y$274)) :rule reordering :premises (t73.t557))
(step t73.t559 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$274) :rule and_pos :args (12))
(step t73.t560 (cl y$274 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t73.t559))
(step t73.t561 (cl y$274) :rule resolution :premises (t73.t560 t73.t402))
(step t73.t562 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) :rule resolution :premises (t73.t558 t73.t561))
(step t73.t563 (cl y$a_too_small_resistance$next_rhs_op) :rule resolution :premises (t73.t446 t73.t556 t73.t562))
(step t73.t564 (cl (not y$1143)) :rule resolution :premises (t73.t199 t73.t444 t73.t563))
(step t73.t565 (cl y$v3_1517448506_562_op) :rule resolution :premises (t73.t197 t73.t564))
(step t73.t566 (cl (and y$1139 y$1140)) :rule resolution :premises (t73.t185 t73.t565))
(step t73.t567 (cl y$1140) :rule resolution :premises (t73.t183 t73.t566))
(step t73.t568 (cl (not y$1131)) :rule resolution :premises (t73.t181 t73.t567))
(step t73.t569 (cl y$v3_1517448506_557_op) :rule resolution :premises (t73.t180 t73.t568))
(step t73.t570 (cl (and y$1127 y$1128)) :rule resolution :premises (t73.t168 t73.t569))
(step t73.t571 (cl y$1128) :rule resolution :premises (t73.t166 t73.t570))
(step t73.t572 (cl (not y$1119)) :rule resolution :premises (t73.t164 t73.t571))
(step t73.t573 (cl y$v3_1517448506_552_op) :rule resolution :premises (t73.t163 t73.t572))
(step t73.t574 (cl (and y$1115 y$1116)) :rule resolution :premises (t73.t151 t73.t573))
(step t73.t575 (cl y$1116) :rule resolution :premises (t73.t149 t73.t574))
(step t73.t576 (cl (not y$1107)) :rule resolution :premises (t73.t147 t73.t575))
(step t73.t577 (cl y$v3_1517448506_547_op) :rule resolution :premises (t73.t146 t73.t576))
(step t73.t578 (cl (and y$1103 y$1104)) :rule resolution :premises (t73.t134 t73.t577))
(step t73.t579 (cl y$1104) :rule resolution :premises (t73.t132 t73.t578))
(step t73.t580 (cl (not y$209)) :rule resolution :premises (t73.t130 t73.t579))
(step t73.t581 (cl y$v3_1517448506_114_op) :rule resolution :premises (t73.t129 t73.t580))
(step t73.t582 (cl (and y$204 y$206)) :rule resolution :premises (t73.t117 t73.t581))
(step t73.t583 (cl y$206) :rule resolution :premises (t73.t115 t73.t582))
(step t73.t584 (cl y$v3_1517448506_174_op (not (and y$206 y$322))) :rule equiv2 :premises (a173))
(step t73.t585 (cl (and y$261 y$v3_1517448506_174_op) (not y$261) (not y$v3_1517448506_174_op)) :rule and_neg)
(step t73.t586 (cl (not y$261) (and y$261 y$v3_1517448506_174_op) (not y$v3_1517448506_174_op)) :rule reordering :premises (t73.t585))
(step t73.t587 (cl (not (and y$259 y$261)) y$261) :rule and_pos :args (1))
(step t73.t588 (cl y$261 (not (and y$259 y$261))) :rule reordering :premises (t73.t587))
(step t73.t589 (cl (not y$v3_1517448506_142_op) (and y$259 y$261)) :rule equiv1 :premises (a130))
(step t73.t590 (cl (and y$259 y$261) (not y$v3_1517448506_142_op)) :rule reordering :premises (t73.t589))
(step t73.t591 (cl (not (= (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op y$264))) (not (or y$v3_1517448506_142_op (not (not y$264)))) (or y$v3_1517448506_142_op y$264)) :rule equiv_pos2)
(step t73.t592 (cl (= y$v3_1517448506_142_op y$v3_1517448506_142_op)) :rule refl)
(step t73.t593 (cl (= (not (not y$264)) y$264)) :rule rare_rewrite :args ("bool-double-not-elim" y$264))
(step t73.t594 (cl (= (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op y$264))) :rule cong :premises (t73.t592 t73.t593))
(step t73.t595 (cl y$v3_1517448506_142_op (not (not y$264))) :rule equiv2 :premises (a131))
(step t73.t596 (cl (or y$v3_1517448506_142_op (not (not y$264))) (not y$v3_1517448506_142_op)) :rule or_neg :args (0))
(step t73.t597 (cl (or y$v3_1517448506_142_op (not (not y$264))) (not (not (not y$264)))) :rule or_neg :args (1))
(step t73.t598 (cl (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op (not (not y$264)))) :rule resolution :premises (t73.t595 t73.t596 t73.t597))
(step t73.t599 (cl (or y$v3_1517448506_142_op (not (not y$264)))) :rule contraction :premises (t73.t598))
(step t73.t600 (cl (or y$v3_1517448506_142_op y$264)) :rule resolution :premises (t73.t591 t73.t594 t73.t599))
(step t73.t601 (cl y$v3_1517448506_142_op y$264) :rule or :premises (t73.t600))
(step t73.t602 (cl y$264 y$v3_1517448506_142_op) :rule reordering :premises (t73.t601))
(step t73.t603 (cl (not y$264) (not y$1139)) :rule equiv1 :premises (a734))
(step t73.t604 (cl (not (and y$1139 y$1140)) y$1139) :rule and_pos :args (0))
(step t73.t605 (cl y$1139 (not (and y$1139 y$1140))) :rule reordering :premises (t73.t604))
(step t73.t606 (cl y$1139) :rule resolution :premises (t73.t605 t73.t566))
(step t73.t607 (cl (not y$264)) :rule resolution :premises (t73.t603 t73.t606))
(step t73.t608 (cl y$v3_1517448506_142_op) :rule resolution :premises (t73.t602 t73.t607))
(step t73.t609 (cl (and y$259 y$261)) :rule resolution :premises (t73.t590 t73.t608))
(step t73.t610 (cl y$261) :rule resolution :premises (t73.t588 t73.t609))
(step t73.t611 (cl y$a_S2$next_rhs_op (not (and y$261 y$v3_1517448506_174_op))) :rule equiv2 :premises (a174))
(step t73.t612 (cl (and y$a_S2$next_rhs_op y$1202) (not y$a_S2$next_rhs_op) (not y$1202)) :rule and_neg)
(step t73.t613 (cl (not y$1202) (and y$a_S2$next_rhs_op y$1202) (not y$a_S2$next_rhs_op)) :rule reordering :premises (t73.t612))
(step t73.t614 (cl (not (= (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op y$1202))) (not (or y$v3_1517448506_587_op (not (not y$1202)))) (or y$v3_1517448506_587_op y$1202)) :rule equiv_pos2)
(step t73.t615 (cl (= y$v3_1517448506_587_op y$v3_1517448506_587_op)) :rule refl)
(step t73.t616 (cl (= (not (not y$1202)) y$1202)) :rule rare_rewrite :args ("bool-double-not-elim" y$1202))
(step t73.t617 (cl (= (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op y$1202))) :rule cong :premises (t73.t615 t73.t616))
(step t73.t618 (cl y$v3_1517448506_587_op (not (not y$1202))) :rule equiv2 :premises (a781))
(step t73.t619 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (not y$v3_1517448506_587_op)) :rule or_neg :args (0))
(step t73.t620 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (not (not (not y$1202)))) :rule or_neg :args (1))
(step t73.t621 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op (not (not y$1202)))) :rule resolution :premises (t73.t618 t73.t619 t73.t620))
(step t73.t622 (cl (or y$v3_1517448506_587_op (not (not y$1202)))) :rule contraction :premises (t73.t621))
(step t73.t623 (cl (or y$v3_1517448506_587_op y$1202)) :rule resolution :premises (t73.t614 t73.t617 t73.t622))
(step t73.t624 (cl y$v3_1517448506_587_op y$1202) :rule or :premises (t73.t623))
(step t73.t625 (cl y$1202 y$v3_1517448506_587_op) :rule reordering :premises (t73.t624))
(step t73.t626 (cl (not y$v3_1517448506_587_op) (and y$1198 y$1199)) :rule equiv1 :premises (a780))
(step t73.t627 (cl (and y$1198 y$1199) (not y$v3_1517448506_587_op)) :rule reordering :premises (t73.t626))
(step t73.t628 (cl (not (and y$1198 y$1199)) y$1199) :rule and_pos :args (1))
(step t73.t629 (cl y$1199 (not (and y$1198 y$1199))) :rule reordering :premises (t73.t628))
(step t73.t630 (cl (not y$1190) (not y$1199)) :rule equiv1 :premises (a779))
(step t73.t631 (cl (not (= (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op y$1190))) (not (or y$v3_1517448506_582_op (not (not y$1190)))) (or y$v3_1517448506_582_op y$1190)) :rule equiv_pos2)
(step t73.t632 (cl (= y$v3_1517448506_582_op y$v3_1517448506_582_op)) :rule refl)
(step t73.t633 (cl (= (not (not y$1190)) y$1190)) :rule rare_rewrite :args ("bool-double-not-elim" y$1190))
(step t73.t634 (cl (= (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op y$1190))) :rule cong :premises (t73.t632 t73.t633))
(step t73.t635 (cl y$v3_1517448506_582_op (not (not y$1190))) :rule equiv2 :premises (a772))
(step t73.t636 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (not y$v3_1517448506_582_op)) :rule or_neg :args (0))
(step t73.t637 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (not (not (not y$1190)))) :rule or_neg :args (1))
(step t73.t638 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op (not (not y$1190)))) :rule resolution :premises (t73.t635 t73.t636 t73.t637))
(step t73.t639 (cl (or y$v3_1517448506_582_op (not (not y$1190)))) :rule contraction :premises (t73.t638))
(step t73.t640 (cl (or y$v3_1517448506_582_op y$1190)) :rule resolution :premises (t73.t631 t73.t634 t73.t639))
(step t73.t641 (cl y$v3_1517448506_582_op y$1190) :rule or :premises (t73.t640))
(step t73.t642 (cl y$1190 y$v3_1517448506_582_op) :rule reordering :premises (t73.t641))
(step t73.t643 (cl (not y$v3_1517448506_582_op) (and y$1186 y$1187)) :rule equiv1 :premises (a771))
(step t73.t644 (cl (and y$1186 y$1187) (not y$v3_1517448506_582_op)) :rule reordering :premises (t73.t643))
(step t73.t645 (cl (not (and y$1186 y$1187)) y$1187) :rule and_pos :args (1))
(step t73.t646 (cl y$1187 (not (and y$1186 y$1187))) :rule reordering :premises (t73.t645))
(step t73.t647 (cl (not y$1178) (not y$1187)) :rule equiv1 :premises (a770))
(step t73.t648 (cl (not (= (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op y$1178))) (not (or y$v3_1517448506_577_op (not (not y$1178)))) (or y$v3_1517448506_577_op y$1178)) :rule equiv_pos2)
(step t73.t649 (cl (= y$v3_1517448506_577_op y$v3_1517448506_577_op)) :rule refl)
(step t73.t650 (cl (= (not (not y$1178)) y$1178)) :rule rare_rewrite :args ("bool-double-not-elim" y$1178))
(step t73.t651 (cl (= (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op y$1178))) :rule cong :premises (t73.t649 t73.t650))
(step t73.t652 (cl y$v3_1517448506_577_op (not (not y$1178))) :rule equiv2 :premises (a763))
(step t73.t653 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (not y$v3_1517448506_577_op)) :rule or_neg :args (0))
(step t73.t654 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (not (not (not y$1178)))) :rule or_neg :args (1))
(step t73.t655 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op (not (not y$1178)))) :rule resolution :premises (t73.t652 t73.t653 t73.t654))
(step t73.t656 (cl (or y$v3_1517448506_577_op (not (not y$1178)))) :rule contraction :premises (t73.t655))
(step t73.t657 (cl (or y$v3_1517448506_577_op y$1178)) :rule resolution :premises (t73.t648 t73.t651 t73.t656))
(step t73.t658 (cl y$v3_1517448506_577_op y$1178) :rule or :premises (t73.t657))
(step t73.t659 (cl y$1178 y$v3_1517448506_577_op) :rule reordering :premises (t73.t658))
(step t73.t660 (cl (not y$v3_1517448506_577_op) (and y$293 y$1175)) :rule equiv1 :premises (a762))
(step t73.t661 (cl (and y$293 y$1175) (not y$v3_1517448506_577_op)) :rule reordering :premises (t73.t660))
(step t73.t662 (cl (not (and y$293 y$1175)) y$1175) :rule and_pos :args (1))
(step t73.t663 (cl y$1175 (not (and y$293 y$1175))) :rule reordering :premises (t73.t662))
(step t73.t664 (cl (not y$1167) (not y$1175)) :rule equiv1 :premises (a761))
(step t73.t665 (cl (not (= (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op y$1167))) (not (or y$v3_1517448506_572_op (not (not y$1167)))) (or y$v3_1517448506_572_op y$1167)) :rule equiv_pos2)
(step t73.t666 (cl (= y$v3_1517448506_572_op y$v3_1517448506_572_op)) :rule refl)
(step t73.t667 (cl (= (not (not y$1167)) y$1167)) :rule rare_rewrite :args ("bool-double-not-elim" y$1167))
(step t73.t668 (cl (= (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op y$1167))) :rule cong :premises (t73.t666 t73.t667))
(step t73.t669 (cl y$v3_1517448506_572_op (not (not y$1167))) :rule equiv2 :premises (a755))
(step t73.t670 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (not y$v3_1517448506_572_op)) :rule or_neg :args (0))
(step t73.t671 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (not (not (not y$1167)))) :rule or_neg :args (1))
(step t73.t672 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op (not (not y$1167)))) :rule resolution :premises (t73.t669 t73.t670 t73.t671))
(step t73.t673 (cl (or y$v3_1517448506_572_op (not (not y$1167)))) :rule contraction :premises (t73.t672))
(step t73.t674 (cl (or y$v3_1517448506_572_op y$1167)) :rule resolution :premises (t73.t665 t73.t668 t73.t673))
(step t73.t675 (cl y$v3_1517448506_572_op y$1167) :rule or :premises (t73.t674))
(step t73.t676 (cl y$1167 y$v3_1517448506_572_op) :rule reordering :premises (t73.t675))
(step t73.t677 (cl (not y$v3_1517448506_572_op) (and y$1163 y$1164)) :rule equiv1 :premises (a754))
(step t73.t678 (cl (and y$1163 y$1164) (not y$v3_1517448506_572_op)) :rule reordering :premises (t73.t677))
(step t73.t679 (cl (not (and y$1163 y$1164)) y$1164) :rule and_pos :args (1))
(step t73.t680 (cl y$1164 (not (and y$1163 y$1164))) :rule reordering :premises (t73.t679))
(step t73.t681 (cl (not y$1155) (not y$1164)) :rule equiv1 :premises (a753))
(step t73.t682 (cl (not (= (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op y$1155))) (not (or y$v3_1517448506_567_op (not (not y$1155)))) (or y$v3_1517448506_567_op y$1155)) :rule equiv_pos2)
(step t73.t683 (cl (= y$v3_1517448506_567_op y$v3_1517448506_567_op)) :rule refl)
(step t73.t684 (cl (= (not (not y$1155)) y$1155)) :rule rare_rewrite :args ("bool-double-not-elim" y$1155))
(step t73.t685 (cl (= (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op y$1155))) :rule cong :premises (t73.t683 t73.t684))
(step t73.t686 (cl y$v3_1517448506_567_op (not (not y$1155))) :rule equiv2 :premises (a746))
(step t73.t687 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (not y$v3_1517448506_567_op)) :rule or_neg :args (0))
(step t73.t688 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (not (not (not y$1155)))) :rule or_neg :args (1))
(step t73.t689 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op (not (not y$1155)))) :rule resolution :premises (t73.t686 t73.t687 t73.t688))
(step t73.t690 (cl (or y$v3_1517448506_567_op (not (not y$1155)))) :rule contraction :premises (t73.t689))
(step t73.t691 (cl (or y$v3_1517448506_567_op y$1155)) :rule resolution :premises (t73.t682 t73.t685 t73.t690))
(step t73.t692 (cl y$v3_1517448506_567_op y$1155) :rule or :premises (t73.t691))
(step t73.t693 (cl y$1155 y$v3_1517448506_567_op) :rule reordering :premises (t73.t692))
(step t73.t694 (cl (not y$v3_1517448506_567_op) (and y$1151 y$1152)) :rule equiv1 :premises (a745))
(step t73.t695 (cl (and y$1151 y$1152) (not y$v3_1517448506_567_op)) :rule reordering :premises (t73.t694))
(step t73.t696 (cl (not (and y$1151 y$1152)) y$1151) :rule and_pos :args (0))
(step t73.t697 (cl y$1151 (not (and y$1151 y$1152))) :rule reordering :premises (t73.t696))
(step t73.t698 (cl (not y$a_too_small_resistance$next_rhs_op) (not y$1151)) :rule equiv1 :premises (a743))
(step t73.t699 (cl (not y$1151) (not y$a_too_small_resistance$next_rhs_op)) :rule reordering :premises (t73.t698))
(step t73.t700 (cl (not y$1151)) :rule resolution :premises (t73.t699 t73.t563))
(step t73.t701 (cl (not (and y$1151 y$1152))) :rule resolution :premises (t73.t697 t73.t700))
(step t73.t702 (cl (not y$v3_1517448506_567_op)) :rule resolution :premises (t73.t695 t73.t701))
(step t73.t703 (cl y$1155) :rule resolution :premises (t73.t693 t73.t702))
(step t73.t704 (cl (not y$1164)) :rule resolution :premises (t73.t681 t73.t703))
(step t73.t705 (cl (not (and y$1163 y$1164))) :rule resolution :premises (t73.t680 t73.t704))
(step t73.t706 (cl (not y$v3_1517448506_572_op)) :rule resolution :premises (t73.t678 t73.t705))
(step t73.t707 (cl y$1167) :rule resolution :premises (t73.t676 t73.t706))
(step t73.t708 (cl (not y$1175)) :rule resolution :premises (t73.t664 t73.t707))
(step t73.t709 (cl (not (and y$293 y$1175))) :rule resolution :premises (t73.t663 t73.t708))
(step t73.t710 (cl (not y$v3_1517448506_577_op)) :rule resolution :premises (t73.t661 t73.t709))
(step t73.t711 (cl y$1178) :rule resolution :premises (t73.t659 t73.t710))
(step t73.t712 (cl (not y$1187)) :rule resolution :premises (t73.t647 t73.t711))
(step t73.t713 (cl (not (and y$1186 y$1187))) :rule resolution :premises (t73.t646 t73.t712))
(step t73.t714 (cl (not y$v3_1517448506_582_op)) :rule resolution :premises (t73.t644 t73.t713))
(step t73.t715 (cl y$1190) :rule resolution :premises (t73.t642 t73.t714))
(step t73.t716 (cl (not y$1199)) :rule resolution :premises (t73.t630 t73.t715))
(step t73.t717 (cl (not (and y$1198 y$1199))) :rule resolution :premises (t73.t629 t73.t716))
(step t73.t718 (cl (not y$v3_1517448506_587_op)) :rule resolution :premises (t73.t627 t73.t717))
(step t73.t719 (cl y$1202) :rule resolution :premises (t73.t625 t73.t718))
(step t73.t720 (cl y$v3_1517448506_588_op (not (and y$a_S2$next_rhs_op y$1202))) :rule equiv2 :premises (a782))
(step t73.t721 (cl (not y$v3_1517448506_588_op) (not y$1205)) :rule equiv1 :premises (a783))
(step t73.t722 (cl (not y$1205) (not y$v3_1517448506_588_op)) :rule reordering :premises (t73.t721))
(step t73.t723 (cl (not (and y$1197 y$1205)) y$1205) :rule and_pos :args (1))
(step t73.t724 (cl y$1205 (not (and y$1197 y$1205))) :rule reordering :premises (t73.t723))
(step t73.t725 (cl y$1205) :rule resolution :premises (t73.t724 t73.t421))
(step t73.t726 (cl (not y$v3_1517448506_588_op)) :rule resolution :premises (t73.t722 t73.t725))
(step t73.t727 (cl (not (and y$a_S2$next_rhs_op y$1202))) :rule resolution :premises (t73.t720 t73.t726))
(step t73.t728 (cl (not y$a_S2$next_rhs_op)) :rule resolution :premises (t73.t613 t73.t719 t73.t727))
(step t73.t729 (cl (not (and y$261 y$v3_1517448506_174_op))) :rule resolution :premises (t73.t611 t73.t728))
(step t73.t730 (cl (not y$v3_1517448506_174_op)) :rule resolution :premises (t73.t586 t73.t610 t73.t729))
(step t73.t731 (cl (not (and y$206 y$322))) :rule resolution :premises (t73.t584 t73.t730))
(step t73.t732 (cl (not y$322)) :rule resolution :premises (t73.t113 t73.t583 t73.t731))
(step t73.t733 (cl y$v3_1517448506_173_op) :rule resolution :premises (t73.t111 t73.t732))
(step t73.t734 (cl (and y$3 y$315)) :rule resolution :premises (t73.t99 t73.t733))
(step t73.t735 (cl y$315) :rule resolution :premises (t73.t97 t73.t734))
(step t73.t736 (cl y$a_S1$next_rhs_op (not (and y$313 y$315))) :rule equiv2 :premises (a168))
(step t73.t737 (cl (and y$a_S1$next_rhs_op y$1190) (not y$a_S1$next_rhs_op) (not y$1190)) :rule and_neg)
(step t73.t738 (cl (not y$1190) (and y$a_S1$next_rhs_op y$1190) (not y$a_S1$next_rhs_op)) :rule reordering :premises (t73.t737))
(step t73.t739 (cl y$v3_1517448506_583_op (not (and y$a_S1$next_rhs_op y$1190))) :rule equiv2 :premises (a773))
(step t73.t740 (cl (not y$v3_1517448506_583_op) (not y$1193)) :rule equiv1 :premises (a774))
(step t73.t741 (cl (not y$1193) (not y$v3_1517448506_583_op)) :rule reordering :premises (t73.t740))
(step t73.t742 (cl (not (and y$1185 y$1193)) y$1193) :rule and_pos :args (1))
(step t73.t743 (cl y$1193 (not (and y$1185 y$1193))) :rule reordering :premises (t73.t742))
(step t73.t744 (cl y$1193) :rule resolution :premises (t73.t743 t73.t425))
(step t73.t745 (cl (not y$v3_1517448506_583_op)) :rule resolution :premises (t73.t741 t73.t744))
(step t73.t746 (cl (not (and y$a_S1$next_rhs_op y$1190))) :rule resolution :premises (t73.t739 t73.t745))
(step t73.t747 (cl (not y$a_S1$next_rhs_op)) :rule resolution :premises (t73.t738 t73.t715 t73.t746))
(step t73.t748 (cl (not (and y$313 y$315))) :rule resolution :premises (t73.t736 t73.t747))
(step t73.t749 (cl (not y$313)) :rule resolution :premises (t73.t95 t73.t735 t73.t748))
(step t73.t750 (cl y$v3_1517448506_168_op) :rule resolution :premises (t73.t93 t73.t749))
(step t73.t751 (cl (and y$1 y$161)) :rule resolution :premises (t73.t81 t73.t750))
(step t73.t752 (cl y$161) :rule resolution :premises (t73.t79 t73.t751))
(step t73.t753 (cl y$a_new_range$next_rhs_op (not (and y$161 y$225))) :rule equiv2 :premises (a107))
(step t73.t754 (cl (not y$a_new_range$next_rhs_op) (not y$1103)) :rule equiv1 :premises (a707))
(step t73.t755 (cl (not y$1103) (not y$a_new_range$next_rhs_op)) :rule reordering :premises (t73.t754))
(step t73.t756 (cl (not (and y$1103 y$1104)) y$1103) :rule and_pos :args (0))
(step t73.t757 (cl y$1103 (not (and y$1103 y$1104))) :rule reordering :premises (t73.t756))
(step t73.t758 (cl y$1103) :rule resolution :premises (t73.t757 t73.t578))
(step t73.t759 (cl (not y$a_new_range$next_rhs_op)) :rule resolution :premises (t73.t755 t73.t758))
(step t73.t760 (cl (not (and y$161 y$225))) :rule resolution :premises (t73.t753 t73.t759))
(step t73.t761 (cl (not y$225)) :rule resolution :premises (t73.t77 t73.t752 t73.t760))
(step t73.t762 (cl y$v3_1517448506_121_op) :rule resolution :premises (t73.t75 t73.t761))
(step t73.t763 (cl (and y$221 y$222)) :rule resolution :premises (t73.t63 t73.t762))
(step t73.t764 (cl y$221) :rule resolution :premises (t73.t61 t73.t763))
(step t73.t765 (cl (not y$220)) :rule resolution :premises (t73.t59 t73.t764))
(step t73.t766 (cl y$v3_1517448506_119_op) :rule resolution :premises (t73.t58 t73.t765))
(step t73.t767 (cl (and y$216 y$217)) :rule resolution :premises (t73.t46 t73.t766))
(step t73.t768 (cl y$216) :rule resolution :premises (t73.t44 t73.t767))
(step t73.t769 (cl (not y$215)) :rule resolution :premises (t73.t42 t73.t768))
(step t73.t770 (cl y$v3_1517448506_117_op) :rule resolution :premises (t73.t41 t73.t769))
(step t73.t771 (cl (and y$25 y$212)) :rule resolution :premises (t73.t29 t73.t770))
(step t73.t772 (cl y$212) :rule resolution :premises (t73.t27 t73.t771))
(step t73.t773 (cl (not y$f09)) :rule resolution :premises (t73.t25 t73.t772))
(step t73.t774 (cl (= y$v_range (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t73.t23 t73.t773))
(step t73.t775 (cl (not (= (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule equiv_pos2)
(step t73.t776 (cl (= y$v3_1517448506_80_op y$v3_1517448506_80_op)) :rule refl)
(step t73.t777 (cl (= (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))) :rule cong :premises (t73.t776 t73.t14))
(step t73.t778 (cl (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t73.t775 t73.t777 t73.a10))
(step t73.t779 (cl (= (ite y$f09 y$132 y$v_range) y$v3_1517448506_80_op)) :rule symm :premises (t73.t778))
(step t73.t780 (cl (not (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) (not (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule equiv_pos2)
(step t73.t781 (cl (= y$f10 y$f10)) :rule refl)
(step t73.t782 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f10 y$127 y$v3_1517448506_80_op) y$127))
(step t73.t783 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))
(step t73.t784 (cl (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule cong :premises (t73.t781 t73.t782 t73.t783))
(step t73.t785 (cl (not (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)))) (not (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule equiv_pos2)
(step t73.t786 (cl (= (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule equiv_simplify)
(step t73.t787 (cl (not (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule equiv1 :premises (t73.t786))
(step t73.t788 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule refl)
(step t73.t789 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule refl)
(step t73.t790 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule cong :premises (t73.t788 t73.t789))
(step t73.t791 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f10 y$127 y$v3_1517448506_80_op)))
(step t73.t792 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) :rule trans :premises (t73.t790 t73.t791))
(step t73.t793 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t73.t787 t73.t792))
(step t73.t794 (cl (= y$127 y$127)) :rule refl)
(step t73.t795 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127))) :rule cong :premises (t73.t793 t73.t794))
(step t73.t796 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule cong :premises (t73.t793 t73.t776))
(step t73.t797 (cl (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)))) :rule cong :premises (t73.t781 t73.t795 t73.t796))
(step t73.t798 (cl (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule hole)
(step t73.t799 (cl (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule resolution :premises (t73.t785 t73.t797 t73.t798))
(step t73.t800 (cl (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule resolution :premises (t73.t780 t73.t784 t73.t799))
(step t73.t801 (cl y$f10 (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule ite1 :premises (t73.t800))
(step t73.t802 (cl (not y$f10) (not y$217)) :rule equiv1 :premises (a100))
(step t73.t803 (cl (not y$217) (not y$f10)) :rule reordering :premises (t73.t802))
(step t73.t804 (cl (not (and y$216 y$217)) y$217) :rule and_pos :args (1))
(step t73.t805 (cl y$217 (not (and y$216 y$217))) :rule reordering :premises (t73.t804))
(step t73.t806 (cl y$217) :rule resolution :premises (t73.t805 t73.t767))
(step t73.t807 (cl (not y$f10)) :rule resolution :premises (t73.t803 t73.t806))
(step t73.t808 (cl (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t73.t801 t73.t807))
(step t73.t809 (cl (not (= (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule equiv_pos2)
(step t73.t810 (cl (= y$v3_1517448506_81_op y$v3_1517448506_81_op)) :rule refl)
(step t73.t811 (cl (= (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule cong :premises (t73.t810 t73.t793))
(step t73.t812 (cl (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t73.t809 t73.t811 t73.a8))
(step t73.t813 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_81_op)) :rule symm :premises (t73.t812))
(step t73.t814 (cl (not (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) (not (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule equiv_pos2)
(step t73.t815 (cl (= y$f11 y$f11)) :rule refl)
(step t73.t816 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16))
(step t73.t817 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))
(step t73.t818 (cl (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule cong :premises (t73.t815 t73.t816 t73.t817))
(step t73.t819 (cl (not (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)))) (not (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule equiv_pos2)
(step t73.t820 (cl (= (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule equiv_simplify)
(step t73.t821 (cl (not (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule equiv1 :premises (t73.t820))
(step t73.t822 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule refl)
(step t73.t823 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule refl)
(step t73.t824 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule cong :premises (t73.t822 t73.t823))
(step t73.t825 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(step t73.t826 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) :rule trans :premises (t73.t824 t73.t825))
(step t73.t827 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t73.t821 t73.t826))
(step t73.t828 (cl (= y$n3s16 y$n3s16)) :rule refl)
(step t73.t829 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16))) :rule cong :premises (t73.t827 t73.t828))
(step t73.t830 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule cong :premises (t73.t827 t73.t810))
(step t73.t831 (cl (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)))) :rule cong :premises (t73.t815 t73.t829 t73.t830))
(step t73.t832 (cl (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule hole)
(step t73.t833 (cl (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule resolution :premises (t73.t819 t73.t831 t73.t832))
(step t73.t834 (cl (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule resolution :premises (t73.t814 t73.t818 t73.t833))
(step t73.t835 (cl y$f11 (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule ite1 :premises (t73.t834))
(step t73.t836 (cl (not y$f11) (not y$222)) :rule equiv1 :premises (a104))
(step t73.t837 (cl (not y$222) (not y$f11)) :rule reordering :premises (t73.t836))
(step t73.t838 (cl (not (and y$221 y$222)) y$222) :rule and_pos :args (1))
(step t73.t839 (cl y$222 (not (and y$221 y$222))) :rule reordering :premises (t73.t838))
(step t73.t840 (cl y$222) :rule resolution :premises (t73.t839 t73.t763))
(step t73.t841 (cl (not y$f11)) :rule resolution :premises (t73.t837 t73.t840))
(step t73.t842 (cl (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t73.t835 t73.t841))
(step t73.t843 (cl (not (= (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule equiv_pos2)
(step t73.t844 (cl (= y$v_range$next_rhs_op y$v_range$next_rhs_op)) :rule refl)
(step t73.t845 (cl (= (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule cong :premises (t73.t844 t73.t827))
(step t73.t846 (cl (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t73.t843 t73.t845 t73.a6))
(step t73.t847 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v_range$next_rhs_op)) :rule symm :premises (t73.t846))
(step t73.t848 (cl (not y$140) (= y$v_range$next y$v_range$next_rhs_op)) :rule equiv1 :premises (a47))
(step t73.t849 (cl (= y$v_range$next y$v_range$next_rhs_op) (not y$140)) :rule reordering :premises (t73.t848))
(step t73.t850 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$140) :rule and_pos :args (2))
(step t73.t851 (cl y$140 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t73.t850))
(step t73.t852 (cl y$140) :rule resolution :premises (t73.t851 t73.t402))
(step t73.t853 (cl (= y$v_range$next y$v_range$next_rhs_op)) :rule resolution :premises (t73.t849 t73.t852))
(step t73.t854 (cl (= y$v_range$next_rhs_op y$v_range$next)) :rule symm :premises (t73.t853))
(step t73.t855 (cl (= y$v_range y$v_range$next)) :rule trans :premises (t73.t774 t73.t779 t73.t808 t73.t813 t73.t842 t73.t847 t73.t854))
(step t73.t856 (cl (= (Concat_32_16_16 y$n0s16 y$v_range) (Concat_32_16_16 y$n0s16 y$v_range$next))) :rule cong :premises (t73.t0 t73.t855))
(step t73.t857 (cl (= (Concat_32_16_16 y$n0s16 y$v_range$next) y$w$6$next_op)) :rule symm :premises (t73.a4))
(step t73.t858 (cl (= y$w$6_op y$w$6$next_op)) :rule trans :premises (t73.a12 t73.t856 t73.t857))
(step t73.t859 (cl (= y$w$6$next_op y$w$6_op)) :rule symm :premises (t73.t858))
(step t73.t860 (cl (= y$n16s32 y$n16s32)) :rule refl)
(step t73.t861 (cl (= (ShiftR_32_32_32 y$w$6$next_op y$n16s32) (ShiftR_32_32_32 y$w$6_op y$n16s32))) :rule cong :premises (t73.t859 t73.t860))
(step t73.t862 (cl (= (ShiftR_32_32_32 y$w$6_op y$n16s32) y$v3_1517448506_31_op)) :rule symm :premises (t73.a18))
(step t73.t863 (cl (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op)) :rule symm :premises (t73.a17))
(step t73.t864 (cl (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule symm :premises (t73.t863))
(step t73.t865 (cl (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op)) :rule symm :premises (t73.a16))
(step t73.t866 (cl (= (BitWiseNot_32_32 y$w$6_op) (BitWiseNot_32_32 y$w$6$next_op))) :rule cong :premises (t73.t858))
(step t73.t867 (cl (= (BitWiseNot_32_32 y$w$6$next_op) y$s$199$next_op)) :rule symm :premises (t73.a3))
(step t73.t868 (cl (= y$s$199_op y$s$199$next_op)) :rule trans :premises (t73.a13 t73.t866 t73.t867))
(step t73.t869 (cl (= (ShiftR_32_32_32 y$s$199_op y$n16s32) (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) :rule cong :premises (t73.t868 t73.t860))
(step t73.t870 (cl (= (ShiftR_32_32_32 y$s$199$next_op y$n16s32) y$v3_1517448506_32$next_op)) :rule symm :premises (t73.a2))
(step t73.t871 (cl (= y$v3_1517448506_32_op y$v3_1517448506_32$next_op)) :rule trans :premises (t73.a14 t73.t869 t73.t870))
(step t73.t872 (cl (= (BitWiseNot_32_32 y$v3_1517448506_32_op) (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) :rule cong :premises (t73.t871))
(step t73.t873 (cl (= (BitWiseNot_32_32 y$v3_1517448506_32$next_op) y$s$198$next_op)) :rule symm :premises (t73.a1))
(step t73.t874 (cl (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$s$198$next_op)) :rule symm :premises (t73.a0))
(step t73.t875 (cl (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule symm :premises (t73.t874))
(step t73.t876 (cl (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule trans :premises (t73.a19 t73.t861 t73.t862 t73.t864 t73.t865 t73.a15 t73.t872 t73.t873 t73.t875))
(step t73 (cl (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule subproof :discharge (t73.a0 t73.a1 t73.a2 t73.a3 t73.a4 t73.a5 t73.a6 t73.a7 t73.a8 t73.a9 t73.a10 t73.a11 t73.a12 t73.a13 t73.a14 t73.a15 t73.a16 t73.a17 t73.a18 t73.a19))
(step t74 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule and_pos :args (0))
(step t75 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) :rule and_pos :args (1))
(step t76 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) :rule and_pos :args (2))
(step t77 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) :rule and_pos :args (3))
(step t78 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) :rule and_pos :args (4))
(step t79 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v_range$next y$v_range$next_rhs_op)) :rule and_pos :args (5))
(step t80 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (6))
(step t81 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (7))
(step t82 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (8))
(step t83 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (9))
(step t84 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule and_pos :args (10))
(step t85 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v_range (ite y$f09 y$132 y$v_range))) :rule and_pos :args (11))
(step t86 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) :rule and_pos :args (12))
(step t87 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) :rule and_pos :args (13))
(step t88 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) :rule and_pos :args (14))
(step t89 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) :rule and_pos :args (15))
(step t90 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule and_pos :args (16))
(step t91 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule and_pos :args (17))
(step t92 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) :rule and_pos :args (18))
(step t93 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) :rule and_pos :args (19))
(step t94 (cl (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))))) :rule resolution :premises (t73 t74 t75 t76 t77 t78 t79 t80 t81 t82 t83 t84 t85 t86 t87 t88 t89 t90 t91 t92 t93))
(step t95 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule reordering :premises (t94))
(step t96 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule contraction :premises (t95))
(step t97 (cl (=> (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule resolution :premises (t72 t96))
(step t98 (cl (=> (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule implies_neg2)
(step t99 (cl (=> (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (=> (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule resolution :premises (t97 t98))
(step t100 (cl (=> (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule contraction :premises (t99))
(step t101 (cl (not (and (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule implies :premises (t100))
(step t102 (cl (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule resolution :premises (t71 t101))
(step t103 (cl (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule reordering :premises (t102))
(step t104 (cl (not (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))))) (not (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule equiv_pos2)
(step t105 (cl (= y$f09 y$f09)) :rule refl)
(step t106 (cl (= (= (ite y$f09 y$132 y$v_range) y$132) (= y$132 (ite y$f09 y$132 y$v_range)))) :rule rare_rewrite :args ("eq-symm" (ite y$f09 y$132 y$v_range) y$132))
(step t107 (cl (= (= (ite y$f09 y$132 y$v_range) y$v_range) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule rare_rewrite :args ("eq-symm" (ite y$f09 y$132 y$v_range) y$v_range))
(step t108 (cl (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))))) :rule cong :premises (t105 t106 t107))
(step t109 (cl (not (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)))) (not (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule equiv_pos2)
(step t110 (cl (= (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)))) :rule equiv_simplify)
(step t111 (cl (not (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule equiv1 :premises (t110))
(step t112 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule refl)
(step t113 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule refl)
(step t114 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)))) :rule cong :premises (t112 t113))
(step t115 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f09 y$132 y$v_range)))
(step t116 (cl (= (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range)) true)) :rule trans :premises (t114 t115))
(step t117 (cl (= (ite y$f09 y$132 y$v_range) (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t111 t116))
(step t118 (cl (= y$132 y$132)) :rule refl)
(step t119 (cl (= (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$132))) :rule cong :premises (t117 t118))
(step t120 (cl (= y$v_range y$v_range)) :rule refl)
(step t121 (cl (= (= (ite y$f09 y$132 y$v_range) y$v_range) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule cong :premises (t117 t120))
(step t122 (cl (= (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)) (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range)))) :rule cong :premises (t105 t119 t121))
(step t123 (cl (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule hole)
(step t124 (cl (ite y$f09 (= (ite y$f09 y$132 y$v_range) y$132) (= (ite y$f09 y$132 y$v_range) y$v_range))) :rule resolution :premises (t109 t122 t123))
(step t125 (cl (ite y$f09 (= y$132 (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule resolution :premises (t104 t108 t124))
(step t126 (cl y$f09 (= y$v_range (ite y$f09 y$132 y$v_range))) :rule ite1 :premises (t125))
(step t127 (cl (not y$f09) (not y$212)) :rule equiv1 :premises (a96))
(step t128 (cl (not y$212) (not y$f09)) :rule reordering :premises (t127))
(step t129 (cl (not (and y$25 y$212)) y$212) :rule and_pos :args (1))
(step t130 (cl y$212 (not (and y$25 y$212))) :rule reordering :premises (t129))
(step t131 (cl (not y$v3_1517448506_117_op) (and y$25 y$212)) :rule equiv1 :premises (a97))
(step t132 (cl (and y$25 y$212) (not y$v3_1517448506_117_op)) :rule reordering :premises (t131))
(step t133 (cl (not (= (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op y$215))) (not (or y$v3_1517448506_117_op (not (not y$215)))) (or y$v3_1517448506_117_op y$215)) :rule equiv_pos2)
(step t134 (cl (= y$v3_1517448506_117_op y$v3_1517448506_117_op)) :rule refl)
(step t135 (cl (= (not (not y$215)) y$215)) :rule rare_rewrite :args ("bool-double-not-elim" y$215))
(step t136 (cl (= (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op y$215))) :rule cong :premises (t134 t135))
(step t137 (cl y$v3_1517448506_117_op (not (not y$215))) :rule equiv2 :premises (a98))
(step t138 (cl (or y$v3_1517448506_117_op (not (not y$215))) (not y$v3_1517448506_117_op)) :rule or_neg :args (0))
(step t139 (cl (or y$v3_1517448506_117_op (not (not y$215))) (not (not (not y$215)))) :rule or_neg :args (1))
(step t140 (cl (or y$v3_1517448506_117_op (not (not y$215))) (or y$v3_1517448506_117_op (not (not y$215)))) :rule resolution :premises (t137 t138 t139))
(step t141 (cl (or y$v3_1517448506_117_op (not (not y$215)))) :rule contraction :premises (t140))
(step t142 (cl (or y$v3_1517448506_117_op y$215)) :rule resolution :premises (t133 t136 t141))
(step t143 (cl y$v3_1517448506_117_op y$215) :rule or :premises (t142))
(step t144 (cl y$215 y$v3_1517448506_117_op) :rule reordering :premises (t143))
(step t145 (cl (not y$215) (not y$216)) :rule equiv1 :premises (a99))
(step t146 (cl (not (and y$216 y$217)) y$216) :rule and_pos :args (0))
(step t147 (cl y$216 (not (and y$216 y$217))) :rule reordering :premises (t146))
(step t148 (cl (not y$v3_1517448506_119_op) (and y$216 y$217)) :rule equiv1 :premises (a101))
(step t149 (cl (and y$216 y$217) (not y$v3_1517448506_119_op)) :rule reordering :premises (t148))
(step t150 (cl (not (= (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op y$220))) (not (or y$v3_1517448506_119_op (not (not y$220)))) (or y$v3_1517448506_119_op y$220)) :rule equiv_pos2)
(step t151 (cl (= y$v3_1517448506_119_op y$v3_1517448506_119_op)) :rule refl)
(step t152 (cl (= (not (not y$220)) y$220)) :rule rare_rewrite :args ("bool-double-not-elim" y$220))
(step t153 (cl (= (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op y$220))) :rule cong :premises (t151 t152))
(step t154 (cl y$v3_1517448506_119_op (not (not y$220))) :rule equiv2 :premises (a102))
(step t155 (cl (or y$v3_1517448506_119_op (not (not y$220))) (not y$v3_1517448506_119_op)) :rule or_neg :args (0))
(step t156 (cl (or y$v3_1517448506_119_op (not (not y$220))) (not (not (not y$220)))) :rule or_neg :args (1))
(step t157 (cl (or y$v3_1517448506_119_op (not (not y$220))) (or y$v3_1517448506_119_op (not (not y$220)))) :rule resolution :premises (t154 t155 t156))
(step t158 (cl (or y$v3_1517448506_119_op (not (not y$220)))) :rule contraction :premises (t157))
(step t159 (cl (or y$v3_1517448506_119_op y$220)) :rule resolution :premises (t150 t153 t158))
(step t160 (cl y$v3_1517448506_119_op y$220) :rule or :premises (t159))
(step t161 (cl y$220 y$v3_1517448506_119_op) :rule reordering :premises (t160))
(step t162 (cl (not y$220) (not y$221)) :rule equiv1 :premises (a103))
(step t163 (cl (not (and y$221 y$222)) y$221) :rule and_pos :args (0))
(step t164 (cl y$221 (not (and y$221 y$222))) :rule reordering :premises (t163))
(step t165 (cl (not y$v3_1517448506_121_op) (and y$221 y$222)) :rule equiv1 :premises (a105))
(step t166 (cl (and y$221 y$222) (not y$v3_1517448506_121_op)) :rule reordering :premises (t165))
(step t167 (cl (not (= (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op y$225))) (not (or y$v3_1517448506_121_op (not (not y$225)))) (or y$v3_1517448506_121_op y$225)) :rule equiv_pos2)
(step t168 (cl (= y$v3_1517448506_121_op y$v3_1517448506_121_op)) :rule refl)
(step t169 (cl (= (not (not y$225)) y$225)) :rule rare_rewrite :args ("bool-double-not-elim" y$225))
(step t170 (cl (= (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op y$225))) :rule cong :premises (t168 t169))
(step t171 (cl y$v3_1517448506_121_op (not (not y$225))) :rule equiv2 :premises (a106))
(step t172 (cl (or y$v3_1517448506_121_op (not (not y$225))) (not y$v3_1517448506_121_op)) :rule or_neg :args (0))
(step t173 (cl (or y$v3_1517448506_121_op (not (not y$225))) (not (not (not y$225)))) :rule or_neg :args (1))
(step t174 (cl (or y$v3_1517448506_121_op (not (not y$225))) (or y$v3_1517448506_121_op (not (not y$225)))) :rule resolution :premises (t171 t172 t173))
(step t175 (cl (or y$v3_1517448506_121_op (not (not y$225)))) :rule contraction :premises (t174))
(step t176 (cl (or y$v3_1517448506_121_op y$225)) :rule resolution :premises (t167 t170 t175))
(step t177 (cl y$v3_1517448506_121_op y$225) :rule or :premises (t176))
(step t178 (cl y$225 y$v3_1517448506_121_op) :rule reordering :premises (t177))
(step t179 (cl (and y$161 y$225) (not y$161) (not y$225)) :rule and_neg)
(step t180 (cl (not y$161) (not y$225) (and y$161 y$225)) :rule reordering :premises (t179))
(step t181 (cl (not (and y$1 y$161)) y$161) :rule and_pos :args (1))
(step t182 (cl y$161 (not (and y$1 y$161))) :rule reordering :premises (t181))
(step t183 (cl (not y$v3_1517448506_168_op) (and y$1 y$161)) :rule equiv1 :premises (a165))
(step t184 (cl (and y$1 y$161) (not y$v3_1517448506_168_op)) :rule reordering :premises (t183))
(step t185 (cl (not (= (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op y$313))) (not (or y$v3_1517448506_168_op (not (not y$313)))) (or y$v3_1517448506_168_op y$313)) :rule equiv_pos2)
(step t186 (cl (= y$v3_1517448506_168_op y$v3_1517448506_168_op)) :rule refl)
(step t187 (cl (= (not (not y$313)) y$313)) :rule rare_rewrite :args ("bool-double-not-elim" y$313))
(step t188 (cl (= (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op y$313))) :rule cong :premises (t186 t187))
(step t189 (cl y$v3_1517448506_168_op (not (not y$313))) :rule equiv2 :premises (a166))
(step t190 (cl (or y$v3_1517448506_168_op (not (not y$313))) (not y$v3_1517448506_168_op)) :rule or_neg :args (0))
(step t191 (cl (or y$v3_1517448506_168_op (not (not y$313))) (not (not (not y$313)))) :rule or_neg :args (1))
(step t192 (cl (or y$v3_1517448506_168_op (not (not y$313))) (or y$v3_1517448506_168_op (not (not y$313)))) :rule resolution :premises (t189 t190 t191))
(step t193 (cl (or y$v3_1517448506_168_op (not (not y$313)))) :rule contraction :premises (t192))
(step t194 (cl (or y$v3_1517448506_168_op y$313)) :rule resolution :premises (t185 t188 t193))
(step t195 (cl y$v3_1517448506_168_op y$313) :rule or :premises (t194))
(step t196 (cl y$313 y$v3_1517448506_168_op) :rule reordering :premises (t195))
(step t197 (cl (and y$313 y$315) (not y$313) (not y$315)) :rule and_neg)
(step t198 (cl (not y$313) (not y$315) (and y$313 y$315)) :rule reordering :premises (t197))
(step t199 (cl (not (and y$3 y$315)) y$315) :rule and_pos :args (1))
(step t200 (cl y$315 (not (and y$3 y$315))) :rule reordering :premises (t199))
(step t201 (cl (not y$v3_1517448506_173_op) (and y$3 y$315)) :rule equiv1 :premises (a171))
(step t202 (cl (and y$3 y$315) (not y$v3_1517448506_173_op)) :rule reordering :premises (t201))
(step t203 (cl (not (= (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op y$322))) (not (or y$v3_1517448506_173_op (not (not y$322)))) (or y$v3_1517448506_173_op y$322)) :rule equiv_pos2)
(step t204 (cl (= y$v3_1517448506_173_op y$v3_1517448506_173_op)) :rule refl)
(step t205 (cl (= (not (not y$322)) y$322)) :rule rare_rewrite :args ("bool-double-not-elim" y$322))
(step t206 (cl (= (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op y$322))) :rule cong :premises (t204 t205))
(step t207 (cl y$v3_1517448506_173_op (not (not y$322))) :rule equiv2 :premises (a172))
(step t208 (cl (or y$v3_1517448506_173_op (not (not y$322))) (not y$v3_1517448506_173_op)) :rule or_neg :args (0))
(step t209 (cl (or y$v3_1517448506_173_op (not (not y$322))) (not (not (not y$322)))) :rule or_neg :args (1))
(step t210 (cl (or y$v3_1517448506_173_op (not (not y$322))) (or y$v3_1517448506_173_op (not (not y$322)))) :rule resolution :premises (t207 t208 t209))
(step t211 (cl (or y$v3_1517448506_173_op (not (not y$322)))) :rule contraction :premises (t210))
(step t212 (cl (or y$v3_1517448506_173_op y$322)) :rule resolution :premises (t203 t206 t211))
(step t213 (cl y$v3_1517448506_173_op y$322) :rule or :premises (t212))
(step t214 (cl y$322 y$v3_1517448506_173_op) :rule reordering :premises (t213))
(step t215 (cl (and y$206 y$322) (not y$206) (not y$322)) :rule and_neg)
(step t216 (cl (not y$206) (not y$322) (and y$206 y$322)) :rule reordering :premises (t215))
(step t217 (cl (not (and y$204 y$206)) y$206) :rule and_pos :args (1))
(step t218 (cl y$206 (not (and y$204 y$206))) :rule reordering :premises (t217))
(step t219 (cl (not y$v3_1517448506_114_op) (and y$204 y$206)) :rule equiv1 :premises (a92))
(step t220 (cl (and y$204 y$206) (not y$v3_1517448506_114_op)) :rule reordering :premises (t219))
(step t221 (cl (not (= (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op y$209))) (not (or y$v3_1517448506_114_op (not (not y$209)))) (or y$v3_1517448506_114_op y$209)) :rule equiv_pos2)
(step t222 (cl (= y$v3_1517448506_114_op y$v3_1517448506_114_op)) :rule refl)
(step t223 (cl (= (not (not y$209)) y$209)) :rule rare_rewrite :args ("bool-double-not-elim" y$209))
(step t224 (cl (= (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op y$209))) :rule cong :premises (t222 t223))
(step t225 (cl y$v3_1517448506_114_op (not (not y$209))) :rule equiv2 :premises (a93))
(step t226 (cl (or y$v3_1517448506_114_op (not (not y$209))) (not y$v3_1517448506_114_op)) :rule or_neg :args (0))
(step t227 (cl (or y$v3_1517448506_114_op (not (not y$209))) (not (not (not y$209)))) :rule or_neg :args (1))
(step t228 (cl (or y$v3_1517448506_114_op (not (not y$209))) (or y$v3_1517448506_114_op (not (not y$209)))) :rule resolution :premises (t225 t226 t227))
(step t229 (cl (or y$v3_1517448506_114_op (not (not y$209)))) :rule contraction :premises (t228))
(step t230 (cl (or y$v3_1517448506_114_op y$209)) :rule resolution :premises (t221 t224 t229))
(step t231 (cl y$v3_1517448506_114_op y$209) :rule or :premises (t230))
(step t232 (cl y$209 y$v3_1517448506_114_op) :rule reordering :premises (t231))
(step t233 (cl (not y$209) (not y$1104)) :rule equiv1 :premises (a708))
(step t234 (cl (not (and y$1103 y$1104)) y$1104) :rule and_pos :args (1))
(step t235 (cl y$1104 (not (and y$1103 y$1104))) :rule reordering :premises (t234))
(step t236 (cl (not y$v3_1517448506_547_op) (and y$1103 y$1104)) :rule equiv1 :premises (a709))
(step t237 (cl (and y$1103 y$1104) (not y$v3_1517448506_547_op)) :rule reordering :premises (t236))
(step t238 (cl (not (= (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op y$1107))) (not (or y$v3_1517448506_547_op (not (not y$1107)))) (or y$v3_1517448506_547_op y$1107)) :rule equiv_pos2)
(step t239 (cl (= y$v3_1517448506_547_op y$v3_1517448506_547_op)) :rule refl)
(step t240 (cl (= (not (not y$1107)) y$1107)) :rule rare_rewrite :args ("bool-double-not-elim" y$1107))
(step t241 (cl (= (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op y$1107))) :rule cong :premises (t239 t240))
(step t242 (cl y$v3_1517448506_547_op (not (not y$1107))) :rule equiv2 :premises (a710))
(step t243 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (not y$v3_1517448506_547_op)) :rule or_neg :args (0))
(step t244 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (not (not (not y$1107)))) :rule or_neg :args (1))
(step t245 (cl (or y$v3_1517448506_547_op (not (not y$1107))) (or y$v3_1517448506_547_op (not (not y$1107)))) :rule resolution :premises (t242 t243 t244))
(step t246 (cl (or y$v3_1517448506_547_op (not (not y$1107)))) :rule contraction :premises (t245))
(step t247 (cl (or y$v3_1517448506_547_op y$1107)) :rule resolution :premises (t238 t241 t246))
(step t248 (cl y$v3_1517448506_547_op y$1107) :rule or :premises (t247))
(step t249 (cl y$1107 y$v3_1517448506_547_op) :rule reordering :premises (t248))
(step t250 (cl (not y$1107) (not y$1116)) :rule equiv1 :premises (a717))
(step t251 (cl (not (and y$1115 y$1116)) y$1116) :rule and_pos :args (1))
(step t252 (cl y$1116 (not (and y$1115 y$1116))) :rule reordering :premises (t251))
(step t253 (cl (not y$v3_1517448506_552_op) (and y$1115 y$1116)) :rule equiv1 :premises (a718))
(step t254 (cl (and y$1115 y$1116) (not y$v3_1517448506_552_op)) :rule reordering :premises (t253))
(step t255 (cl (not (= (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op y$1119))) (not (or y$v3_1517448506_552_op (not (not y$1119)))) (or y$v3_1517448506_552_op y$1119)) :rule equiv_pos2)
(step t256 (cl (= y$v3_1517448506_552_op y$v3_1517448506_552_op)) :rule refl)
(step t257 (cl (= (not (not y$1119)) y$1119)) :rule rare_rewrite :args ("bool-double-not-elim" y$1119))
(step t258 (cl (= (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op y$1119))) :rule cong :premises (t256 t257))
(step t259 (cl y$v3_1517448506_552_op (not (not y$1119))) :rule equiv2 :premises (a719))
(step t260 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (not y$v3_1517448506_552_op)) :rule or_neg :args (0))
(step t261 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (not (not (not y$1119)))) :rule or_neg :args (1))
(step t262 (cl (or y$v3_1517448506_552_op (not (not y$1119))) (or y$v3_1517448506_552_op (not (not y$1119)))) :rule resolution :premises (t259 t260 t261))
(step t263 (cl (or y$v3_1517448506_552_op (not (not y$1119)))) :rule contraction :premises (t262))
(step t264 (cl (or y$v3_1517448506_552_op y$1119)) :rule resolution :premises (t255 t258 t263))
(step t265 (cl y$v3_1517448506_552_op y$1119) :rule or :premises (t264))
(step t266 (cl y$1119 y$v3_1517448506_552_op) :rule reordering :premises (t265))
(step t267 (cl (not y$1119) (not y$1128)) :rule equiv1 :premises (a726))
(step t268 (cl (not (and y$1127 y$1128)) y$1128) :rule and_pos :args (1))
(step t269 (cl y$1128 (not (and y$1127 y$1128))) :rule reordering :premises (t268))
(step t270 (cl (not y$v3_1517448506_557_op) (and y$1127 y$1128)) :rule equiv1 :premises (a727))
(step t271 (cl (and y$1127 y$1128) (not y$v3_1517448506_557_op)) :rule reordering :premises (t270))
(step t272 (cl (not (= (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op y$1131))) (not (or y$v3_1517448506_557_op (not (not y$1131)))) (or y$v3_1517448506_557_op y$1131)) :rule equiv_pos2)
(step t273 (cl (= y$v3_1517448506_557_op y$v3_1517448506_557_op)) :rule refl)
(step t274 (cl (= (not (not y$1131)) y$1131)) :rule rare_rewrite :args ("bool-double-not-elim" y$1131))
(step t275 (cl (= (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op y$1131))) :rule cong :premises (t273 t274))
(step t276 (cl y$v3_1517448506_557_op (not (not y$1131))) :rule equiv2 :premises (a728))
(step t277 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (not y$v3_1517448506_557_op)) :rule or_neg :args (0))
(step t278 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (not (not (not y$1131)))) :rule or_neg :args (1))
(step t279 (cl (or y$v3_1517448506_557_op (not (not y$1131))) (or y$v3_1517448506_557_op (not (not y$1131)))) :rule resolution :premises (t276 t277 t278))
(step t280 (cl (or y$v3_1517448506_557_op (not (not y$1131)))) :rule contraction :premises (t279))
(step t281 (cl (or y$v3_1517448506_557_op y$1131)) :rule resolution :premises (t272 t275 t280))
(step t282 (cl y$v3_1517448506_557_op y$1131) :rule or :premises (t281))
(step t283 (cl y$1131 y$v3_1517448506_557_op) :rule reordering :premises (t282))
(step t284 (cl (not y$1131) (not y$1140)) :rule equiv1 :premises (a735))
(step t285 (cl (not (and y$1139 y$1140)) y$1140) :rule and_pos :args (1))
(step t286 (cl y$1140 (not (and y$1139 y$1140))) :rule reordering :premises (t285))
(step t287 (cl (not y$v3_1517448506_562_op) (and y$1139 y$1140)) :rule equiv1 :premises (a736))
(step t288 (cl (and y$1139 y$1140) (not y$v3_1517448506_562_op)) :rule reordering :premises (t287))
(step t289 (cl (not (= (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op y$1143))) (not (or y$v3_1517448506_562_op (not (not y$1143)))) (or y$v3_1517448506_562_op y$1143)) :rule equiv_pos2)
(step t290 (cl (= y$v3_1517448506_562_op y$v3_1517448506_562_op)) :rule refl)
(step t291 (cl (= (not (not y$1143)) y$1143)) :rule rare_rewrite :args ("bool-double-not-elim" y$1143))
(step t292 (cl (= (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op y$1143))) :rule cong :premises (t290 t291))
(step t293 (cl y$v3_1517448506_562_op (not (not y$1143))) :rule equiv2 :premises (a737))
(step t294 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (not y$v3_1517448506_562_op)) :rule or_neg :args (0))
(step t295 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (not (not (not y$1143)))) :rule or_neg :args (1))
(step t296 (cl (or y$v3_1517448506_562_op (not (not y$1143))) (or y$v3_1517448506_562_op (not (not y$1143)))) :rule resolution :premises (t293 t294 t295))
(step t297 (cl (or y$v3_1517448506_562_op (not (not y$1143)))) :rule contraction :premises (t296))
(step t298 (cl (or y$v3_1517448506_562_op y$1143)) :rule resolution :premises (t289 t292 t297))
(step t299 (cl y$v3_1517448506_562_op y$1143) :rule or :premises (t298))
(step t300 (cl y$1143 y$v3_1517448506_562_op) :rule reordering :premises (t299))
(step t301 (cl (and y$a_too_small_resistance$next_rhs_op y$1143) (not y$a_too_small_resistance$next_rhs_op) (not y$1143)) :rule and_neg)
(step t302 (cl (not y$1143) (and y$a_too_small_resistance$next_rhs_op y$1143) (not y$a_too_small_resistance$next_rhs_op)) :rule reordering :premises (t301))
(step t303 (cl y$v3_1517448506_563_op (not (and y$a_too_small_resistance$next_rhs_op y$1143))) :rule equiv2 :premises (a738))
(step t304 (cl (not y$v3_1517448506_563_op) (not y$1146)) :rule equiv1 :premises (a739))
(step t305 (cl (not y$1146) (not y$v3_1517448506_563_op)) :rule reordering :premises (t304))
(step t306 (cl (not (and y$1138 y$1146)) y$1146) :rule and_pos :args (1))
(step t307 (cl y$1146 (not (and y$1138 y$1146))) :rule reordering :premises (t306))
(step t308 (cl (not y$v3_1517448506_565_op) (and y$1138 y$1146)) :rule equiv1 :premises (a740))
(step t309 (cl (and y$1138 y$1146) (not y$v3_1517448506_565_op)) :rule reordering :premises (t308))
(step t310 (cl (not (= (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op y$1149))) (not (or y$v3_1517448506_565_op (not (not y$1149)))) (or y$v3_1517448506_565_op y$1149)) :rule equiv_pos2)
(step t311 (cl (= y$v3_1517448506_565_op y$v3_1517448506_565_op)) :rule refl)
(step t312 (cl (= (not (not y$1149)) y$1149)) :rule rare_rewrite :args ("bool-double-not-elim" y$1149))
(step t313 (cl (= (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op y$1149))) :rule cong :premises (t311 t312))
(step t314 (cl y$v3_1517448506_565_op (not (not y$1149))) :rule equiv2 :premises (a741))
(step t315 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (not y$v3_1517448506_565_op)) :rule or_neg :args (0))
(step t316 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (not (not (not y$1149)))) :rule or_neg :args (1))
(step t317 (cl (or y$v3_1517448506_565_op (not (not y$1149))) (or y$v3_1517448506_565_op (not (not y$1149)))) :rule resolution :premises (t314 t315 t316))
(step t318 (cl (or y$v3_1517448506_565_op (not (not y$1149)))) :rule contraction :premises (t317))
(step t319 (cl (or y$v3_1517448506_565_op y$1149)) :rule resolution :premises (t310 t313 t318))
(step t320 (cl y$v3_1517448506_565_op y$1149) :rule or :premises (t319))
(step t321 (cl y$1149 y$v3_1517448506_565_op) :rule reordering :premises (t320))
(step t322 (cl (not y$1149) (not y$1150)) :rule equiv1 :premises (a742))
(step t323 (cl (not (and y$1150 y$1158)) y$1150) :rule and_pos :args (0))
(step t324 (cl y$1150 (not (and y$1150 y$1158))) :rule reordering :premises (t323))
(step t325 (cl (not y$v3_1517448506_570_op) (and y$1150 y$1158)) :rule equiv1 :premises (a749))
(step t326 (cl (and y$1150 y$1158) (not y$v3_1517448506_570_op)) :rule reordering :premises (t325))
(step t327 (cl (not (= (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op y$1161))) (not (or y$v3_1517448506_570_op (not (not y$1161)))) (or y$v3_1517448506_570_op y$1161)) :rule equiv_pos2)
(step t328 (cl (= y$v3_1517448506_570_op y$v3_1517448506_570_op)) :rule refl)
(step t329 (cl (= (not (not y$1161)) y$1161)) :rule rare_rewrite :args ("bool-double-not-elim" y$1161))
(step t330 (cl (= (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op y$1161))) :rule cong :premises (t328 t329))
(step t331 (cl y$v3_1517448506_570_op (not (not y$1161))) :rule equiv2 :premises (a750))
(step t332 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (not y$v3_1517448506_570_op)) :rule or_neg :args (0))
(step t333 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (not (not (not y$1161)))) :rule or_neg :args (1))
(step t334 (cl (or y$v3_1517448506_570_op (not (not y$1161))) (or y$v3_1517448506_570_op (not (not y$1161)))) :rule resolution :premises (t331 t332 t333))
(step t335 (cl (or y$v3_1517448506_570_op (not (not y$1161)))) :rule contraction :premises (t334))
(step t336 (cl (or y$v3_1517448506_570_op y$1161)) :rule resolution :premises (t327 t330 t335))
(step t337 (cl y$v3_1517448506_570_op y$1161) :rule or :premises (t336))
(step t338 (cl y$1161 y$v3_1517448506_570_op) :rule reordering :premises (t337))
(step t339 (cl (not y$1161) (not y$1162)) :rule equiv1 :premises (a751))
(step t340 (cl (not (and y$1162 y$1170)) y$1162) :rule and_pos :args (0))
(step t341 (cl y$1162 (not (and y$1162 y$1170))) :rule reordering :premises (t340))
(step t342 (cl (not y$v3_1517448506_575_op) (and y$1162 y$1170)) :rule equiv1 :premises (a758))
(step t343 (cl (and y$1162 y$1170) (not y$v3_1517448506_575_op)) :rule reordering :premises (t342))
(step t344 (cl (not (= (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op y$1173))) (not (or y$v3_1517448506_575_op (not (not y$1173)))) (or y$v3_1517448506_575_op y$1173)) :rule equiv_pos2)
(step t345 (cl (= y$v3_1517448506_575_op y$v3_1517448506_575_op)) :rule refl)
(step t346 (cl (= (not (not y$1173)) y$1173)) :rule rare_rewrite :args ("bool-double-not-elim" y$1173))
(step t347 (cl (= (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op y$1173))) :rule cong :premises (t345 t346))
(step t348 (cl y$v3_1517448506_575_op (not (not y$1173))) :rule equiv2 :premises (a759))
(step t349 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (not y$v3_1517448506_575_op)) :rule or_neg :args (0))
(step t350 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (not (not (not y$1173)))) :rule or_neg :args (1))
(step t351 (cl (or y$v3_1517448506_575_op (not (not y$1173))) (or y$v3_1517448506_575_op (not (not y$1173)))) :rule resolution :premises (t348 t349 t350))
(step t352 (cl (or y$v3_1517448506_575_op (not (not y$1173)))) :rule contraction :premises (t351))
(step t353 (cl (or y$v3_1517448506_575_op y$1173)) :rule resolution :premises (t344 t347 t352))
(step t354 (cl y$v3_1517448506_575_op y$1173) :rule or :premises (t353))
(step t355 (cl y$1173 y$v3_1517448506_575_op) :rule reordering :premises (t354))
(step t356 (cl (not y$1173) (not y$1174)) :rule equiv1 :premises (a760))
(step t357 (cl (not (and y$1174 y$1181)) y$1174) :rule and_pos :args (0))
(step t358 (cl y$1174 (not (and y$1174 y$1181))) :rule reordering :premises (t357))
(step t359 (cl (not y$v3_1517448506_580_op) (and y$1174 y$1181)) :rule equiv1 :premises (a766))
(step t360 (cl (and y$1174 y$1181) (not y$v3_1517448506_580_op)) :rule reordering :premises (t359))
(step t361 (cl (not (= (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op y$1184))) (not (or y$v3_1517448506_580_op (not (not y$1184)))) (or y$v3_1517448506_580_op y$1184)) :rule equiv_pos2)
(step t362 (cl (= y$v3_1517448506_580_op y$v3_1517448506_580_op)) :rule refl)
(step t363 (cl (= (not (not y$1184)) y$1184)) :rule rare_rewrite :args ("bool-double-not-elim" y$1184))
(step t364 (cl (= (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op y$1184))) :rule cong :premises (t362 t363))
(step t365 (cl y$v3_1517448506_580_op (not (not y$1184))) :rule equiv2 :premises (a767))
(step t366 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (not y$v3_1517448506_580_op)) :rule or_neg :args (0))
(step t367 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (not (not (not y$1184)))) :rule or_neg :args (1))
(step t368 (cl (or y$v3_1517448506_580_op (not (not y$1184))) (or y$v3_1517448506_580_op (not (not y$1184)))) :rule resolution :premises (t365 t366 t367))
(step t369 (cl (or y$v3_1517448506_580_op (not (not y$1184)))) :rule contraction :premises (t368))
(step t370 (cl (or y$v3_1517448506_580_op y$1184)) :rule resolution :premises (t361 t364 t369))
(step t371 (cl y$v3_1517448506_580_op y$1184) :rule or :premises (t370))
(step t372 (cl y$1184 y$v3_1517448506_580_op) :rule reordering :premises (t371))
(step t373 (cl (not y$1184) (not y$1185)) :rule equiv1 :premises (a768))
(step t374 (cl (not (and y$1185 y$1193)) y$1185) :rule and_pos :args (0))
(step t375 (cl y$1185 (not (and y$1185 y$1193))) :rule reordering :premises (t374))
(step t376 (cl (not y$v3_1517448506_585_op) (and y$1185 y$1193)) :rule equiv1 :premises (a775))
(step t377 (cl (and y$1185 y$1193) (not y$v3_1517448506_585_op)) :rule reordering :premises (t376))
(step t378 (cl (not (= (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op y$1196))) (not (or y$v3_1517448506_585_op (not (not y$1196)))) (or y$v3_1517448506_585_op y$1196)) :rule equiv_pos2)
(step t379 (cl (= y$v3_1517448506_585_op y$v3_1517448506_585_op)) :rule refl)
(step t380 (cl (= (not (not y$1196)) y$1196)) :rule rare_rewrite :args ("bool-double-not-elim" y$1196))
(step t381 (cl (= (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op y$1196))) :rule cong :premises (t379 t380))
(step t382 (cl y$v3_1517448506_585_op (not (not y$1196))) :rule equiv2 :premises (a776))
(step t383 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (not y$v3_1517448506_585_op)) :rule or_neg :args (0))
(step t384 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (not (not (not y$1196)))) :rule or_neg :args (1))
(step t385 (cl (or y$v3_1517448506_585_op (not (not y$1196))) (or y$v3_1517448506_585_op (not (not y$1196)))) :rule resolution :premises (t382 t383 t384))
(step t386 (cl (or y$v3_1517448506_585_op (not (not y$1196)))) :rule contraction :premises (t385))
(step t387 (cl (or y$v3_1517448506_585_op y$1196)) :rule resolution :premises (t378 t381 t386))
(step t388 (cl y$v3_1517448506_585_op y$1196) :rule or :premises (t387))
(step t389 (cl y$1196 y$v3_1517448506_585_op) :rule reordering :premises (t388))
(step t390 (cl (not y$1196) (not y$1197)) :rule equiv1 :premises (a777))
(step t391 (cl (not (and y$1197 y$1205)) y$1197) :rule and_pos :args (0))
(step t392 (cl y$1197 (not (and y$1197 y$1205))) :rule reordering :premises (t391))
(step t393 (cl (not y$v3_1517448506_590_op) (and y$1197 y$1205)) :rule equiv1 :premises (a784))
(step t394 (cl (and y$1197 y$1205) (not y$v3_1517448506_590_op)) :rule reordering :premises (t393))
(step t395 (cl (not (= (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op y$1208))) (not (or y$v3_1517448506_590_op (not (not y$1208)))) (or y$v3_1517448506_590_op y$1208)) :rule equiv_pos2)
(step t396 (cl (= y$v3_1517448506_590_op y$v3_1517448506_590_op)) :rule refl)
(step t397 (cl (= (not (not y$1208)) y$1208)) :rule rare_rewrite :args ("bool-double-not-elim" y$1208))
(step t398 (cl (= (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op y$1208))) :rule cong :premises (t396 t397))
(step t399 (cl y$v3_1517448506_590_op (not (not y$1208))) :rule equiv2 :premises (a785))
(step t400 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (not y$v3_1517448506_590_op)) :rule or_neg :args (0))
(step t401 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (not (not (not y$1208)))) :rule or_neg :args (1))
(step t402 (cl (or y$v3_1517448506_590_op (not (not y$1208))) (or y$v3_1517448506_590_op (not (not y$1208)))) :rule resolution :premises (t399 t400 t401))
(step t403 (cl (or y$v3_1517448506_590_op (not (not y$1208)))) :rule contraction :premises (t402))
(step t404 (cl (or y$v3_1517448506_590_op y$1208)) :rule resolution :premises (t395 t398 t403))
(step t405 (cl y$v3_1517448506_590_op y$1208) :rule or :premises (t404))
(step t406 (cl y$1208 y$v3_1517448506_590_op) :rule reordering :premises (t405))
(step t407 (cl (not y$1208) (not y$1209)) :rule equiv1 :premises (a786))
(step t408 (cl (not (and y$v3_1517448506_544_op y$1209)) y$1209) :rule and_pos :args (1))
(step t409 (cl y$1209 (not (and y$v3_1517448506_544_op y$1209))) :rule reordering :premises (t408))
(step t410 (cl (not y$v3_1517448506_591_op) (and y$v3_1517448506_544_op y$1209)) :rule equiv1 :premises (a787))
(step t411 (cl (and y$v3_1517448506_544_op y$1209) (not y$v3_1517448506_591_op)) :rule reordering :premises (t410))
(step t412 (cl (not (and y$v3_1517448506_591_op y$1216)) y$v3_1517448506_591_op) :rule and_pos :args (0))
(step t413 (cl y$v3_1517448506_591_op (not (and y$v3_1517448506_591_op y$1216))) :rule reordering :premises (t412))
(step t414 (cl (not y$v3_1517448506_594_op) (and y$v3_1517448506_591_op y$1216)) :rule equiv1 :premises (a792))
(step t415 (cl (and y$v3_1517448506_591_op y$1216) (not y$v3_1517448506_594_op)) :rule reordering :premises (t414))
(step t416 (cl (not (and y$v3_1517448506_594_op y$1256)) y$v3_1517448506_594_op) :rule and_pos :args (0))
(step t417 (cl y$v3_1517448506_594_op (not (and y$v3_1517448506_594_op y$1256))) :rule reordering :premises (t416))
(step t418 (cl (not y$v3_1517448506_611_op) (and y$v3_1517448506_594_op y$1256)) :rule equiv1 :premises (a821))
(step t419 (cl (and y$v3_1517448506_594_op y$1256) (not y$v3_1517448506_611_op)) :rule reordering :premises (t418))
(step t420 (cl (not (and y$v3_1517448506_611_op y$1263)) y$v3_1517448506_611_op) :rule and_pos :args (0))
(step t421 (cl y$v3_1517448506_611_op (not (and y$v3_1517448506_611_op y$1263))) :rule reordering :premises (t420))
(step t422 (cl (not y$v3_1517448506_614_op) (and y$v3_1517448506_611_op y$1263)) :rule equiv1 :premises (a826))
(step t423 (cl (and y$v3_1517448506_611_op y$1263) (not y$v3_1517448506_614_op)) :rule reordering :premises (t422))
(step t424 (cl (not (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) y$v3_1517448506_614_op) :rule and_pos :args (1))
(step t425 (cl y$v3_1517448506_614_op (not (and y$v3_1517448506_530_op y$v3_1517448506_614_op))) :rule reordering :premises (t424))
(step t426 (cl (not y$v3_1517448506_615_op) (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) :rule equiv1 :premises (a827))
(step t427 (cl (and y$v3_1517448506_530_op y$v3_1517448506_614_op) (not y$v3_1517448506_615_op)) :rule reordering :premises (t426))
(step t428 (cl (not (and y$41 y$v3_1517448506_615_op)) y$v3_1517448506_615_op) :rule and_pos :args (1))
(step t429 (cl y$v3_1517448506_615_op (not (and y$41 y$v3_1517448506_615_op))) :rule reordering :premises (t428))
(step t430 (cl (not y$v3_1517448506_616_op) (and y$41 y$v3_1517448506_615_op)) :rule equiv1 :premises (a828))
(step t431 (cl (and y$41 y$v3_1517448506_615_op) (not y$v3_1517448506_616_op)) :rule reordering :premises (t430))
(step t432 (cl (not (= (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op y$1270))) (not (or y$v3_1517448506_616_op (not (not y$1270)))) (or y$v3_1517448506_616_op y$1270)) :rule equiv_pos2)
(step t433 (cl (= y$v3_1517448506_616_op y$v3_1517448506_616_op)) :rule refl)
(step t434 (cl (= (not (not y$1270)) y$1270)) :rule rare_rewrite :args ("bool-double-not-elim" y$1270))
(step t435 (cl (= (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op y$1270))) :rule cong :premises (t433 t434))
(step t436 (cl y$v3_1517448506_616_op (not (not y$1270))) :rule equiv2 :premises (a829))
(step t437 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (not y$v3_1517448506_616_op)) :rule or_neg :args (0))
(step t438 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (not (not (not y$1270)))) :rule or_neg :args (1))
(step t439 (cl (or y$v3_1517448506_616_op (not (not y$1270))) (or y$v3_1517448506_616_op (not (not y$1270)))) :rule resolution :premises (t436 t437 t438))
(step t440 (cl (or y$v3_1517448506_616_op (not (not y$1270)))) :rule contraction :premises (t439))
(step t441 (cl (or y$v3_1517448506_616_op y$1270)) :rule resolution :premises (t432 t435 t440))
(step t442 (cl y$v3_1517448506_616_op y$1270) :rule or :premises (t441))
(step t443 (cl y$1270 y$v3_1517448506_616_op) :rule reordering :premises (t442))
(step t444 (cl (not (= y$1270 y$dve_invalid$next)) (not y$1270) y$dve_invalid$next) :rule equiv_pos2)
(step t445 (cl y$dve_invalid$next (not y$1270) (not (= y$1270 y$dve_invalid$next))) :rule reordering :premises (t444))
(step t446 (cl (not y$dve_invalid$next) (not y$1319)) :rule equiv1 :premises (a832))
(step t447 (cl (not y$1319) (not y$dve_invalid$next)) :rule reordering :premises (t446))
(step t448 (cl (not (and y$1319 y$v3_1517448506_38$next_op)) y$1319) :rule and_pos :args (0))
(step t449 (cl y$1319 (not (and y$1319 y$v3_1517448506_38$next_op))) :rule reordering :premises (t448))
(step t450 (cl (not y$id37$next_op) (and y$1319 y$v3_1517448506_38$next_op)) :rule equiv1 :premises (a844))
(step t451 (cl (and y$1319 y$v3_1517448506_38$next_op) (not y$id37$next_op)) :rule reordering :premises (t450))
(step t452 (cl (not (= (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op y$1342))) (not (or y$id37$next_op (not (not y$1342)))) (or y$id37$next_op y$1342)) :rule equiv_pos2)
(step t453 (cl (= y$id37$next_op y$id37$next_op)) :rule refl)
(step t454 (cl (= (not (not y$1342)) y$1342)) :rule rare_rewrite :args ("bool-double-not-elim" y$1342))
(step t455 (cl (= (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op y$1342))) :rule cong :premises (t453 t454))
(step t456 (cl y$id37$next_op (not (not y$1342))) :rule equiv2 :premises (a845))
(step t457 (cl (or y$id37$next_op (not (not y$1342))) (not y$id37$next_op)) :rule or_neg :args (0))
(step t458 (cl (or y$id37$next_op (not (not y$1342))) (not (not (not y$1342)))) :rule or_neg :args (1))
(step t459 (cl (or y$id37$next_op (not (not y$1342))) (or y$id37$next_op (not (not y$1342)))) :rule resolution :premises (t456 t457 t458))
(step t460 (cl (or y$id37$next_op (not (not y$1342)))) :rule contraction :premises (t459))
(step t461 (cl (or y$id37$next_op y$1342)) :rule resolution :premises (t452 t455 t460))
(step t462 (cl y$id37$next_op y$1342) :rule or :premises (t461))
(step t463 (cl y$1342 y$id37$next_op) :rule reordering :premises (t462))
(step t464 (cl (not (= y$1342 y$prop$next)) (not y$1342) y$prop$next) :rule equiv_pos2)
(step t465 (cl y$prop$next (not y$1342) (not (= y$1342 y$prop$next))) :rule reordering :premises (t464))
(step t466 (cl (not y$prop$next) (not y$1318)) :rule equiv1 :premises (a847))
(step t467 (cl (not y$1318) (not y$prop$next)) :rule reordering :premises (t466))
(step t468 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1318) :rule and_pos :args (4))
(step t469 (cl y$1318 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t468))
(step t470 (cl (not y$1346) (and y$prop y$1284 y$1272 y$1343 y$1318)) :rule equiv1 :premises (a848))
(step t471 (cl (and y$prop y$1284 y$1272 y$1343 y$1318) (not y$1346)) :rule reordering :premises (t470))
(step t472 (cl (and y$prop y$1284 y$1272 y$1343 y$1318)) :rule resolution :premises (t471 a849))
(step t473 (cl y$1318) :rule resolution :premises (t469 t472))
(step t474 (cl (not y$prop$next)) :rule resolution :premises (t467 t473))
(step t475 (cl (not (= (= y$1343 (= y$prop$next y$1342)) (= y$1343 (= y$1342 y$prop$next)))) (not (= y$1343 (= y$prop$next y$1342))) (= y$1343 (= y$1342 y$prop$next))) :rule equiv_pos2)
(step t476 (cl (= y$1343 y$1343)) :rule refl)
(step t477 (cl (= (= y$prop$next y$1342) (= y$1342 y$prop$next))) :rule rare_rewrite :args ("eq-symm" y$prop$next y$1342))
(step t478 (cl (= (= y$1343 (= y$prop$next y$1342)) (= y$1343 (= y$1342 y$prop$next)))) :rule cong :premises (t476 t477))
(step t479 (cl (= y$1343 (= y$1342 y$prop$next))) :rule resolution :premises (t475 t478 a846))
(step t480 (cl (not y$1343) (= y$1342 y$prop$next)) :rule equiv1 :premises (t479))
(step t481 (cl (= y$1342 y$prop$next) (not y$1343)) :rule reordering :premises (t480))
(step t482 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1343) :rule and_pos :args (3))
(step t483 (cl y$1343 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t482))
(step t484 (cl y$1343) :rule resolution :premises (t483 t472))
(step t485 (cl (= y$1342 y$prop$next)) :rule resolution :premises (t481 t484))
(step t486 (cl (not y$1342)) :rule resolution :premises (t465 t474 t485))
(step t487 (cl y$id37$next_op) :rule resolution :premises (t463 t486))
(step t488 (cl (and y$1319 y$v3_1517448506_38$next_op)) :rule resolution :premises (t451 t487))
(step t489 (cl y$1319) :rule resolution :premises (t449 t488))
(step t490 (cl (not y$dve_invalid$next)) :rule resolution :premises (t447 t489))
(step t491 (cl (not (= (= y$1271 (= y$dve_invalid$next y$1270)) (= y$1271 (= y$1270 y$dve_invalid$next)))) (not (= y$1271 (= y$dve_invalid$next y$1270))) (= y$1271 (= y$1270 y$dve_invalid$next))) :rule equiv_pos2)
(step t492 (cl (= y$1271 y$1271)) :rule refl)
(step t493 (cl (= (= y$dve_invalid$next y$1270) (= y$1270 y$dve_invalid$next))) :rule rare_rewrite :args ("eq-symm" y$dve_invalid$next y$1270))
(step t494 (cl (= (= y$1271 (= y$dve_invalid$next y$1270)) (= y$1271 (= y$1270 y$dve_invalid$next)))) :rule cong :premises (t492 t493))
(step t495 (cl (= y$1271 (= y$1270 y$dve_invalid$next))) :rule resolution :premises (t491 t494 a830))
(step t496 (cl (not y$1271) (= y$1270 y$dve_invalid$next)) :rule equiv1 :premises (t495))
(step t497 (cl (= y$1270 y$dve_invalid$next) (not y$1271)) :rule reordering :premises (t496))
(step t498 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$1271) :rule and_pos :args (23))
(step t499 (cl y$1271 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t498))
(step t500 (cl (not y$1272) (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) :rule equiv1 :premises (a831))
(step t501 (cl (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271) (not y$1272)) :rule reordering :premises (t500))
(step t502 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1272) :rule and_pos :args (2))
(step t503 (cl y$1272 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t502))
(step t504 (cl y$1272) :rule resolution :premises (t503 t472))
(step t505 (cl (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) :rule resolution :premises (t501 t504))
(step t506 (cl y$1271) :rule resolution :premises (t499 t505))
(step t507 (cl (= y$1270 y$dve_invalid$next)) :rule resolution :premises (t497 t506))
(step t508 (cl (not y$1270)) :rule resolution :premises (t445 t490 t507))
(step t509 (cl y$v3_1517448506_616_op) :rule resolution :premises (t443 t508))
(step t510 (cl (and y$41 y$v3_1517448506_615_op)) :rule resolution :premises (t431 t509))
(step t511 (cl y$v3_1517448506_615_op) :rule resolution :premises (t429 t510))
(step t512 (cl (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) :rule resolution :premises (t427 t511))
(step t513 (cl y$v3_1517448506_614_op) :rule resolution :premises (t425 t512))
(step t514 (cl (and y$v3_1517448506_611_op y$1263)) :rule resolution :premises (t423 t513))
(step t515 (cl y$v3_1517448506_611_op) :rule resolution :premises (t421 t514))
(step t516 (cl (and y$v3_1517448506_594_op y$1256)) :rule resolution :premises (t419 t515))
(step t517 (cl y$v3_1517448506_594_op) :rule resolution :premises (t417 t516))
(step t518 (cl (and y$v3_1517448506_591_op y$1216)) :rule resolution :premises (t415 t517))
(step t519 (cl y$v3_1517448506_591_op) :rule resolution :premises (t413 t518))
(step t520 (cl (and y$v3_1517448506_544_op y$1209)) :rule resolution :premises (t411 t519))
(step t521 (cl y$1209) :rule resolution :premises (t409 t520))
(step t522 (cl (not y$1208)) :rule resolution :premises (t407 t521))
(step t523 (cl y$v3_1517448506_590_op) :rule resolution :premises (t406 t522))
(step t524 (cl (and y$1197 y$1205)) :rule resolution :premises (t394 t523))
(step t525 (cl y$1197) :rule resolution :premises (t392 t524))
(step t526 (cl (not y$1196)) :rule resolution :premises (t390 t525))
(step t527 (cl y$v3_1517448506_585_op) :rule resolution :premises (t389 t526))
(step t528 (cl (and y$1185 y$1193)) :rule resolution :premises (t377 t527))
(step t529 (cl y$1185) :rule resolution :premises (t375 t528))
(step t530 (cl (not y$1184)) :rule resolution :premises (t373 t529))
(step t531 (cl y$v3_1517448506_580_op) :rule resolution :premises (t372 t530))
(step t532 (cl (and y$1174 y$1181)) :rule resolution :premises (t360 t531))
(step t533 (cl y$1174) :rule resolution :premises (t358 t532))
(step t534 (cl (not y$1173)) :rule resolution :premises (t356 t533))
(step t535 (cl y$v3_1517448506_575_op) :rule resolution :premises (t355 t534))
(step t536 (cl (and y$1162 y$1170)) :rule resolution :premises (t343 t535))
(step t537 (cl y$1162) :rule resolution :premises (t341 t536))
(step t538 (cl (not y$1161)) :rule resolution :premises (t339 t537))
(step t539 (cl y$v3_1517448506_570_op) :rule resolution :premises (t338 t538))
(step t540 (cl (and y$1150 y$1158)) :rule resolution :premises (t326 t539))
(step t541 (cl y$1150) :rule resolution :premises (t324 t540))
(step t542 (cl (not y$1149)) :rule resolution :premises (t322 t541))
(step t543 (cl y$v3_1517448506_565_op) :rule resolution :premises (t321 t542))
(step t544 (cl (and y$1138 y$1146)) :rule resolution :premises (t309 t543))
(step t545 (cl y$1146) :rule resolution :premises (t307 t544))
(step t546 (cl (not y$v3_1517448506_563_op)) :rule resolution :premises (t305 t545))
(step t547 (cl (not (and y$a_too_small_resistance$next_rhs_op y$1143))) :rule resolution :premises (t303 t546))
(step t548 (cl (not (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) (not y$a_too_small_resistance$next) y$a_too_small_resistance$next_rhs_op) :rule equiv_pos2)
(step t549 (cl y$a_too_small_resistance$next_rhs_op (not y$a_too_small_resistance$next) (not (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op))) :rule reordering :premises (t548))
(step t550 (cl (not (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) (not (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv_pos2)
(step t551 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next)) :rule refl)
(step t552 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule rare_rewrite :args ("eq-symm" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32))
(step t553 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule rare_rewrite :args ("eq-symm" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))
(step t554 (cl (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) :rule cong :premises (t551 t552 t553))
(step t555 (cl (not (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)))) (not (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule equiv_pos2)
(step t556 (cl (= (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv_simplify)
(step t557 (cl (not (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule equiv1 :premises (t556))
(step t558 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule refl)
(step t559 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule refl)
(step t560 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule cong :premises (t558 t559))
(step t561 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))
(step t562 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) true)) :rule trans :premises (t560 t561))
(step t563 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule resolution :premises (t557 t562))
(step t564 (cl (= y$n1s32 y$n1s32)) :rule refl)
(step t565 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32))) :rule cong :premises (t563 t564))
(step t566 (cl (= y$n0s32 y$n0s32)) :rule refl)
(step t567 (cl (= (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule cong :premises (t563 t566))
(step t568 (cl (= (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32)))) :rule cong :premises (t551 t565 t567))
(step t569 (cl (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule hole)
(step t570 (cl (ite y$a_too_small_resistance$next (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n0s32))) :rule resolution :premises (t555 t568 t569))
(step t571 (cl (ite y$a_too_small_resistance$next (= y$n1s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule resolution :premises (t550 t554 t570))
(step t572 (cl y$a_too_small_resistance$next (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule ite1 :premises (t571))
(step t573 (cl (= (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false) (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))))) :rule equiv_simplify)
(step t574 (cl (not (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false)) (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule equiv1 :premises (t573))
(step t575 (cl (not (= (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) (not (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule equiv_pos2)
(step t576 (cl (= y$v3_1517448506_36$next_op y$v3_1517448506_36$next_op)) :rule refl)
(step t577 (cl (= (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule cong :premises (t576 t563))
(step t578 (cl (= y$v3_1517448506_36$next_op (ite y$a_too_small_resistance$next y$n1s32 y$n0s32))) :rule resolution :premises (t575 t577 a841))
(step t579 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$v3_1517448506_36$next_op)) :rule symm :premises (t578))
(step t580 (cl (not y$1337) (= y$n1s32 y$v3_1517448506_36$next_op)) :rule equiv1 :premises (a842))
(step t581 (cl (= y$n1s32 y$v3_1517448506_36$next_op) (not y$1337)) :rule reordering :premises (t580))
(step t582 (cl (not (and y$1334 y$1337)) y$1337) :rule and_pos :args (1))
(step t583 (cl y$1337 (not (and y$1334 y$1337))) :rule reordering :premises (t582))
(step t584 (cl (not y$v3_1517448506_38$next_op) (and y$1334 y$1337)) :rule equiv1 :premises (a843))
(step t585 (cl (and y$1334 y$1337) (not y$v3_1517448506_38$next_op)) :rule reordering :premises (t584))
(step t586 (cl (not (and y$1319 y$v3_1517448506_38$next_op)) y$v3_1517448506_38$next_op) :rule and_pos :args (1))
(step t587 (cl y$v3_1517448506_38$next_op (not (and y$1319 y$v3_1517448506_38$next_op))) :rule reordering :premises (t586))
(step t588 (cl y$v3_1517448506_38$next_op) :rule resolution :premises (t587 t488))
(step t589 (cl (and y$1334 y$1337)) :rule resolution :premises (t585 t588))
(step t590 (cl y$1337) :rule resolution :premises (t583 t589))
(step t591 (cl (= y$n1s32 y$v3_1517448506_36$next_op)) :rule resolution :premises (t581 t590))
(step t592 (cl (= y$v3_1517448506_36$next_op y$n1s32)) :rule symm :premises (t591))
(step t593 (cl (= (ite y$a_too_small_resistance$next y$n1s32 y$n0s32) y$n1s32)) :rule trans :premises (t579 t592))
(step t594 (cl (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) (= y$n0s32 y$n1s32))) :rule cong :premises (t566 t593))
(step t595 (cl (= (= (= y$n0s32 y$n1s32) false) (not (= y$n0s32 y$n1s32)))) :rule equiv_simplify)
(step t596 (cl (= (= y$n0s32 y$n1s32) false) (not (not (= y$n0s32 y$n1s32)))) :rule equiv2 :premises (t595))
(step t597 (cl (not (not (not (= y$n0s32 y$n1s32)))) (= y$n0s32 y$n1s32)) :rule not_not)
(step t598 (cl (= (= y$n0s32 y$n1s32) false) (= y$n0s32 y$n1s32)) :rule resolution :premises (t596 t597))
(step t599 (cl (not (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) (not (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32)) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule equiv_pos2)
(step t600 (cl (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n0s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n1000s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n0s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n1000s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n6200s32 y$n0s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n5999s32)) (not (= y$n6200s32 y$n1000s32)) (not (= y$n6200s32 y$n5800s32)) (not (= y$n6200s32 y$n5s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n999s32 y$n5999s32)) (not (= y$n999s32 y$n1000s32)) (not (= y$n999s32 y$n5800s32)) (not (= y$n999s32 y$n5s32)) (not (= y$n5999s32 y$n1000s32)) (not (= y$n5999s32 y$n5800s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule distinct_elim)
(step t601 (cl (= (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n1s32)))) :rule refl)
(step t602 (cl (= (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n6200s32)))) :rule refl)
(step t603 (cl (= (= y$n16s32 y$n0s32) (= y$n0s32 y$n16s32))) :rule rare_rewrite :args ("eq-symm" y$n16s32 y$n0s32))
(step t604 (cl (= (not (= y$n16s32 y$n0s32)) (not (= y$n0s32 y$n16s32)))) :rule cong :premises (t603))
(step t605 (cl (= (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n999s32)))) :rule refl)
(step t606 (cl (= (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n5999s32)))) :rule refl)
(step t607 (cl (= (= y$n16s32 y$n1000s32) (= y$n1000s32 y$n16s32))) :rule rare_rewrite :args ("eq-symm" y$n16s32 y$n1000s32))
(step t608 (cl (= (not (= y$n16s32 y$n1000s32)) (not (= y$n1000s32 y$n16s32)))) :rule cong :premises (t607))
(step t609 (cl (= (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5800s32)))) :rule refl)
(step t610 (cl (= (not (= y$n16s32 y$n5s32)) (not (= y$n16s32 y$n5s32)))) :rule refl)
(step t611 (cl (= (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n6200s32)))) :rule refl)
(step t612 (cl (= (= y$n1s32 y$n0s32) (= y$n0s32 y$n1s32))) :rule rare_rewrite :args ("eq-symm" y$n1s32 y$n0s32))
(step t613 (cl (= (not (= y$n1s32 y$n0s32)) (not (= y$n0s32 y$n1s32)))) :rule cong :premises (t612))
(step t614 (cl (= (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n999s32)))) :rule refl)
(step t615 (cl (= (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n5999s32)))) :rule refl)
(step t616 (cl (= (= y$n1s32 y$n1000s32) (= y$n1000s32 y$n1s32))) :rule rare_rewrite :args ("eq-symm" y$n1s32 y$n1000s32))
(step t617 (cl (= (not (= y$n1s32 y$n1000s32)) (not (= y$n1000s32 y$n1s32)))) :rule cong :premises (t616))
(step t618 (cl (= (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5800s32)))) :rule refl)
(step t619 (cl (= (not (= y$n1s32 y$n5s32)) (not (= y$n1s32 y$n5s32)))) :rule refl)
(step t620 (cl (= (= y$n6200s32 y$n0s32) (= y$n0s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n0s32))
(step t621 (cl (= (not (= y$n6200s32 y$n0s32)) (not (= y$n0s32 y$n6200s32)))) :rule cong :premises (t620))
(step t622 (cl (= (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n999s32)))) :rule refl)
(step t623 (cl (= (= y$n6200s32 y$n5999s32) (= y$n5999s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5999s32))
(step t624 (cl (= (not (= y$n6200s32 y$n5999s32)) (not (= y$n5999s32 y$n6200s32)))) :rule cong :premises (t623))
(step t625 (cl (= (= y$n6200s32 y$n1000s32) (= y$n1000s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n1000s32))
(step t626 (cl (= (not (= y$n6200s32 y$n1000s32)) (not (= y$n1000s32 y$n6200s32)))) :rule cong :premises (t625))
(step t627 (cl (= (= y$n6200s32 y$n5800s32) (= y$n5800s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5800s32))
(step t628 (cl (= (not (= y$n6200s32 y$n5800s32)) (not (= y$n5800s32 y$n6200s32)))) :rule cong :premises (t627))
(step t629 (cl (= (= y$n6200s32 y$n5s32) (= y$n5s32 y$n6200s32))) :rule rare_rewrite :args ("eq-symm" y$n6200s32 y$n5s32))
(step t630 (cl (= (not (= y$n6200s32 y$n5s32)) (not (= y$n5s32 y$n6200s32)))) :rule cong :premises (t629))
(step t631 (cl (= (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n999s32)))) :rule refl)
(step t632 (cl (= (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n5999s32)))) :rule refl)
(step t633 (cl (= (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n1000s32)))) :rule refl)
(step t634 (cl (= (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5800s32)))) :rule refl)
(step t635 (cl (= (not (= y$n0s32 y$n5s32)) (not (= y$n0s32 y$n5s32)))) :rule refl)
(step t636 (cl (= (= y$n999s32 y$n5999s32) (= y$n5999s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5999s32))
(step t637 (cl (= (not (= y$n999s32 y$n5999s32)) (not (= y$n5999s32 y$n999s32)))) :rule cong :premises (t636))
(step t638 (cl (= (= y$n999s32 y$n1000s32) (= y$n1000s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n1000s32))
(step t639 (cl (= (not (= y$n999s32 y$n1000s32)) (not (= y$n1000s32 y$n999s32)))) :rule cong :premises (t638))
(step t640 (cl (= (= y$n999s32 y$n5800s32) (= y$n5800s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5800s32))
(step t641 (cl (= (not (= y$n999s32 y$n5800s32)) (not (= y$n5800s32 y$n999s32)))) :rule cong :premises (t640))
(step t642 (cl (= (= y$n999s32 y$n5s32) (= y$n5s32 y$n999s32))) :rule rare_rewrite :args ("eq-symm" y$n999s32 y$n5s32))
(step t643 (cl (= (not (= y$n999s32 y$n5s32)) (not (= y$n5s32 y$n999s32)))) :rule cong :premises (t642))
(step t644 (cl (= (= y$n5999s32 y$n1000s32) (= y$n1000s32 y$n5999s32))) :rule rare_rewrite :args ("eq-symm" y$n5999s32 y$n1000s32))
(step t645 (cl (= (not (= y$n5999s32 y$n1000s32)) (not (= y$n1000s32 y$n5999s32)))) :rule cong :premises (t644))
(step t646 (cl (= (= y$n5999s32 y$n5800s32) (= y$n5800s32 y$n5999s32))) :rule rare_rewrite :args ("eq-symm" y$n5999s32 y$n5800s32))
(step t647 (cl (= (not (= y$n5999s32 y$n5800s32)) (not (= y$n5800s32 y$n5999s32)))) :rule cong :premises (t646))
(step t648 (cl (= (not (= y$n5999s32 y$n5s32)) (not (= y$n5999s32 y$n5s32)))) :rule refl)
(step t649 (cl (= (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5800s32)))) :rule refl)
(step t650 (cl (= (not (= y$n1000s32 y$n5s32)) (not (= y$n1000s32 y$n5s32)))) :rule refl)
(step t651 (cl (= (not (= y$n5800s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule refl)
(step t652 (cl (= (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n16s32 y$n0s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n16s32 y$n1000s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n1s32 y$n0s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1s32 y$n1000s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n6200s32 y$n0s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n6200s32 y$n5999s32)) (not (= y$n6200s32 y$n1000s32)) (not (= y$n6200s32 y$n5800s32)) (not (= y$n6200s32 y$n5s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n999s32 y$n5999s32)) (not (= y$n999s32 y$n1000s32)) (not (= y$n999s32 y$n5800s32)) (not (= y$n999s32 y$n5s32)) (not (= y$n5999s32 y$n1000s32)) (not (= y$n5999s32 y$n5800s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule cong :premises (t601 t602 t604 t605 t606 t608 t609 t610 t611 t613 t614 t615 t617 t618 t619 t621 t622 t624 t626 t628 t630 t631 t632 t633 t634 t635 t637 t639 t641 t643 t645 t647 t648 t649 t650 t651))
(step t653 (cl (= (distinct y$n16s32 y$n1s32 y$n6200s32 y$n0s32 y$n999s32 y$n5999s32 y$n1000s32 y$n5800s32 y$n5s32) (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32))))) :rule trans :premises (t600 t652))
(step t654 (cl (and (not (= y$n16s32 y$n1s32)) (not (= y$n16s32 y$n6200s32)) (not (= y$n0s32 y$n16s32)) (not (= y$n16s32 y$n999s32)) (not (= y$n16s32 y$n5999s32)) (not (= y$n1000s32 y$n16s32)) (not (= y$n16s32 y$n5800s32)) (not (= y$n16s32 y$n5s32)) (not (= y$n1s32 y$n6200s32)) (not (= y$n0s32 y$n1s32)) (not (= y$n1s32 y$n999s32)) (not (= y$n1s32 y$n5999s32)) (not (= y$n1000s32 y$n1s32)) (not (= y$n1s32 y$n5800s32)) (not (= y$n1s32 y$n5s32)) (not (= y$n0s32 y$n6200s32)) (not (= y$n6200s32 y$n999s32)) (not (= y$n5999s32 y$n6200s32)) (not (= y$n1000s32 y$n6200s32)) (not (= y$n5800s32 y$n6200s32)) (not (= y$n5s32 y$n6200s32)) (not (= y$n0s32 y$n999s32)) (not (= y$n0s32 y$n5999s32)) (not (= y$n0s32 y$n1000s32)) (not (= y$n0s32 y$n5800s32)) (not (= y$n0s32 y$n5s32)) (not (= y$n5999s32 y$n999s32)) (not (= y$n1000s32 y$n999s32)) (not (= y$n5800s32 y$n999s32)) (not (= y$n5s32 y$n999s32)) (not (= y$n1000s32 y$n5999s32)) (not (= y$n5800s32 y$n5999s32)) (not (= y$n5999s32 y$n5s32)) (not (= y$n1000s32 y$n5800s32)) (not (= y$n1000s32 y$n5s32)) (not (= y$n5800s32 y$n5s32)))) :rule resolution :premises (t599 t653 a1))
(step t655 (cl (not (= y$n0s32 y$n1s32))) :rule and :premises (t654) :args (9))
(step t656 (cl (= (= y$n0s32 y$n1s32) false)) :rule resolution :premises (t598 t655))
(step t657 (cl (= (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)) false)) :rule trans :premises (t594 t656))
(step t658 (cl (not (= y$n0s32 (ite y$a_too_small_resistance$next y$n1s32 y$n0s32)))) :rule resolution :premises (t574 t657))
(step t659 (cl y$a_too_small_resistance$next) :rule resolution :premises (t572 t658))
(step t660 (cl (not y$274) (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) :rule equiv1 :premises (a138))
(step t661 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op) (not y$274)) :rule reordering :premises (t660))
(step t662 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$274) :rule and_pos :args (12))
(step t663 (cl y$274 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t662))
(step t664 (cl y$274) :rule resolution :premises (t663 t505))
(step t665 (cl (= y$a_too_small_resistance$next y$a_too_small_resistance$next_rhs_op)) :rule resolution :premises (t661 t664))
(step t666 (cl y$a_too_small_resistance$next_rhs_op) :rule resolution :premises (t549 t659 t665))
(step t667 (cl (not y$1143)) :rule resolution :premises (t302 t547 t666))
(step t668 (cl y$v3_1517448506_562_op) :rule resolution :premises (t300 t667))
(step t669 (cl (and y$1139 y$1140)) :rule resolution :premises (t288 t668))
(step t670 (cl y$1140) :rule resolution :premises (t286 t669))
(step t671 (cl (not y$1131)) :rule resolution :premises (t284 t670))
(step t672 (cl y$v3_1517448506_557_op) :rule resolution :premises (t283 t671))
(step t673 (cl (and y$1127 y$1128)) :rule resolution :premises (t271 t672))
(step t674 (cl y$1128) :rule resolution :premises (t269 t673))
(step t675 (cl (not y$1119)) :rule resolution :premises (t267 t674))
(step t676 (cl y$v3_1517448506_552_op) :rule resolution :premises (t266 t675))
(step t677 (cl (and y$1115 y$1116)) :rule resolution :premises (t254 t676))
(step t678 (cl y$1116) :rule resolution :premises (t252 t677))
(step t679 (cl (not y$1107)) :rule resolution :premises (t250 t678))
(step t680 (cl y$v3_1517448506_547_op) :rule resolution :premises (t249 t679))
(step t681 (cl (and y$1103 y$1104)) :rule resolution :premises (t237 t680))
(step t682 (cl y$1104) :rule resolution :premises (t235 t681))
(step t683 (cl (not y$209)) :rule resolution :premises (t233 t682))
(step t684 (cl y$v3_1517448506_114_op) :rule resolution :premises (t232 t683))
(step t685 (cl (and y$204 y$206)) :rule resolution :premises (t220 t684))
(step t686 (cl y$206) :rule resolution :premises (t218 t685))
(step t687 (cl y$v3_1517448506_174_op (not (and y$206 y$322))) :rule equiv2 :premises (a173))
(step t688 (cl (and y$261 y$v3_1517448506_174_op) (not y$261) (not y$v3_1517448506_174_op)) :rule and_neg)
(step t689 (cl (not y$261) (and y$261 y$v3_1517448506_174_op) (not y$v3_1517448506_174_op)) :rule reordering :premises (t688))
(step t690 (cl (not (and y$259 y$261)) y$261) :rule and_pos :args (1))
(step t691 (cl y$261 (not (and y$259 y$261))) :rule reordering :premises (t690))
(step t692 (cl (not y$v3_1517448506_142_op) (and y$259 y$261)) :rule equiv1 :premises (a130))
(step t693 (cl (and y$259 y$261) (not y$v3_1517448506_142_op)) :rule reordering :premises (t692))
(step t694 (cl (not (= (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op y$264))) (not (or y$v3_1517448506_142_op (not (not y$264)))) (or y$v3_1517448506_142_op y$264)) :rule equiv_pos2)
(step t695 (cl (= y$v3_1517448506_142_op y$v3_1517448506_142_op)) :rule refl)
(step t696 (cl (= (not (not y$264)) y$264)) :rule rare_rewrite :args ("bool-double-not-elim" y$264))
(step t697 (cl (= (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op y$264))) :rule cong :premises (t695 t696))
(step t698 (cl y$v3_1517448506_142_op (not (not y$264))) :rule equiv2 :premises (a131))
(step t699 (cl (or y$v3_1517448506_142_op (not (not y$264))) (not y$v3_1517448506_142_op)) :rule or_neg :args (0))
(step t700 (cl (or y$v3_1517448506_142_op (not (not y$264))) (not (not (not y$264)))) :rule or_neg :args (1))
(step t701 (cl (or y$v3_1517448506_142_op (not (not y$264))) (or y$v3_1517448506_142_op (not (not y$264)))) :rule resolution :premises (t698 t699 t700))
(step t702 (cl (or y$v3_1517448506_142_op (not (not y$264)))) :rule contraction :premises (t701))
(step t703 (cl (or y$v3_1517448506_142_op y$264)) :rule resolution :premises (t694 t697 t702))
(step t704 (cl y$v3_1517448506_142_op y$264) :rule or :premises (t703))
(step t705 (cl y$264 y$v3_1517448506_142_op) :rule reordering :premises (t704))
(step t706 (cl (not y$264) (not y$1139)) :rule equiv1 :premises (a734))
(step t707 (cl (not (and y$1139 y$1140)) y$1139) :rule and_pos :args (0))
(step t708 (cl y$1139 (not (and y$1139 y$1140))) :rule reordering :premises (t707))
(step t709 (cl y$1139) :rule resolution :premises (t708 t669))
(step t710 (cl (not y$264)) :rule resolution :premises (t706 t709))
(step t711 (cl y$v3_1517448506_142_op) :rule resolution :premises (t705 t710))
(step t712 (cl (and y$259 y$261)) :rule resolution :premises (t693 t711))
(step t713 (cl y$261) :rule resolution :premises (t691 t712))
(step t714 (cl y$a_S2$next_rhs_op (not (and y$261 y$v3_1517448506_174_op))) :rule equiv2 :premises (a174))
(step t715 (cl (and y$a_S2$next_rhs_op y$1202) (not y$a_S2$next_rhs_op) (not y$1202)) :rule and_neg)
(step t716 (cl (not y$1202) (and y$a_S2$next_rhs_op y$1202) (not y$a_S2$next_rhs_op)) :rule reordering :premises (t715))
(step t717 (cl (not (= (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op y$1202))) (not (or y$v3_1517448506_587_op (not (not y$1202)))) (or y$v3_1517448506_587_op y$1202)) :rule equiv_pos2)
(step t718 (cl (= y$v3_1517448506_587_op y$v3_1517448506_587_op)) :rule refl)
(step t719 (cl (= (not (not y$1202)) y$1202)) :rule rare_rewrite :args ("bool-double-not-elim" y$1202))
(step t720 (cl (= (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op y$1202))) :rule cong :premises (t718 t719))
(step t721 (cl y$v3_1517448506_587_op (not (not y$1202))) :rule equiv2 :premises (a781))
(step t722 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (not y$v3_1517448506_587_op)) :rule or_neg :args (0))
(step t723 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (not (not (not y$1202)))) :rule or_neg :args (1))
(step t724 (cl (or y$v3_1517448506_587_op (not (not y$1202))) (or y$v3_1517448506_587_op (not (not y$1202)))) :rule resolution :premises (t721 t722 t723))
(step t725 (cl (or y$v3_1517448506_587_op (not (not y$1202)))) :rule contraction :premises (t724))
(step t726 (cl (or y$v3_1517448506_587_op y$1202)) :rule resolution :premises (t717 t720 t725))
(step t727 (cl y$v3_1517448506_587_op y$1202) :rule or :premises (t726))
(step t728 (cl y$1202 y$v3_1517448506_587_op) :rule reordering :premises (t727))
(step t729 (cl (not y$v3_1517448506_587_op) (and y$1198 y$1199)) :rule equiv1 :premises (a780))
(step t730 (cl (and y$1198 y$1199) (not y$v3_1517448506_587_op)) :rule reordering :premises (t729))
(step t731 (cl (not (and y$1198 y$1199)) y$1199) :rule and_pos :args (1))
(step t732 (cl y$1199 (not (and y$1198 y$1199))) :rule reordering :premises (t731))
(step t733 (cl (not y$1190) (not y$1199)) :rule equiv1 :premises (a779))
(step t734 (cl (not (= (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op y$1190))) (not (or y$v3_1517448506_582_op (not (not y$1190)))) (or y$v3_1517448506_582_op y$1190)) :rule equiv_pos2)
(step t735 (cl (= y$v3_1517448506_582_op y$v3_1517448506_582_op)) :rule refl)
(step t736 (cl (= (not (not y$1190)) y$1190)) :rule rare_rewrite :args ("bool-double-not-elim" y$1190))
(step t737 (cl (= (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op y$1190))) :rule cong :premises (t735 t736))
(step t738 (cl y$v3_1517448506_582_op (not (not y$1190))) :rule equiv2 :premises (a772))
(step t739 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (not y$v3_1517448506_582_op)) :rule or_neg :args (0))
(step t740 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (not (not (not y$1190)))) :rule or_neg :args (1))
(step t741 (cl (or y$v3_1517448506_582_op (not (not y$1190))) (or y$v3_1517448506_582_op (not (not y$1190)))) :rule resolution :premises (t738 t739 t740))
(step t742 (cl (or y$v3_1517448506_582_op (not (not y$1190)))) :rule contraction :premises (t741))
(step t743 (cl (or y$v3_1517448506_582_op y$1190)) :rule resolution :premises (t734 t737 t742))
(step t744 (cl y$v3_1517448506_582_op y$1190) :rule or :premises (t743))
(step t745 (cl y$1190 y$v3_1517448506_582_op) :rule reordering :premises (t744))
(step t746 (cl (not y$v3_1517448506_582_op) (and y$1186 y$1187)) :rule equiv1 :premises (a771))
(step t747 (cl (and y$1186 y$1187) (not y$v3_1517448506_582_op)) :rule reordering :premises (t746))
(step t748 (cl (not (and y$1186 y$1187)) y$1187) :rule and_pos :args (1))
(step t749 (cl y$1187 (not (and y$1186 y$1187))) :rule reordering :premises (t748))
(step t750 (cl (not y$1178) (not y$1187)) :rule equiv1 :premises (a770))
(step t751 (cl (not (= (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op y$1178))) (not (or y$v3_1517448506_577_op (not (not y$1178)))) (or y$v3_1517448506_577_op y$1178)) :rule equiv_pos2)
(step t752 (cl (= y$v3_1517448506_577_op y$v3_1517448506_577_op)) :rule refl)
(step t753 (cl (= (not (not y$1178)) y$1178)) :rule rare_rewrite :args ("bool-double-not-elim" y$1178))
(step t754 (cl (= (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op y$1178))) :rule cong :premises (t752 t753))
(step t755 (cl y$v3_1517448506_577_op (not (not y$1178))) :rule equiv2 :premises (a763))
(step t756 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (not y$v3_1517448506_577_op)) :rule or_neg :args (0))
(step t757 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (not (not (not y$1178)))) :rule or_neg :args (1))
(step t758 (cl (or y$v3_1517448506_577_op (not (not y$1178))) (or y$v3_1517448506_577_op (not (not y$1178)))) :rule resolution :premises (t755 t756 t757))
(step t759 (cl (or y$v3_1517448506_577_op (not (not y$1178)))) :rule contraction :premises (t758))
(step t760 (cl (or y$v3_1517448506_577_op y$1178)) :rule resolution :premises (t751 t754 t759))
(step t761 (cl y$v3_1517448506_577_op y$1178) :rule or :premises (t760))
(step t762 (cl y$1178 y$v3_1517448506_577_op) :rule reordering :premises (t761))
(step t763 (cl (not y$v3_1517448506_577_op) (and y$293 y$1175)) :rule equiv1 :premises (a762))
(step t764 (cl (and y$293 y$1175) (not y$v3_1517448506_577_op)) :rule reordering :premises (t763))
(step t765 (cl (not (and y$293 y$1175)) y$1175) :rule and_pos :args (1))
(step t766 (cl y$1175 (not (and y$293 y$1175))) :rule reordering :premises (t765))
(step t767 (cl (not y$1167) (not y$1175)) :rule equiv1 :premises (a761))
(step t768 (cl (not (= (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op y$1167))) (not (or y$v3_1517448506_572_op (not (not y$1167)))) (or y$v3_1517448506_572_op y$1167)) :rule equiv_pos2)
(step t769 (cl (= y$v3_1517448506_572_op y$v3_1517448506_572_op)) :rule refl)
(step t770 (cl (= (not (not y$1167)) y$1167)) :rule rare_rewrite :args ("bool-double-not-elim" y$1167))
(step t771 (cl (= (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op y$1167))) :rule cong :premises (t769 t770))
(step t772 (cl y$v3_1517448506_572_op (not (not y$1167))) :rule equiv2 :premises (a755))
(step t773 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (not y$v3_1517448506_572_op)) :rule or_neg :args (0))
(step t774 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (not (not (not y$1167)))) :rule or_neg :args (1))
(step t775 (cl (or y$v3_1517448506_572_op (not (not y$1167))) (or y$v3_1517448506_572_op (not (not y$1167)))) :rule resolution :premises (t772 t773 t774))
(step t776 (cl (or y$v3_1517448506_572_op (not (not y$1167)))) :rule contraction :premises (t775))
(step t777 (cl (or y$v3_1517448506_572_op y$1167)) :rule resolution :premises (t768 t771 t776))
(step t778 (cl y$v3_1517448506_572_op y$1167) :rule or :premises (t777))
(step t779 (cl y$1167 y$v3_1517448506_572_op) :rule reordering :premises (t778))
(step t780 (cl (not y$v3_1517448506_572_op) (and y$1163 y$1164)) :rule equiv1 :premises (a754))
(step t781 (cl (and y$1163 y$1164) (not y$v3_1517448506_572_op)) :rule reordering :premises (t780))
(step t782 (cl (not (and y$1163 y$1164)) y$1164) :rule and_pos :args (1))
(step t783 (cl y$1164 (not (and y$1163 y$1164))) :rule reordering :premises (t782))
(step t784 (cl (not y$1155) (not y$1164)) :rule equiv1 :premises (a753))
(step t785 (cl (not (= (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op y$1155))) (not (or y$v3_1517448506_567_op (not (not y$1155)))) (or y$v3_1517448506_567_op y$1155)) :rule equiv_pos2)
(step t786 (cl (= y$v3_1517448506_567_op y$v3_1517448506_567_op)) :rule refl)
(step t787 (cl (= (not (not y$1155)) y$1155)) :rule rare_rewrite :args ("bool-double-not-elim" y$1155))
(step t788 (cl (= (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op y$1155))) :rule cong :premises (t786 t787))
(step t789 (cl y$v3_1517448506_567_op (not (not y$1155))) :rule equiv2 :premises (a746))
(step t790 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (not y$v3_1517448506_567_op)) :rule or_neg :args (0))
(step t791 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (not (not (not y$1155)))) :rule or_neg :args (1))
(step t792 (cl (or y$v3_1517448506_567_op (not (not y$1155))) (or y$v3_1517448506_567_op (not (not y$1155)))) :rule resolution :premises (t789 t790 t791))
(step t793 (cl (or y$v3_1517448506_567_op (not (not y$1155)))) :rule contraction :premises (t792))
(step t794 (cl (or y$v3_1517448506_567_op y$1155)) :rule resolution :premises (t785 t788 t793))
(step t795 (cl y$v3_1517448506_567_op y$1155) :rule or :premises (t794))
(step t796 (cl y$1155 y$v3_1517448506_567_op) :rule reordering :premises (t795))
(step t797 (cl (not y$v3_1517448506_567_op) (and y$1151 y$1152)) :rule equiv1 :premises (a745))
(step t798 (cl (and y$1151 y$1152) (not y$v3_1517448506_567_op)) :rule reordering :premises (t797))
(step t799 (cl (not (and y$1151 y$1152)) y$1151) :rule and_pos :args (0))
(step t800 (cl y$1151 (not (and y$1151 y$1152))) :rule reordering :premises (t799))
(step t801 (cl (not y$a_too_small_resistance$next_rhs_op) (not y$1151)) :rule equiv1 :premises (a743))
(step t802 (cl (not y$1151) (not y$a_too_small_resistance$next_rhs_op)) :rule reordering :premises (t801))
(step t803 (cl (not y$1151)) :rule resolution :premises (t802 t666))
(step t804 (cl (not (and y$1151 y$1152))) :rule resolution :premises (t800 t803))
(step t805 (cl (not y$v3_1517448506_567_op)) :rule resolution :premises (t798 t804))
(step t806 (cl y$1155) :rule resolution :premises (t796 t805))
(step t807 (cl (not y$1164)) :rule resolution :premises (t784 t806))
(step t808 (cl (not (and y$1163 y$1164))) :rule resolution :premises (t783 t807))
(step t809 (cl (not y$v3_1517448506_572_op)) :rule resolution :premises (t781 t808))
(step t810 (cl y$1167) :rule resolution :premises (t779 t809))
(step t811 (cl (not y$1175)) :rule resolution :premises (t767 t810))
(step t812 (cl (not (and y$293 y$1175))) :rule resolution :premises (t766 t811))
(step t813 (cl (not y$v3_1517448506_577_op)) :rule resolution :premises (t764 t812))
(step t814 (cl y$1178) :rule resolution :premises (t762 t813))
(step t815 (cl (not y$1187)) :rule resolution :premises (t750 t814))
(step t816 (cl (not (and y$1186 y$1187))) :rule resolution :premises (t749 t815))
(step t817 (cl (not y$v3_1517448506_582_op)) :rule resolution :premises (t747 t816))
(step t818 (cl y$1190) :rule resolution :premises (t745 t817))
(step t819 (cl (not y$1199)) :rule resolution :premises (t733 t818))
(step t820 (cl (not (and y$1198 y$1199))) :rule resolution :premises (t732 t819))
(step t821 (cl (not y$v3_1517448506_587_op)) :rule resolution :premises (t730 t820))
(step t822 (cl y$1202) :rule resolution :premises (t728 t821))
(step t823 (cl y$v3_1517448506_588_op (not (and y$a_S2$next_rhs_op y$1202))) :rule equiv2 :premises (a782))
(step t824 (cl (not y$v3_1517448506_588_op) (not y$1205)) :rule equiv1 :premises (a783))
(step t825 (cl (not y$1205) (not y$v3_1517448506_588_op)) :rule reordering :premises (t824))
(step t826 (cl (not (and y$1197 y$1205)) y$1205) :rule and_pos :args (1))
(step t827 (cl y$1205 (not (and y$1197 y$1205))) :rule reordering :premises (t826))
(step t828 (cl y$1205) :rule resolution :premises (t827 t524))
(step t829 (cl (not y$v3_1517448506_588_op)) :rule resolution :premises (t825 t828))
(step t830 (cl (not (and y$a_S2$next_rhs_op y$1202))) :rule resolution :premises (t823 t829))
(step t831 (cl (not y$a_S2$next_rhs_op)) :rule resolution :premises (t716 t822 t830))
(step t832 (cl (not (and y$261 y$v3_1517448506_174_op))) :rule resolution :premises (t714 t831))
(step t833 (cl (not y$v3_1517448506_174_op)) :rule resolution :premises (t689 t713 t832))
(step t834 (cl (not (and y$206 y$322))) :rule resolution :premises (t687 t833))
(step t835 (cl (not y$322)) :rule resolution :premises (t216 t686 t834))
(step t836 (cl y$v3_1517448506_173_op) :rule resolution :premises (t214 t835))
(step t837 (cl (and y$3 y$315)) :rule resolution :premises (t202 t836))
(step t838 (cl y$315) :rule resolution :premises (t200 t837))
(step t839 (cl y$a_S1$next_rhs_op (not (and y$313 y$315))) :rule equiv2 :premises (a168))
(step t840 (cl (and y$a_S1$next_rhs_op y$1190) (not y$a_S1$next_rhs_op) (not y$1190)) :rule and_neg)
(step t841 (cl (not y$1190) (and y$a_S1$next_rhs_op y$1190) (not y$a_S1$next_rhs_op)) :rule reordering :premises (t840))
(step t842 (cl y$v3_1517448506_583_op (not (and y$a_S1$next_rhs_op y$1190))) :rule equiv2 :premises (a773))
(step t843 (cl (not y$v3_1517448506_583_op) (not y$1193)) :rule equiv1 :premises (a774))
(step t844 (cl (not y$1193) (not y$v3_1517448506_583_op)) :rule reordering :premises (t843))
(step t845 (cl (not (and y$1185 y$1193)) y$1193) :rule and_pos :args (1))
(step t846 (cl y$1193 (not (and y$1185 y$1193))) :rule reordering :premises (t845))
(step t847 (cl y$1193) :rule resolution :premises (t846 t528))
(step t848 (cl (not y$v3_1517448506_583_op)) :rule resolution :premises (t844 t847))
(step t849 (cl (not (and y$a_S1$next_rhs_op y$1190))) :rule resolution :premises (t842 t848))
(step t850 (cl (not y$a_S1$next_rhs_op)) :rule resolution :premises (t841 t818 t849))
(step t851 (cl (not (and y$313 y$315))) :rule resolution :premises (t839 t850))
(step t852 (cl (not y$313)) :rule resolution :premises (t198 t838 t851))
(step t853 (cl y$v3_1517448506_168_op) :rule resolution :premises (t196 t852))
(step t854 (cl (and y$1 y$161)) :rule resolution :premises (t184 t853))
(step t855 (cl y$161) :rule resolution :premises (t182 t854))
(step t856 (cl y$a_new_range$next_rhs_op (not (and y$161 y$225))) :rule equiv2 :premises (a107))
(step t857 (cl (not y$a_new_range$next_rhs_op) (not y$1103)) :rule equiv1 :premises (a707))
(step t858 (cl (not y$1103) (not y$a_new_range$next_rhs_op)) :rule reordering :premises (t857))
(step t859 (cl (not (and y$1103 y$1104)) y$1103) :rule and_pos :args (0))
(step t860 (cl y$1103 (not (and y$1103 y$1104))) :rule reordering :premises (t859))
(step t861 (cl y$1103) :rule resolution :premises (t860 t681))
(step t862 (cl (not y$a_new_range$next_rhs_op)) :rule resolution :premises (t858 t861))
(step t863 (cl (not (and y$161 y$225))) :rule resolution :premises (t856 t862))
(step t864 (cl (not y$225)) :rule resolution :premises (t180 t855 t863))
(step t865 (cl y$v3_1517448506_121_op) :rule resolution :premises (t178 t864))
(step t866 (cl (and y$221 y$222)) :rule resolution :premises (t166 t865))
(step t867 (cl y$221) :rule resolution :premises (t164 t866))
(step t868 (cl (not y$220)) :rule resolution :premises (t162 t867))
(step t869 (cl y$v3_1517448506_119_op) :rule resolution :premises (t161 t868))
(step t870 (cl (and y$216 y$217)) :rule resolution :premises (t149 t869))
(step t871 (cl y$216) :rule resolution :premises (t147 t870))
(step t872 (cl (not y$215)) :rule resolution :premises (t145 t871))
(step t873 (cl y$v3_1517448506_117_op) :rule resolution :premises (t144 t872))
(step t874 (cl (and y$25 y$212)) :rule resolution :premises (t132 t873))
(step t875 (cl y$212) :rule resolution :premises (t130 t874))
(step t876 (cl (not y$f09)) :rule resolution :premises (t128 t875))
(step t877 (cl (= y$v_range (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t126 t876))
(step t878 (cl (not (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) (not (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule equiv_pos2)
(step t879 (cl (= y$f10 y$f10)) :rule refl)
(step t880 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f10 y$127 y$v3_1517448506_80_op) y$127))
(step t881 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))
(step t882 (cl (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule cong :premises (t879 t880 t881))
(step t883 (cl (not (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)))) (not (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule equiv_pos2)
(step t884 (cl (= (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule equiv_simplify)
(step t885 (cl (not (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule equiv1 :premises (t884))
(step t886 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule refl)
(step t887 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule refl)
(step t888 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule cong :premises (t886 t887))
(step t889 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f10 y$127 y$v3_1517448506_80_op)))
(step t890 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op)) true)) :rule trans :premises (t888 t889))
(step t891 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t885 t890))
(step t892 (cl (= y$127 y$127)) :rule refl)
(step t893 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127))) :rule cong :premises (t891 t892))
(step t894 (cl (= y$v3_1517448506_80_op y$v3_1517448506_80_op)) :rule refl)
(step t895 (cl (= (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule cong :premises (t891 t894))
(step t896 (cl (= (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)) (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op)))) :rule cong :premises (t879 t893 t895))
(step t897 (cl (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule hole)
(step t898 (cl (ite y$f10 (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$127) (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_80_op))) :rule resolution :premises (t883 t896 t897))
(step t899 (cl (ite y$f10 (= y$127 (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule resolution :premises (t878 t882 t898))
(step t900 (cl y$f10 (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule ite1 :premises (t899))
(step t901 (cl (not y$f10) (not y$217)) :rule equiv1 :premises (a100))
(step t902 (cl (not y$217) (not y$f10)) :rule reordering :premises (t901))
(step t903 (cl (not (and y$216 y$217)) y$217) :rule and_pos :args (1))
(step t904 (cl y$217 (not (and y$216 y$217))) :rule reordering :premises (t903))
(step t905 (cl y$217) :rule resolution :premises (t904 t870))
(step t906 (cl (not y$f10)) :rule resolution :premises (t902 t905))
(step t907 (cl (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t900 t906))
(step t908 (cl (not (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) (not (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule equiv_pos2)
(step t909 (cl (= y$f11 y$f11)) :rule refl)
(step t910 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16))
(step t911 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule rare_rewrite :args ("eq-symm" (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))
(step t912 (cl (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule cong :premises (t909 t910 t911))
(step t913 (cl (not (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)))) (not (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule equiv_pos2)
(step t914 (cl (= (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule equiv_simplify)
(step t915 (cl (not (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule equiv1 :premises (t914))
(step t916 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule refl)
(step t917 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule refl)
(step t918 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule cong :premises (t916 t917))
(step t919 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(step t920 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) true)) :rule trans :premises (t918 t919))
(step t921 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t915 t920))
(step t922 (cl (= y$n3s16 y$n3s16)) :rule refl)
(step t923 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16))) :rule cong :premises (t921 t922))
(step t924 (cl (= y$v3_1517448506_81_op y$v3_1517448506_81_op)) :rule refl)
(step t925 (cl (= (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule cong :premises (t921 t924))
(step t926 (cl (= (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)) (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op)))) :rule cong :premises (t909 t923 t925))
(step t927 (cl (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule hole)
(step t928 (cl (ite y$f11 (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$n3s16) (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v3_1517448506_81_op))) :rule resolution :premises (t913 t926 t927))
(step t929 (cl (ite y$f11 (= y$n3s16 (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule resolution :premises (t908 t912 t928))
(step t930 (cl y$f11 (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule ite1 :premises (t929))
(step t931 (cl (not y$f11) (not y$222)) :rule equiv1 :premises (a104))
(step t932 (cl (not y$222) (not y$f11)) :rule reordering :premises (t931))
(step t933 (cl (not (and y$221 y$222)) y$222) :rule and_pos :args (1))
(step t934 (cl y$222 (not (and y$221 y$222))) :rule reordering :premises (t933))
(step t935 (cl y$222) :rule resolution :premises (t934 t866))
(step t936 (cl (not y$f11)) :rule resolution :premises (t932 t935))
(step t937 (cl (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t930 t936))
(step t938 (cl (not y$140) (= y$v_range$next y$v_range$next_rhs_op)) :rule equiv1 :premises (a47))
(step t939 (cl (= y$v_range$next y$v_range$next_rhs_op) (not y$140)) :rule reordering :premises (t938))
(step t940 (cl (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271)) y$140) :rule and_pos :args (2))
(step t941 (cl y$140 (not (and y$64 y$106 y$140 y$164 y$176 y$188 y$193 y$210 y$228 y$239 y$248 y$265 y$274 y$283 y$294 y$309 y$318 y$327 y$341 y$353 y$364 y$377 y$384 y$1271))) :rule reordering :premises (t940))
(step t942 (cl y$140) :rule resolution :premises (t941 t505))
(step t943 (cl (= y$v_range$next y$v_range$next_rhs_op)) :rule resolution :premises (t939 t942))
(step t944 (cl (not (= (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule equiv_pos2)
(step t945 (cl (= y$v_range$next_rhs_op y$v_range$next_rhs_op)) :rule refl)
(step t946 (cl (= (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule cong :premises (t945 t921))
(step t947 (cl (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule resolution :premises (t944 t946 a46))
(step t948 (cl (not (= (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule equiv_pos2)
(step t949 (cl (= (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))) :rule cong :premises (t924 t891))
(step t950 (cl (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule resolution :premises (t948 t949 a45))
(step t951 (cl (not (= (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule equiv_pos2)
(step t952 (cl (= (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))) :rule cong :premises (t894 t117))
(step t953 (cl (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule resolution :premises (t951 t952 a44))
(step t954 (cl (not y$111) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule ite2 :premises (t69))
(step t955 (cl (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (not y$111)) :rule reordering :premises (t954))
(step t956 (cl y$111 (not (Extract_1_15_15_16 y$v_range))) :rule equiv2 :premises (a3))
(step t957 (cl (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))) (not (Extract_1_15_15_16 y$v_range$next)) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule and_neg)
(step t958 (cl (=> (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))) (Extract_1_15_15_16 y$v_range)) (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) :rule implies_neg1)
(anchor :step t959)
(assume t959.a0 (Extract_1_15_15_16 y$v_range$next))
(assume t959.a1 (= y$v_range$next y$v_range$next_rhs_op))
(assume t959.a2 (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t959.a3 (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t959.a4 (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t959.a5 (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t959.a6 (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))
(assume t959.a7 (= y$v_range (ite y$f09 y$132 y$v_range)))
(step t959.t0 (cl (= (= (Extract_1_15_15_16 y$v_range) true) (Extract_1_15_15_16 y$v_range))) :rule equiv_simplify)
(step t959.t1 (cl (not (= (Extract_1_15_15_16 y$v_range) true)) (Extract_1_15_15_16 y$v_range)) :rule equiv1 :premises (t959.t0))
(step t959.t2 (cl (= (ite y$f09 y$132 y$v_range) y$v3_1517448506_80_op)) :rule symm :premises (t953))
(step t959.t3 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_81_op)) :rule symm :premises (t950))
(step t959.t4 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v_range$next_rhs_op)) :rule symm :premises (t947))
(step t959.t5 (cl (= y$v_range$next_rhs_op y$v_range$next)) :rule symm :premises (t943))
(step t959.t6 (cl (= y$v_range y$v_range$next)) :rule trans :premises (t877 t959.t2 t907 t959.t3 t937 t959.t4 t959.t5))
(step t959.t7 (cl (= (Extract_1_15_15_16 y$v_range) (Extract_1_15_15_16 y$v_range$next))) :rule cong :premises (t959.t6))
(step t959.t8 (cl (= (= (Extract_1_15_15_16 y$v_range$next) true) (Extract_1_15_15_16 y$v_range$next))) :rule equiv_simplify)
(step t959.t9 (cl (= (Extract_1_15_15_16 y$v_range$next) true) (not (Extract_1_15_15_16 y$v_range$next))) :rule equiv2 :premises (t959.t8))
(step t959.t10 (cl (= (Extract_1_15_15_16 y$v_range$next) true)) :rule resolution :premises (t959.t9 t959.a0))
(step t959.t11 (cl (= (Extract_1_15_15_16 y$v_range) true)) :rule trans :premises (t959.t7 t959.t10))
(step t959.t12 (cl (Extract_1_15_15_16 y$v_range)) :rule resolution :premises (t959.t1 t959.t11))
(step t959 (cl (not (Extract_1_15_15_16 y$v_range$next)) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (Extract_1_15_15_16 y$v_range)) :rule subproof :discharge (t959.a0 t959.a1 t959.a2 t959.a3 t959.a4 t959.a5 t959.a6 t959.a7))
(step t960 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (Extract_1_15_15_16 y$v_range$next)) :rule and_pos :args (0))
(step t961 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (= y$v_range$next y$v_range$next_rhs_op)) :rule and_pos :args (1))
(step t962 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (2))
(step t963 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (3))
(step t964 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (4))
(step t965 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (5))
(step t966 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule and_pos :args (6))
(step t967 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (= y$v_range (ite y$f09 y$132 y$v_range))) :rule and_pos :args (7))
(step t968 (cl (Extract_1_15_15_16 y$v_range) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))))) :rule resolution :premises (t959 t960 t961 t962 t963 t964 t965 t966 t967))
(step t969 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (Extract_1_15_15_16 y$v_range)) :rule reordering :premises (t968))
(step t970 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (Extract_1_15_15_16 y$v_range)) :rule contraction :premises (t969))
(step t971 (cl (=> (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))) (Extract_1_15_15_16 y$v_range)) (Extract_1_15_15_16 y$v_range)) :rule resolution :premises (t958 t970))
(step t972 (cl (=> (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))) (Extract_1_15_15_16 y$v_range)) (not (Extract_1_15_15_16 y$v_range))) :rule implies_neg2)
(step t973 (cl (=> (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))) (Extract_1_15_15_16 y$v_range)) (=> (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))) (Extract_1_15_15_16 y$v_range))) :rule resolution :premises (t971 t972))
(step t974 (cl (=> (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range))) (Extract_1_15_15_16 y$v_range))) :rule contraction :premises (t973))
(step t975 (cl (not (and (Extract_1_15_15_16 y$v_range$next) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)))) (Extract_1_15_15_16 y$v_range)) :rule implies :premises (t974))
(step t976 (cl (not (Extract_1_15_15_16 y$v_range$next)) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (Extract_1_15_15_16 y$v_range)) :rule resolution :premises (t957 t975))
(step t977 (cl (Extract_1_15_15_16 y$v_range) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (Extract_1_15_15_16 y$v_range$next)) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))) :rule reordering :premises (t976))
(step t978 (cl (not y$1320) (Extract_1_15_15_16 y$v_range$next)) :rule equiv1 :premises (a833))
(step t979 (cl (Extract_1_15_15_16 y$v_range$next) (not y$1320)) :rule reordering :premises (t978))
(step t980 (cl y$1320 (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule ite1 :premises (t21))
(step t981 (cl (not (= (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))))) (not (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op)))) (or (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule equiv_pos2)
(step t982 (cl (= (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (= y$n0s32 y$v3_1517448506_30$next_op))) :rule rare_rewrite :args ("bool-double-not-elim" (= y$n0s32 y$v3_1517448506_30$next_op)))
(step t983 (cl (= (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule refl)
(step t984 (cl (= (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule refl)
(step t985 (cl (= (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))))) :rule refl)
(step t986 (cl (= (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))))) :rule refl)
(step t987 (cl (= (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next y$v_range$next_rhs_op)))) :rule refl)
(step t988 (cl (= (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule refl)
(step t989 (cl (= (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule refl)
(step t990 (cl (= (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule refl)
(step t991 (cl (= (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule refl)
(step t992 (cl (= (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))))) :rule refl)
(step t993 (cl (= (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))))) :rule refl)
(step t994 (cl (= (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))))) :rule refl)
(step t995 (cl (= (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))))) :rule refl)
(step t996 (cl (= (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule refl)
(step t997 (cl (= (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule refl)
(step t998 (cl (= (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule refl)
(step t999 (cl (= (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))))) :rule cong :premises (t982 t983 t984 t985 t986 t987 t988 t989 t990 t991 t992 t993 t994 t995 t996 t997 t998))
(step t1000 (cl (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule and_neg)
(step t1001 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule implies_neg1)
(anchor :step t1002)
(assume t1002.a0 (not (= y$n0s32 y$v3_1517448506_30$next_op)))
(assume t1002.a1 (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))
(assume t1002.a2 (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))
(assume t1002.a3 (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)))
(assume t1002.a4 (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)))
(assume t1002.a5 (= y$v_range$next y$v_range$next_rhs_op))
(assume t1002.a6 (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t1002.a7 (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t1002.a8 (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t1002.a9 (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t1002.a10 (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))
(assume t1002.a11 (= y$v_range (ite y$f09 y$132 y$v_range)))
(assume t1002.a12 (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)))
(assume t1002.a13 (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)))
(assume t1002.a14 (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))
(assume t1002.a15 (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))
(step t1002.t0 (cl (= (= (= y$n0s32 y$v3_1517448506_30_op) false) (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule equiv_simplify)
(step t1002.t1 (cl (not (= (= y$n0s32 y$v3_1517448506_30_op) false)) (not (= y$n0s32 y$v3_1517448506_30_op))) :rule equiv1 :premises (t1002.t0))
(step t1002.t2 (cl (not (= (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule equiv_pos2)
(step t1002.t3 (cl (= y$v3_1517448506_30_op y$v3_1517448506_30_op)) :rule refl)
(step t1002.t4 (cl (= (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule cong :premises (t1002.t3 t61))
(step t1002.t5 (cl (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule resolution :premises (t1002.t2 t1002.t4 t1002.a15))
(step t1002.t6 (cl (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_31_op)) :rule symm :premises (t1002.a14))
(step t1002.t7 (cl (= y$n0s16 y$n0s16)) :rule refl)
(step t1002.t8 (cl (= (ite y$f09 y$132 y$v_range) y$v3_1517448506_80_op)) :rule symm :premises (t953))
(step t1002.t9 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_81_op)) :rule symm :premises (t950))
(step t1002.t10 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v_range$next_rhs_op)) :rule symm :premises (t947))
(step t1002.t11 (cl (= y$v_range$next_rhs_op y$v_range$next)) :rule symm :premises (t943))
(step t1002.t12 (cl (= y$v_range y$v_range$next)) :rule trans :premises (t877 t1002.t8 t907 t1002.t9 t937 t1002.t10 t1002.t11))
(step t1002.t13 (cl (= (Concat_32_16_16 y$n0s16 y$v_range) (Concat_32_16_16 y$n0s16 y$v_range$next))) :rule cong :premises (t1002.t7 t1002.t12))
(step t1002.t14 (cl (= (Concat_32_16_16 y$n0s16 y$v_range$next) y$w$6$next_op)) :rule symm :premises (t1002.a4))
(step t1002.t15 (cl (= y$w$6_op y$w$6$next_op)) :rule trans :premises (t1002.a12 t1002.t13 t1002.t14))
(step t1002.t16 (cl (= y$n16s32 y$n16s32)) :rule refl)
(step t1002.t17 (cl (= (ShiftR_32_32_32 y$w$6_op y$n16s32) (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) :rule cong :premises (t1002.t15 t1002.t16))
(step t1002.t18 (cl (= (ShiftR_32_32_32 y$w$6$next_op y$n16s32) y$v3_1517448506_31$next_op)) :rule symm :premises (t1002.a3))
(step t1002.t19 (cl (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_31$next_op)) :rule symm :premises (t1002.a2))
(step t1002.t20 (cl (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule symm :premises (t1002.t19))
(step t1002.t21 (cl (not (= (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule equiv_pos2)
(step t1002.t22 (cl (= y$v3_1517448506_30$next_op y$v3_1517448506_30$next_op)) :rule refl)
(step t1002.t23 (cl (= (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule cong :premises (t1002.t22 t13))
(step t1002.t24 (cl (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule resolution :premises (t1002.t21 t1002.t23 t1002.a1))
(step t1002.t25 (cl (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_30$next_op)) :rule symm :premises (t1002.t24))
(step t1002.t26 (cl (= y$v3_1517448506_30_op y$v3_1517448506_30$next_op)) :rule trans :premises (t1002.t5 t1002.t6 t1002.a13 t1002.t17 t1002.t18 t1002.t20 t1002.t25))
(step t1002.t27 (cl (= (= y$n0s32 y$v3_1517448506_30_op) (= y$n0s32 y$v3_1517448506_30$next_op))) :rule cong :premises (t566 t1002.t26))
(step t1002.t28 (cl (= (= (= y$n0s32 y$v3_1517448506_30$next_op) false) (not (= y$n0s32 y$v3_1517448506_30$next_op)))) :rule equiv_simplify)
(step t1002.t29 (cl (= (= y$n0s32 y$v3_1517448506_30$next_op) false) (not (not (= y$n0s32 y$v3_1517448506_30$next_op)))) :rule equiv2 :premises (t1002.t28))
(step t1002.t30 (cl (not (not (not (= y$n0s32 y$v3_1517448506_30$next_op)))) (= y$n0s32 y$v3_1517448506_30$next_op)) :rule not_not)
(step t1002.t31 (cl (= (= y$n0s32 y$v3_1517448506_30$next_op) false) (= y$n0s32 y$v3_1517448506_30$next_op)) :rule resolution :premises (t1002.t29 t1002.t30))
(step t1002.t32 (cl (not y$1334) (not (= y$n0s32 y$v3_1517448506_30$next_op))) :rule equiv1 :premises (a840))
(step t1002.t33 (cl (not (= y$n0s32 y$v3_1517448506_30$next_op)) (not y$1334)) :rule reordering :premises (t1002.t32))
(step t1002.t34 (cl (not (and y$1334 y$1337)) y$1334) :rule and_pos :args (0))
(step t1002.t35 (cl y$1334 (not (and y$1334 y$1337))) :rule reordering :premises (t1002.t34))
(step t1002.t36 (cl y$1334) :rule resolution :premises (t1002.t35 t589))
(step t1002.t37 (cl (not (= y$n0s32 y$v3_1517448506_30$next_op))) :rule resolution :premises (t1002.t33 t1002.t36))
(step t1002.t38 (cl (= (= y$n0s32 y$v3_1517448506_30$next_op) false)) :rule resolution :premises (t1002.t31 t1002.t37))
(step t1002.t39 (cl (= (= y$n0s32 y$v3_1517448506_30_op) false)) :rule trans :premises (t1002.t27 t1002.t38))
(step t1002.t40 (cl (not (= y$n0s32 y$v3_1517448506_30_op))) :rule resolution :premises (t1002.t1 t1002.t39))
(step t1002 (cl (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) :rule subproof :discharge (t1002.a0 t1002.a1 t1002.a2 t1002.a3 t1002.a4 t1002.a5 t1002.a6 t1002.a7 t1002.a8 t1002.a9 t1002.a10 t1002.a11 t1002.a12 t1002.a13 t1002.a14 t1002.a15))
(step t1003 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (= y$n0s32 y$v3_1517448506_30$next_op))) :rule and_pos :args (0))
(step t1004 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule and_pos :args (1))
(step t1005 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule and_pos :args (2))
(step t1006 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) :rule and_pos :args (3))
(step t1007 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) :rule and_pos :args (4))
(step t1008 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v_range$next y$v_range$next_rhs_op)) :rule and_pos :args (5))
(step t1009 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (6))
(step t1010 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (7))
(step t1011 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (8))
(step t1012 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (9))
(step t1013 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule and_pos :args (10))
(step t1014 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v_range (ite y$f09 y$132 y$v_range))) :rule and_pos :args (11))
(step t1015 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) :rule and_pos :args (12))
(step t1016 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) :rule and_pos :args (13))
(step t1017 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule and_pos :args (14))
(step t1018 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule and_pos :args (15))
(step t1019 (cl (not (= y$n0s32 y$v3_1517448506_30_op)) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule resolution :premises (t1002 t1003 t1004 t1005 t1006 t1007 t1008 t1009 t1010 t1011 t1012 t1013 t1014 t1015 t1016 t1017 t1018))
(step t1020 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (= y$n0s32 y$v3_1517448506_30_op))) :rule reordering :premises (t1019))
(step t1021 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (= y$n0s32 y$v3_1517448506_30_op))) :rule contraction :premises (t1020))
(step t1022 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) :rule resolution :premises (t1001 t1021))
(step t1023 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule implies_neg2)
(step t1024 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule resolution :premises (t1022 t1023))
(step t1025 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule contraction :premises (t1024))
(step t1026 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32)) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$v_range$next y$v_range$next_rhs_op) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32)) (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (= y$n0s32 y$v3_1517448506_30_op))) :rule implies :premises (t1025))
(step t1027 (cl (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) :rule resolution :premises (t1000 t1026))
(step t1028 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (not (= y$n0s32 y$v3_1517448506_30$next_op))))) :rule or_neg :args (0))
(step t1029 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule or_neg :args (1))
(step t1030 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule or_neg :args (2))
(step t1031 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))))) :rule or_neg :args (3))
(step t1032 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))))) :rule or_neg :args (4))
(step t1033 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v_range$next y$v_range$next_rhs_op)))) :rule or_neg :args (5))
(step t1034 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule or_neg :args (6))
(step t1035 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule or_neg :args (7))
(step t1036 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule or_neg :args (8))
(step t1037 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule or_neg :args (9))
(step t1038 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))))) :rule or_neg :args (10))
(step t1039 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v_range (ite y$f09 y$132 y$v_range))))) :rule or_neg :args (11))
(step t1040 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))))) :rule or_neg :args (12))
(step t1041 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))))) :rule or_neg :args (13))
(step t1042 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule or_neg :args (14))
(step t1043 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule or_neg :args (15))
(step t1044 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (not (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule or_neg :args (16))
(step t1045 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule resolution :premises (t1027 t1028 t1029 t1030 t1031 t1032 t1033 t1034 t1035 t1036 t1037 t1038 t1039 t1040 t1041 t1042 t1043 t1044))
(step t1046 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule contraction :premises (t1045))
(step t1047 (cl (or (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule resolution :premises (t981 t999 t1046))
(step t1048 (cl (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$n0s32 y$v3_1517448506_30_op))) :rule or :premises (t1047))
(step t1049 (cl (not (= y$n0s32 y$v3_1517448506_30_op)) (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v3_1517448506_31_op (ShiftR_32_32_32 y$w$6_op y$n16s32))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$v3_1517448506_31$next_op (ShiftR_32_32_32 y$w$6$next_op y$n16s32))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v3_1517448506_31$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule reordering :premises (t1048))
(step t1050 (cl (not (= (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (or (= y$n1s32 y$v3_1517448506_36_op) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))))) (not (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) (or (= y$n1s32 y$v3_1517448506_36_op) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule equiv_pos2)
(step t1051 (cl (= (not (not (= y$n1s32 y$v3_1517448506_36_op))) (= y$n1s32 y$v3_1517448506_36_op))) :rule rare_rewrite :args ("bool-double-not-elim" (= y$n1s32 y$v3_1517448506_36_op)))
(step t1052 (cl (= (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule refl)
(step t1053 (cl (= (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule refl)
(step t1054 (cl (= (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (or (= y$n1s32 y$v3_1517448506_36_op) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))))) :rule cong :premises (t1051 t1052 t1053))
(step t1055 (cl (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule and_neg)
(step t1056 (cl (=> (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule implies_neg1)
(anchor :step t1057)
(assume t1057.a0 (not (= y$n1s32 y$v3_1517448506_36_op)))
(assume t1057.a1 (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))
(step t1057.t0 (cl (= (= (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) false) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule equiv_simplify)
(step t1057.t1 (cl (not (= (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) false)) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule equiv1 :premises (t1057.t0))
(step t1057.t2 (cl (not (= (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule equiv_pos2)
(step t1057.t3 (cl (= y$v3_1517448506_36_op y$v3_1517448506_36_op)) :rule refl)
(step t1057.t4 (cl (= (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) true) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule equiv_simplify)
(step t1057.t5 (cl (not (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) true)) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule equiv1 :premises (t1057.t4))
(step t1057.t6 (cl (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule refl)
(step t1057.t7 (cl (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule refl)
(step t1057.t8 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule cong :premises (t1057.t6 t1057.t7))
(step t1057.t9 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$a_too_small_resistance y$n1s32 y$n0s32)))
(step t1057.t10 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) true)) :rule trans :premises (t1057.t8 t1057.t9))
(step t1057.t11 (cl (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule resolution :premises (t1057.t5 t1057.t10))
(step t1057.t12 (cl (= (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule cong :premises (t1057.t3 t1057.t11))
(step t1057.t13 (cl (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule resolution :premises (t1057.t2 t1057.t12 t1057.a1))
(step t1057.t14 (cl (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$v3_1517448506_36_op)) :rule symm :premises (t1057.t13))
(step t1057.t15 (cl (= (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$n1s32 y$v3_1517448506_36_op))) :rule cong :premises (t564 t1057.t14))
(step t1057.t16 (cl (= (= (= y$n1s32 y$v3_1517448506_36_op) false) (not (= y$n1s32 y$v3_1517448506_36_op)))) :rule equiv_simplify)
(step t1057.t17 (cl (= (= y$n1s32 y$v3_1517448506_36_op) false) (not (not (= y$n1s32 y$v3_1517448506_36_op)))) :rule equiv2 :premises (t1057.t16))
(step t1057.t18 (cl (not (not (not (= y$n1s32 y$v3_1517448506_36_op)))) (= y$n1s32 y$v3_1517448506_36_op)) :rule not_not)
(step t1057.t19 (cl (= (= y$n1s32 y$v3_1517448506_36_op) false) (= y$n1s32 y$v3_1517448506_36_op)) :rule resolution :premises (t1057.t17 t1057.t18))
(step t1057.t20 (cl (= (= y$n1s32 y$v3_1517448506_36_op) false)) :rule resolution :premises (t1057.t19 t1057.a0))
(step t1057.t21 (cl (= (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) false)) :rule trans :premises (t1057.t15 t1057.t20))
(step t1057.t22 (cl (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule resolution :premises (t1057.t1 t1057.t21))
(step t1057 (cl (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule subproof :discharge (t1057.a0 t1057.a1))
(step t1058 (cl (not (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (= y$n1s32 y$v3_1517448506_36_op))) :rule and_pos :args (0))
(step t1059 (cl (not (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule and_pos :args (1))
(step t1060 (cl (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule resolution :premises (t1057 t1058 t1059))
(step t1061 (cl (not (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule reordering :premises (t1060))
(step t1062 (cl (not (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule contraction :premises (t1061))
(step t1063 (cl (=> (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule resolution :premises (t1056 t1062))
(step t1064 (cl (=> (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule implies_neg2)
(step t1065 (cl (=> (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (=> (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule resolution :premises (t1063 t1064))
(step t1066 (cl (=> (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule contraction :premises (t1065))
(step t1067 (cl (not (and (not (= y$n1s32 y$v3_1517448506_36_op)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule implies :premises (t1066))
(step t1068 (cl (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule resolution :premises (t1055 t1067))
(step t1069 (cl (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (not (not (= y$n1s32 y$v3_1517448506_36_op))))) :rule or_neg :args (0))
(step t1070 (cl (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule or_neg :args (1))
(step t1071 (cl (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule or_neg :args (2))
(step t1072 (cl (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule resolution :premises (t1068 t1069 t1070 t1071))
(step t1073 (cl (or (not (not (= y$n1s32 y$v3_1517448506_36_op))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule contraction :premises (t1072))
(step t1074 (cl (or (= y$n1s32 y$v3_1517448506_36_op) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule resolution :premises (t1050 t1054 t1073))
(step t1075 (cl (= y$n1s32 y$v3_1517448506_36_op) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (not (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule or :premises (t1074))
(step t1076 (cl (not (= (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) (not (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule equiv_pos2)
(step t1077 (cl (= y$v3_1517448506_36_op y$v3_1517448506_36_op)) :rule refl)
(step t1078 (cl (= (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) true) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule equiv_simplify)
(step t1079 (cl (not (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) true)) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule equiv1 :premises (t1078))
(step t1080 (cl (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule refl)
(step t1081 (cl (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule refl)
(step t1082 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule cong :premises (t1080 t1081))
(step t1083 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) true)) :rule rare_rewrite :args ("eq-refl" (ite y$a_too_small_resistance y$n1s32 y$n0s32)))
(step t1084 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32)) true)) :rule trans :premises (t1082 t1083))
(step t1085 (cl (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule resolution :premises (t1079 t1084))
(step t1086 (cl (= (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule cong :premises (t1077 t1085))
(step t1087 (cl (= y$v3_1517448506_36_op (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule resolution :premises (t1076 t1086 a11))
(step t1088 (cl (not (= (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) (not (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32))) (ite y$a_too_small_resistance (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule equiv_pos2)
(step t1089 (cl (= y$a_too_small_resistance y$a_too_small_resistance)) :rule refl)
(step t1090 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule rare_rewrite :args ("eq-symm" (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32))
(step t1091 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32) (= y$n0s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule rare_rewrite :args ("eq-symm" (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32))
(step t1092 (cl (= (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))))) :rule cong :premises (t1089 t1090 t1091))
(step t1093 (cl (not (= (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32)))) (not (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32))) (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32))) :rule equiv_pos2)
(step t1094 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32))) :rule cong :premises (t1085 t564))
(step t1095 (cl (= (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32))) :rule cong :premises (t1085 t566))
(step t1096 (cl (= (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32)) (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32)))) :rule cong :premises (t1089 t1094 t1095))
(step t1097 (cl (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32))) :rule hole)
(step t1098 (cl (ite y$a_too_small_resistance (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n1s32) (= (ite y$a_too_small_resistance y$n1s32 y$n0s32) y$n0s32))) :rule resolution :premises (t1093 t1096 t1097))
(step t1099 (cl (ite y$a_too_small_resistance (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (= y$n0s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)))) :rule resolution :premises (t1088 t1092 t1098))
(step t1100 (cl (not y$a_too_small_resistance) (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32))) :rule ite2 :premises (t1099))
(step t1101 (cl (= y$n1s32 (ite y$a_too_small_resistance y$n1s32 y$n0s32)) (not y$a_too_small_resistance)) :rule reordering :premises (t1100))
(step t1102 (cl (not (= (or y$a_too_small_resistance (not (not y$39))) (or y$a_too_small_resistance y$39))) (not (or y$a_too_small_resistance (not (not y$39)))) (or y$a_too_small_resistance y$39)) :rule equiv_pos2)
(step t1103 (cl (= (not (not y$39)) y$39)) :rule rare_rewrite :args ("bool-double-not-elim" y$39))
(step t1104 (cl (= (or y$a_too_small_resistance (not (not y$39))) (or y$a_too_small_resistance y$39))) :rule cong :premises (t1089 t1103))
(step t1105 (cl y$a_too_small_resistance (not (not y$39))) :rule equiv2 :premises (a133))
(step t1106 (cl (or y$a_too_small_resistance (not (not y$39))) (not y$a_too_small_resistance)) :rule or_neg :args (0))
(step t1107 (cl (or y$a_too_small_resistance (not (not y$39))) (not (not (not y$39)))) :rule or_neg :args (1))
(step t1108 (cl (or y$a_too_small_resistance (not (not y$39))) (or y$a_too_small_resistance (not (not y$39)))) :rule resolution :premises (t1105 t1106 t1107))
(step t1109 (cl (or y$a_too_small_resistance (not (not y$39)))) :rule contraction :premises (t1108))
(step t1110 (cl (or y$a_too_small_resistance y$39)) :rule resolution :premises (t1102 t1104 t1109))
(step t1111 (cl y$a_too_small_resistance y$39) :rule or :premises (t1110))
(step t1112 (cl y$39 y$a_too_small_resistance) :rule reordering :premises (t1111))
(step t1113 (cl (and y$39 y$236) (not y$39) (not y$236)) :rule and_neg)
(step t1114 (cl (not y$236) (not y$39) (and y$39 y$236)) :rule reordering :premises (t1113))
(step t1115 (cl y$v3_1517448506_145_op (not (and y$39 y$236))) :rule equiv2 :premises (a134))
(step t1116 (cl (not y$v3_1517448506_145_op) (not y$269)) :rule equiv1 :premises (a135))
(step t1117 (cl (not y$269) (not y$v3_1517448506_145_op)) :rule reordering :premises (t1116))
(step t1118 (cl (not (and y$269 y$271)) y$269) :rule and_pos :args (0))
(step t1119 (cl y$269 (not (and y$269 y$271))) :rule reordering :premises (t1118))
(step t1120 (cl (not y$a_too_small_resistance$next_rhs_op) (and y$269 y$271)) :rule equiv1 :premises (a137))
(step t1121 (cl (and y$269 y$271) (not y$a_too_small_resistance$next_rhs_op)) :rule reordering :premises (t1120))
(step t1122 (cl (and y$269 y$271)) :rule resolution :premises (t1121 t666))
(step t1123 (cl y$269) :rule resolution :premises (t1119 t1122))
(step t1124 (cl (not y$v3_1517448506_145_op)) :rule resolution :premises (t1117 t1123))
(step t1125 (cl (not (and y$39 y$236))) :rule resolution :premises (t1115 t1124))
(step t1126 (cl (not (= (or y$f12 (not (not y$236))) (or y$f12 y$236))) (not (or y$f12 (not (not y$236)))) (or y$f12 y$236)) :rule equiv_pos2)
(step t1127 (cl (= y$f12 y$f12)) :rule refl)
(step t1128 (cl (= (not (not y$236)) y$236)) :rule rare_rewrite :args ("bool-double-not-elim" y$236))
(step t1129 (cl (= (or y$f12 (not (not y$236))) (or y$f12 y$236))) :rule cong :premises (t1127 t1128))
(step t1130 (cl y$f12 (not (not y$236))) :rule equiv2 :premises (a113))
(step t1131 (cl (or y$f12 (not (not y$236))) (not y$f12)) :rule or_neg :args (0))
(step t1132 (cl (or y$f12 (not (not y$236))) (not (not (not y$236)))) :rule or_neg :args (1))
(step t1133 (cl (or y$f12 (not (not y$236))) (or y$f12 (not (not y$236)))) :rule resolution :premises (t1130 t1131 t1132))
(step t1134 (cl (or y$f12 (not (not y$236)))) :rule contraction :premises (t1133))
(step t1135 (cl (or y$f12 y$236)) :rule resolution :premises (t1126 t1129 t1134))
(step t1136 (cl y$f12 y$236) :rule or :premises (t1135))
(step t1137 (cl y$236 y$f12) :rule reordering :premises (t1136))
(step t1138 (cl (and y$f12 y$519) (not y$f12) (not y$519)) :rule and_neg)
(step t1139 (cl (not y$519) (and y$f12 y$519) (not y$f12)) :rule reordering :premises (t1138))
(step t1140 (cl y$v3_1517448506_273_op (not (and y$f12 y$519))) :rule equiv2 :premises (a295))
(step t1141 (cl (not y$v3_1517448506_273_op) (not y$522)) :rule equiv1 :premises (a296))
(step t1142 (cl (not y$522) (not y$v3_1517448506_273_op)) :rule reordering :premises (t1141))
(step t1143 (cl (not (and y$v3_1517448506_270_op y$522)) y$522) :rule and_pos :args (1))
(step t1144 (cl y$522 (not (and y$v3_1517448506_270_op y$522))) :rule reordering :premises (t1143))
(step t1145 (cl (not y$v3_1517448506_274_op) (and y$v3_1517448506_270_op y$522)) :rule equiv1 :premises (a297))
(step t1146 (cl (and y$v3_1517448506_270_op y$522) (not y$v3_1517448506_274_op)) :rule reordering :premises (t1145))
(step t1147 (cl (not (and y$v3_1517448506_274_op y$531)) y$v3_1517448506_274_op) :rule and_pos :args (0))
(step t1148 (cl y$v3_1517448506_274_op (not (and y$v3_1517448506_274_op y$531))) :rule reordering :premises (t1147))
(step t1149 (cl (not y$v3_1517448506_279_op) (and y$v3_1517448506_274_op y$531)) :rule equiv1 :premises (a303))
(step t1150 (cl (and y$v3_1517448506_274_op y$531) (not y$v3_1517448506_279_op)) :rule reordering :premises (t1149))
(step t1151 (cl (not (and y$v3_1517448506_279_op y$536)) y$v3_1517448506_279_op) :rule and_pos :args (0))
(step t1152 (cl y$v3_1517448506_279_op (not (and y$v3_1517448506_279_op y$536))) :rule reordering :premises (t1151))
(step t1153 (cl (not y$v3_1517448506_282_op) (and y$v3_1517448506_279_op y$536)) :rule equiv1 :premises (a306))
(step t1154 (cl (and y$v3_1517448506_279_op y$536) (not y$v3_1517448506_282_op)) :rule reordering :premises (t1153))
(step t1155 (cl (not (and y$v3_1517448506_282_op y$541)) y$v3_1517448506_282_op) :rule and_pos :args (0))
(step t1156 (cl y$v3_1517448506_282_op (not (and y$v3_1517448506_282_op y$541))) :rule reordering :premises (t1155))
(step t1157 (cl (not y$v3_1517448506_285_op) (and y$v3_1517448506_282_op y$541)) :rule equiv1 :premises (a309))
(step t1158 (cl (and y$v3_1517448506_282_op y$541) (not y$v3_1517448506_285_op)) :rule reordering :premises (t1157))
(step t1159 (cl (not (and y$v3_1517448506_285_op y$546)) y$v3_1517448506_285_op) :rule and_pos :args (0))
(step t1160 (cl y$v3_1517448506_285_op (not (and y$v3_1517448506_285_op y$546))) :rule reordering :premises (t1159))
(step t1161 (cl (not y$v3_1517448506_288_op) (and y$v3_1517448506_285_op y$546)) :rule equiv1 :premises (a312))
(step t1162 (cl (and y$v3_1517448506_285_op y$546) (not y$v3_1517448506_288_op)) :rule reordering :premises (t1161))
(step t1163 (cl (not (and y$v3_1517448506_288_op y$551)) y$v3_1517448506_288_op) :rule and_pos :args (0))
(step t1164 (cl y$v3_1517448506_288_op (not (and y$v3_1517448506_288_op y$551))) :rule reordering :premises (t1163))
(step t1165 (cl (not y$v3_1517448506_291_op) (and y$v3_1517448506_288_op y$551)) :rule equiv1 :premises (a315))
(step t1166 (cl (and y$v3_1517448506_288_op y$551) (not y$v3_1517448506_291_op)) :rule reordering :premises (t1165))
(step t1167 (cl (not (and y$v3_1517448506_291_op y$556)) y$v3_1517448506_291_op) :rule and_pos :args (0))
(step t1168 (cl y$v3_1517448506_291_op (not (and y$v3_1517448506_291_op y$556))) :rule reordering :premises (t1167))
(step t1169 (cl (not y$v3_1517448506_294_op) (and y$v3_1517448506_291_op y$556)) :rule equiv1 :premises (a319))
(step t1170 (cl (and y$v3_1517448506_291_op y$556) (not y$v3_1517448506_294_op)) :rule reordering :premises (t1169))
(step t1171 (cl (not (and y$v3_1517448506_294_op y$561)) y$v3_1517448506_294_op) :rule and_pos :args (0))
(step t1172 (cl y$v3_1517448506_294_op (not (and y$v3_1517448506_294_op y$561))) :rule reordering :premises (t1171))
(step t1173 (cl (not y$v3_1517448506_297_op) (and y$v3_1517448506_294_op y$561)) :rule equiv1 :premises (a322))
(step t1174 (cl (and y$v3_1517448506_294_op y$561) (not y$v3_1517448506_297_op)) :rule reordering :premises (t1173))
(step t1175 (cl (not (and y$v3_1517448506_297_op y$566)) y$v3_1517448506_297_op) :rule and_pos :args (0))
(step t1176 (cl y$v3_1517448506_297_op (not (and y$v3_1517448506_297_op y$566))) :rule reordering :premises (t1175))
(step t1177 (cl (not y$v3_1517448506_300_op) (and y$v3_1517448506_297_op y$566)) :rule equiv1 :premises (a325))
(step t1178 (cl (and y$v3_1517448506_297_op y$566) (not y$v3_1517448506_300_op)) :rule reordering :premises (t1177))
(step t1179 (cl (not (and y$v3_1517448506_300_op y$571)) y$v3_1517448506_300_op) :rule and_pos :args (0))
(step t1180 (cl y$v3_1517448506_300_op (not (and y$v3_1517448506_300_op y$571))) :rule reordering :premises (t1179))
(step t1181 (cl (not y$v3_1517448506_303_op) (and y$v3_1517448506_300_op y$571)) :rule equiv1 :premises (a328))
(step t1182 (cl (and y$v3_1517448506_300_op y$571) (not y$v3_1517448506_303_op)) :rule reordering :premises (t1181))
(step t1183 (cl (not (and y$v3_1517448506_303_op y$581)) y$v3_1517448506_303_op) :rule and_pos :args (0))
(step t1184 (cl y$v3_1517448506_303_op (not (and y$v3_1517448506_303_op y$581))) :rule reordering :premises (t1183))
(step t1185 (cl (not y$v3_1517448506_308_op) (and y$v3_1517448506_303_op y$581)) :rule equiv1 :premises (a334))
(step t1186 (cl (and y$v3_1517448506_303_op y$581) (not y$v3_1517448506_308_op)) :rule reordering :premises (t1185))
(step t1187 (cl (not (and y$v3_1517448506_308_op y$589)) y$v3_1517448506_308_op) :rule and_pos :args (0))
(step t1188 (cl y$v3_1517448506_308_op (not (and y$v3_1517448506_308_op y$589))) :rule reordering :premises (t1187))
(step t1189 (cl (not y$v3_1517448506_312_op) (and y$v3_1517448506_308_op y$589)) :rule equiv1 :premises (a339))
(step t1190 (cl (and y$v3_1517448506_308_op y$589) (not y$v3_1517448506_312_op)) :rule reordering :premises (t1189))
(step t1191 (cl (not (and y$v3_1517448506_312_op y$597)) y$v3_1517448506_312_op) :rule and_pos :args (0))
(step t1192 (cl y$v3_1517448506_312_op (not (and y$v3_1517448506_312_op y$597))) :rule reordering :premises (t1191))
(step t1193 (cl (not y$v3_1517448506_316_op) (and y$v3_1517448506_312_op y$597)) :rule equiv1 :premises (a344))
(step t1194 (cl (and y$v3_1517448506_312_op y$597) (not y$v3_1517448506_316_op)) :rule reordering :premises (t1193))
(step t1195 (cl (not (and y$v3_1517448506_316_op y$605)) y$v3_1517448506_316_op) :rule and_pos :args (0))
(step t1196 (cl y$v3_1517448506_316_op (not (and y$v3_1517448506_316_op y$605))) :rule reordering :premises (t1195))
(step t1197 (cl (not y$v3_1517448506_320_op) (and y$v3_1517448506_316_op y$605)) :rule equiv1 :premises (a349))
(step t1198 (cl (and y$v3_1517448506_316_op y$605) (not y$v3_1517448506_320_op)) :rule reordering :premises (t1197))
(step t1199 (cl (not (and y$v3_1517448506_320_op y$707)) y$v3_1517448506_320_op) :rule and_pos :args (0))
(step t1200 (cl y$v3_1517448506_320_op (not (and y$v3_1517448506_320_op y$707))) :rule reordering :premises (t1199))
(step t1201 (cl (not y$v3_1517448506_371_op) (and y$v3_1517448506_320_op y$707)) :rule equiv1 :premises (a425))
(step t1202 (cl (and y$v3_1517448506_320_op y$707) (not y$v3_1517448506_371_op)) :rule reordering :premises (t1201))
(step t1203 (cl (not (and y$v3_1517448506_371_op y$880)) y$v3_1517448506_371_op) :rule and_pos :args (0))
(step t1204 (cl y$v3_1517448506_371_op (not (and y$v3_1517448506_371_op y$880))) :rule reordering :premises (t1203))
(step t1205 (cl (not y$v3_1517448506_445_op) (and y$v3_1517448506_371_op y$880)) :rule equiv1 :premises (a548))
(step t1206 (cl (and y$v3_1517448506_371_op y$880) (not y$v3_1517448506_445_op)) :rule reordering :premises (t1205))
(step t1207 (cl (not (and y$v3_1517448506_445_op y$v3_1517448506_529_op)) y$v3_1517448506_445_op) :rule and_pos :args (0))
(step t1208 (cl y$v3_1517448506_445_op (not (and y$v3_1517448506_445_op y$v3_1517448506_529_op))) :rule reordering :premises (t1207))
(step t1209 (cl (not y$v3_1517448506_530_op) (and y$v3_1517448506_445_op y$v3_1517448506_529_op)) :rule equiv1 :premises (a680))
(step t1210 (cl (and y$v3_1517448506_445_op y$v3_1517448506_529_op) (not y$v3_1517448506_530_op)) :rule reordering :premises (t1209))
(step t1211 (cl (not (and y$v3_1517448506_530_op y$v3_1517448506_614_op)) y$v3_1517448506_530_op) :rule and_pos :args (0))
(step t1212 (cl y$v3_1517448506_530_op (not (and y$v3_1517448506_530_op y$v3_1517448506_614_op))) :rule reordering :premises (t1211))
(step t1213 (cl y$v3_1517448506_530_op) :rule resolution :premises (t1212 t512))
(step t1214 (cl (and y$v3_1517448506_445_op y$v3_1517448506_529_op)) :rule resolution :premises (t1210 t1213))
(step t1215 (cl y$v3_1517448506_445_op) :rule resolution :premises (t1208 t1214))
(step t1216 (cl (and y$v3_1517448506_371_op y$880)) :rule resolution :premises (t1206 t1215))
(step t1217 (cl y$v3_1517448506_371_op) :rule resolution :premises (t1204 t1216))
(step t1218 (cl (and y$v3_1517448506_320_op y$707)) :rule resolution :premises (t1202 t1217))
(step t1219 (cl y$v3_1517448506_320_op) :rule resolution :premises (t1200 t1218))
(step t1220 (cl (and y$v3_1517448506_316_op y$605)) :rule resolution :premises (t1198 t1219))
(step t1221 (cl y$v3_1517448506_316_op) :rule resolution :premises (t1196 t1220))
(step t1222 (cl (and y$v3_1517448506_312_op y$597)) :rule resolution :premises (t1194 t1221))
(step t1223 (cl y$v3_1517448506_312_op) :rule resolution :premises (t1192 t1222))
(step t1224 (cl (and y$v3_1517448506_308_op y$589)) :rule resolution :premises (t1190 t1223))
(step t1225 (cl y$v3_1517448506_308_op) :rule resolution :premises (t1188 t1224))
(step t1226 (cl (and y$v3_1517448506_303_op y$581)) :rule resolution :premises (t1186 t1225))
(step t1227 (cl y$v3_1517448506_303_op) :rule resolution :premises (t1184 t1226))
(step t1228 (cl (and y$v3_1517448506_300_op y$571)) :rule resolution :premises (t1182 t1227))
(step t1229 (cl y$v3_1517448506_300_op) :rule resolution :premises (t1180 t1228))
(step t1230 (cl (and y$v3_1517448506_297_op y$566)) :rule resolution :premises (t1178 t1229))
(step t1231 (cl y$v3_1517448506_297_op) :rule resolution :premises (t1176 t1230))
(step t1232 (cl (and y$v3_1517448506_294_op y$561)) :rule resolution :premises (t1174 t1231))
(step t1233 (cl y$v3_1517448506_294_op) :rule resolution :premises (t1172 t1232))
(step t1234 (cl (and y$v3_1517448506_291_op y$556)) :rule resolution :premises (t1170 t1233))
(step t1235 (cl y$v3_1517448506_291_op) :rule resolution :premises (t1168 t1234))
(step t1236 (cl (and y$v3_1517448506_288_op y$551)) :rule resolution :premises (t1166 t1235))
(step t1237 (cl y$v3_1517448506_288_op) :rule resolution :premises (t1164 t1236))
(step t1238 (cl (and y$v3_1517448506_285_op y$546)) :rule resolution :premises (t1162 t1237))
(step t1239 (cl y$v3_1517448506_285_op) :rule resolution :premises (t1160 t1238))
(step t1240 (cl (and y$v3_1517448506_282_op y$541)) :rule resolution :premises (t1158 t1239))
(step t1241 (cl y$v3_1517448506_282_op) :rule resolution :premises (t1156 t1240))
(step t1242 (cl (and y$v3_1517448506_279_op y$536)) :rule resolution :premises (t1154 t1241))
(step t1243 (cl y$v3_1517448506_279_op) :rule resolution :premises (t1152 t1242))
(step t1244 (cl (and y$v3_1517448506_274_op y$531)) :rule resolution :premises (t1150 t1243))
(step t1245 (cl y$v3_1517448506_274_op) :rule resolution :premises (t1148 t1244))
(step t1246 (cl (and y$v3_1517448506_270_op y$522)) :rule resolution :premises (t1146 t1245))
(step t1247 (cl y$522) :rule resolution :premises (t1144 t1246))
(step t1248 (cl (not y$v3_1517448506_273_op)) :rule resolution :premises (t1142 t1247))
(step t1249 (cl (not (and y$f12 y$519))) :rule resolution :premises (t1140 t1248))
(step t1250 (cl (not (= (or y$v3_1517448506_271_op (not (not y$519))) (or y$v3_1517448506_271_op y$519))) (not (or y$v3_1517448506_271_op (not (not y$519)))) (or y$v3_1517448506_271_op y$519)) :rule equiv_pos2)
(step t1251 (cl (= y$v3_1517448506_271_op y$v3_1517448506_271_op)) :rule refl)
(step t1252 (cl (= (not (not y$519)) y$519)) :rule rare_rewrite :args ("bool-double-not-elim" y$519))
(step t1253 (cl (= (or y$v3_1517448506_271_op (not (not y$519))) (or y$v3_1517448506_271_op y$519))) :rule cong :premises (t1251 t1252))
(step t1254 (cl y$v3_1517448506_271_op (not (not y$519))) :rule equiv2 :premises (a294))
(step t1255 (cl (or y$v3_1517448506_271_op (not (not y$519))) (not y$v3_1517448506_271_op)) :rule or_neg :args (0))
(step t1256 (cl (or y$v3_1517448506_271_op (not (not y$519))) (not (not (not y$519)))) :rule or_neg :args (1))
(step t1257 (cl (or y$v3_1517448506_271_op (not (not y$519))) (or y$v3_1517448506_271_op (not (not y$519)))) :rule resolution :premises (t1254 t1255 t1256))
(step t1258 (cl (or y$v3_1517448506_271_op (not (not y$519)))) :rule contraction :premises (t1257))
(step t1259 (cl (or y$v3_1517448506_271_op y$519)) :rule resolution :premises (t1250 t1253 t1258))
(step t1260 (cl y$v3_1517448506_271_op y$519) :rule or :premises (t1259))
(step t1261 (cl y$519 y$v3_1517448506_271_op) :rule reordering :premises (t1260))
(step t1262 (cl (not y$v3_1517448506_271_op) (and y$a_small_resistance y$516)) :rule equiv1 :premises (a293))
(step t1263 (cl (and y$a_small_resistance y$516) (not y$v3_1517448506_271_op)) :rule reordering :premises (t1262))
(step t1264 (cl y$1278 (not (= y$n1s32 y$v3_1517448506_36_op))) :rule equiv2 :premises (a12))
(step t1265 (cl (not (and y$a_small_resistance y$516)) y$516) :rule and_pos :args (1))
(step t1266 (cl y$516 (not (and y$a_small_resistance y$516))) :rule reordering :premises (t1265))
(step t1267 (cl (and y$1275 y$1278) (not y$1275) (not y$1278)) :rule and_neg)
(step t1268 (cl y$v3_1517448506_38_op (not (and y$1275 y$1278))) :rule equiv2 :premises (a13))
(step t1269 (cl (and y$41 y$v3_1517448506_38_op) (not y$41) (not y$v3_1517448506_38_op)) :rule and_neg)
(step t1270 (cl (not y$41) (and y$41 y$v3_1517448506_38_op) (not y$v3_1517448506_38_op)) :rule reordering :premises (t1269))
(step t1271 (cl (not (and y$41 y$v3_1517448506_615_op)) y$41) :rule and_pos :args (0))
(step t1272 (cl y$41 (not (and y$41 y$v3_1517448506_615_op))) :rule reordering :premises (t1271))
(step t1273 (cl y$41) :rule resolution :premises (t1272 t510))
(step t1274 (cl y$id37_op (not (and y$41 y$v3_1517448506_38_op))) :rule equiv2 :premises (a14))
(step t1275 (cl (not y$id37_op) (not y$1283)) :rule equiv1 :premises (a15))
(step t1276 (cl (not y$1283) (not y$id37_op)) :rule reordering :premises (t1275))
(step t1277 (cl (not (= y$1283 y$prop)) y$1283 (not y$prop)) :rule equiv_pos1)
(step t1278 (cl y$1283 (not y$prop) (not (= y$1283 y$prop))) :rule reordering :premises (t1277))
(step t1279 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$prop) :rule and_pos :args (0))
(step t1280 (cl y$prop (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t1279))
(step t1281 (cl y$prop) :rule resolution :premises (t1280 t472))
(step t1282 (cl (not (= (= y$1284 (= y$prop y$1283)) (= y$1284 (= y$1283 y$prop)))) (not (= y$1284 (= y$prop y$1283))) (= y$1284 (= y$1283 y$prop))) :rule equiv_pos2)
(step t1283 (cl (= y$1284 y$1284)) :rule refl)
(step t1284 (cl (= (= y$prop y$1283) (= y$1283 y$prop))) :rule rare_rewrite :args ("eq-symm" y$prop y$1283))
(step t1285 (cl (= (= y$1284 (= y$prop y$1283)) (= y$1284 (= y$1283 y$prop)))) :rule cong :premises (t1283 t1284))
(step t1286 (cl (= y$1284 (= y$1283 y$prop))) :rule resolution :premises (t1282 t1285 a16))
(step t1287 (cl (not y$1284) (= y$1283 y$prop)) :rule equiv1 :premises (t1286))
(step t1288 (cl (= y$1283 y$prop) (not y$1284)) :rule reordering :premises (t1287))
(step t1289 (cl (not (and y$prop y$1284 y$1272 y$1343 y$1318)) y$1284) :rule and_pos :args (1))
(step t1290 (cl y$1284 (not (and y$prop y$1284 y$1272 y$1343 y$1318))) :rule reordering :premises (t1289))
(step t1291 (cl y$1284) :rule resolution :premises (t1290 t472))
(step t1292 (cl (= y$1283 y$prop)) :rule resolution :premises (t1288 t1291))
(step t1293 (cl y$1283) :rule resolution :premises (t1278 t1281 t1292))
(step t1294 (cl (not y$id37_op)) :rule resolution :premises (t1276 t1293))
(step t1295 (cl (not (and y$41 y$v3_1517448506_38_op))) :rule resolution :premises (t1274 t1294))
(step t1296 (cl (not y$v3_1517448506_38_op)) :rule resolution :premises (t1270 t1273 t1295))
(step t1297 (cl (not (and y$1275 y$1278))) :rule resolution :premises (t1268 t1296))
(step t1298 (cl (= y$516 (= y$n0s32 y$v3_1517448506_30_op))) :rule symm :premises (a292))
(step t1299 (cl (not y$516) (= y$n0s32 y$v3_1517448506_30_op)) :rule equiv1 :premises (t1298))
(step t1300 (cl (= y$n0s32 y$v3_1517448506_30_op) (not y$516)) :rule reordering :premises (t1299))
(step t1301 (cl (not (= (or y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op)))) (or y$1275 (= y$n0s32 y$v3_1517448506_30_op)))) (not (or y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op))))) (or y$1275 (= y$n0s32 y$v3_1517448506_30_op))) :rule equiv_pos2)
(step t1302 (cl (= y$1275 y$1275)) :rule refl)
(step t1303 (cl (= (not (not (= y$n0s32 y$v3_1517448506_30_op))) (= y$n0s32 y$v3_1517448506_30_op))) :rule rare_rewrite :args ("bool-double-not-elim" (= y$n0s32 y$v3_1517448506_30_op)))
(step t1304 (cl (= (or y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op)))) (or y$1275 (= y$n0s32 y$v3_1517448506_30_op)))) :rule cong :premises (t1302 t1303))
(step t1305 (cl y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule equiv2 :premises (a10))
(step t1306 (cl (or y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op)))) (not y$1275)) :rule or_neg :args (0))
(step t1307 (cl (or y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op)))) (not (not (not (= y$n0s32 y$v3_1517448506_30_op))))) :rule or_neg :args (1))
(step t1308 (cl (or y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op)))) (or y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op))))) :rule resolution :premises (t1305 t1306 t1307))
(step t1309 (cl (or y$1275 (not (not (= y$n0s32 y$v3_1517448506_30_op))))) :rule contraction :premises (t1308))
(step t1310 (cl (or y$1275 (= y$n0s32 y$v3_1517448506_30_op))) :rule resolution :premises (t1301 t1304 t1309))
(step t1311 (cl y$1275 (= y$n0s32 y$v3_1517448506_30_op)) :rule or :premises (t1310))
(step t1312 (cl (= y$n0s32 y$v3_1517448506_30_op)) :rule resolution :premises (t1075 t1087 t1101 t1112 t1114 t1125 t1137 t1139 t1249 t1261 t1263 t1264 t1266 t1267 t1297 t1300 t1311))
(step t1313 (cl (not y$1334) (not (= y$n0s32 y$v3_1517448506_30$next_op))) :rule equiv1 :premises (a840))
(step t1314 (cl (not (= y$n0s32 y$v3_1517448506_30$next_op)) (not y$1334)) :rule reordering :premises (t1313))
(step t1315 (cl (not (and y$1334 y$1337)) y$1334) :rule and_pos :args (0))
(step t1316 (cl y$1334 (not (and y$1334 y$1337))) :rule reordering :premises (t1315))
(step t1317 (cl y$1334) :rule resolution :premises (t1316 t589))
(step t1318 (cl (not (= y$n0s32 y$v3_1517448506_30$next_op))) :rule resolution :premises (t1314 t1317))
(step t1319 (cl (not (= (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule equiv_pos2)
(step t1320 (cl (= y$v3_1517448506_30$next_op y$v3_1517448506_30$next_op)) :rule refl)
(step t1321 (cl (= (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule cong :premises (t1320 t13))
(step t1322 (cl (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule resolution :premises (t1319 t1321 a839))
(step t1323 (cl (not (= (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule equiv_pos2)
(step t1324 (cl (= y$v3_1517448506_30_op y$v3_1517448506_30_op)) :rule refl)
(step t1325 (cl (= (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule cong :premises (t1324 t61))
(step t1326 (cl (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule resolution :premises (t1323 t1325 a9))
(step t1327 (cl (not (= y$v3_1517448506_31_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))) :rule resolution :premises (t103 t877 t907 t937 t943 a838 a837 a836 a835 a834 t947 t950 t953 a8 a7 a6 a5 a4 t955 t956 t977 t877 t907 t937 t943 t947 t950 t953 t23 t979 t980 t1049 t1312 t877 t907 t937 t1318 t943 t1322 a838 a834 t947 t950 t953 t1326 a8 a4))
(step t1328 (cl y$111) :rule resolution :premises (t70 t1327))
(step t1329 (cl (Extract_1_15_15_16 y$v_range)) :rule resolution :premises (t47 t1328))
(step t1330 (cl (Extract_1_15_15_16 y$v_range$next)) :rule resolution :premises (t45 t1329 t953 t950 t947 t943 t877 t907 t937))
(step t1331 (cl y$1320) :rule resolution :premises (t24 t1330))
(step t1332 (cl (not (= (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))))) (not (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) (or (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule equiv_pos2)
(step t1333 (cl (= (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule refl)
(step t1334 (cl (= (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))))) :rule refl)
(step t1335 (cl (= (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))))) :rule refl)
(step t1336 (cl (= (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))))) :rule refl)
(step t1337 (cl (= (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))))) :rule refl)
(step t1338 (cl (= (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))))) :rule refl)
(step t1339 (cl (= (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))))) :rule refl)
(step t1340 (cl (= (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule refl)
(step t1341 (cl (= (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))))) :rule cong :premises (t982 t998 t997 t1333 t1334 t1335 t1336 t994 t993 t992 t991 t990 t989 t988 t987 t986 t1337 t1338 t1339 t983 t1340))
(step t1342 (cl (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule and_neg)
(step t1343 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule implies_neg1)
(anchor :step t1344)
(assume t1344.a0 (not (= y$n0s32 y$v3_1517448506_30$next_op)))
(assume t1344.a1 (= y$n0s32 y$v3_1517448506_30_op))
(assume t1344.a2 (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))
(assume t1344.a3 (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)))
(assume t1344.a4 (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)))
(assume t1344.a5 (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)))
(assume t1344.a6 (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)))
(assume t1344.a7 (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)))
(assume t1344.a8 (= y$v_range (ite y$f09 y$132 y$v_range)))
(assume t1344.a9 (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)))
(assume t1344.a10 (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t1344.a11 (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)))
(assume t1344.a12 (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t1344.a13 (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)))
(assume t1344.a14 (= y$v_range$next y$v_range$next_rhs_op))
(assume t1344.a15 (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)))
(assume t1344.a16 (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)))
(assume t1344.a17 (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)))
(assume t1344.a18 (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)))
(assume t1344.a19 (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))
(step t1344.t0 (cl (= (= (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) false) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule equiv_simplify)
(step t1344.t1 (cl (not (= (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) false)) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule equiv1 :premises (t1344.t0))
(step t1344.t2 (cl (= y$n0s16 y$n0s16)) :rule refl)
(step t1344.t3 (cl (= (ite y$f09 y$132 y$v_range) y$v3_1517448506_80_op)) :rule symm :premises (t953))
(step t1344.t4 (cl (= (ite y$f10 y$127 y$v3_1517448506_80_op) y$v3_1517448506_81_op)) :rule symm :premises (t950))
(step t1344.t5 (cl (= (ite y$f11 y$n3s16 y$v3_1517448506_81_op) y$v_range$next_rhs_op)) :rule symm :premises (t947))
(step t1344.t6 (cl (= y$v_range$next_rhs_op y$v_range$next)) :rule symm :premises (t943))
(step t1344.t7 (cl (= y$v_range y$v_range$next)) :rule trans :premises (t877 t1344.t3 t907 t1344.t4 t937 t1344.t5 t1344.t6))
(step t1344.t8 (cl (= (Concat_32_16_16 y$n0s16 y$v_range) (Concat_32_16_16 y$n0s16 y$v_range$next))) :rule cong :premises (t1344.t2 t1344.t7))
(step t1344.t9 (cl (= (Concat_32_16_16 y$n0s16 y$v_range$next) y$w$6$next_op)) :rule symm :premises (t1344.a15))
(step t1344.t10 (cl (= y$w$6_op y$w$6$next_op)) :rule trans :premises (t1344.a7 t1344.t8 t1344.t9))
(step t1344.t11 (cl (= y$w$6$next_op y$w$6_op)) :rule symm :premises (t1344.t10))
(step t1344.t12 (cl (= (BitWiseNot_32_32 y$w$6$next_op) (BitWiseNot_32_32 y$w$6_op))) :rule cong :premises (t1344.t11))
(step t1344.t13 (cl (= (BitWiseNot_32_32 y$w$6_op) y$s$199_op)) :rule symm :premises (t1344.a6))
(step t1344.t14 (cl (= y$s$199$next_op y$s$199_op)) :rule trans :premises (t1344.a16 t1344.t12 t1344.t13))
(step t1344.t15 (cl (= y$n16s32 y$n16s32)) :rule refl)
(step t1344.t16 (cl (= (ShiftR_32_32_32 y$s$199$next_op y$n16s32) (ShiftR_32_32_32 y$s$199_op y$n16s32))) :rule cong :premises (t1344.t14 t1344.t15))
(step t1344.t17 (cl (= (ShiftR_32_32_32 y$s$199_op y$n16s32) y$v3_1517448506_32_op)) :rule symm :premises (t1344.a5))
(step t1344.t18 (cl (= y$v3_1517448506_32$next_op y$v3_1517448506_32_op)) :rule trans :premises (t1344.a17 t1344.t16 t1344.t17))
(step t1344.t19 (cl (= (BitWiseNot_32_32 y$v3_1517448506_32$next_op) (BitWiseNot_32_32 y$v3_1517448506_32_op))) :rule cong :premises (t1344.t18))
(step t1344.t20 (cl (= (BitWiseNot_32_32 y$v3_1517448506_32_op) y$s$198_op)) :rule symm :premises (t1344.a4))
(step t1344.t21 (cl (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$s$198_op)) :rule symm :premises (t1344.a3))
(step t1344.t22 (cl (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule symm :premises (t1344.t21))
(step t1344.t23 (cl (= (ite y$111 y$s$198_op y$v3_1517448506_31_op) y$v3_1517448506_30_op)) :rule symm :premises (t1326))
(step t1344.t24 (cl (= y$v3_1517448506_30_op y$n0s32)) :rule symm :premises (t1344.a1))
(step t1344.t25 (cl (= y$s$198$next_op y$n0s32)) :rule trans :premises (t1344.a18 t1344.t19 t1344.t20 t1344.t22 t1344.t23 t1344.t24))
(step t1344.t26 (cl (= (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op) y$v3_1517448506_30$next_op)) :rule symm :premises (t1322))
(step t1344.t27 (cl (= (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) (= y$n0s32 y$v3_1517448506_30$next_op))) :rule cong :premises (t1344.t25 t1344.t26))
(step t1344.t28 (cl (= (= (= y$n0s32 y$v3_1517448506_30$next_op) false) (not (= y$n0s32 y$v3_1517448506_30$next_op)))) :rule equiv_simplify)
(step t1344.t29 (cl (= (= y$n0s32 y$v3_1517448506_30$next_op) false) (not (not (= y$n0s32 y$v3_1517448506_30$next_op)))) :rule equiv2 :premises (t1344.t28))
(step t1344.t30 (cl (not (not (not (= y$n0s32 y$v3_1517448506_30$next_op)))) (= y$n0s32 y$v3_1517448506_30$next_op)) :rule not_not)
(step t1344.t31 (cl (= (= y$n0s32 y$v3_1517448506_30$next_op) false) (= y$n0s32 y$v3_1517448506_30$next_op)) :rule resolution :premises (t1344.t29 t1344.t30))
(step t1344.t32 (cl (= (= y$n0s32 y$v3_1517448506_30$next_op) false)) :rule resolution :premises (t1344.t31 t1318))
(step t1344.t33 (cl (= (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)) false)) :rule trans :premises (t1344.t27 t1344.t32))
(step t1344.t34 (cl (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule resolution :premises (t1344.t1 t1344.t33))
(step t1344 (cl (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule subproof :discharge (t1344.a0 t1344.a1 t1344.a2 t1344.a3 t1344.a4 t1344.a5 t1344.a6 t1344.a7 t1344.a8 t1344.a9 t1344.a10 t1344.a11 t1344.a12 t1344.a13 t1344.a14 t1344.a15 t1344.a16 t1344.a17 t1344.a18 t1344.a19))
(step t1345 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (= y$n0s32 y$v3_1517448506_30$next_op))) :rule and_pos :args (0))
(step t1346 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$n0s32 y$v3_1517448506_30_op)) :rule and_pos :args (1))
(step t1347 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule and_pos :args (2))
(step t1348 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule and_pos :args (3))
(step t1349 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) :rule and_pos :args (4))
(step t1350 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) :rule and_pos :args (5))
(step t1351 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) :rule and_pos :args (6))
(step t1352 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) :rule and_pos :args (7))
(step t1353 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v_range (ite y$f09 y$132 y$v_range))) :rule and_pos :args (8))
(step t1354 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) :rule and_pos :args (9))
(step t1355 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (10))
(step t1356 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) :rule and_pos :args (11))
(step t1357 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (12))
(step t1358 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) :rule and_pos :args (13))
(step t1359 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v_range$next y$v_range$next_rhs_op)) :rule and_pos :args (14))
(step t1360 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) :rule and_pos :args (15))
(step t1361 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) :rule and_pos :args (16))
(step t1362 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) :rule and_pos :args (17))
(step t1363 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) :rule and_pos :args (18))
(step t1364 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) :rule and_pos :args (19))
(step t1365 (cl (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule resolution :premises (t1344 t1345 t1346 t1347 t1348 t1349 t1350 t1351 t1352 t1353 t1354 t1355 t1356 t1357 t1358 t1359 t1360 t1361 t1362 t1363 t1364))
(step t1366 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule reordering :premises (t1365))
(step t1367 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule contraction :premises (t1366))
(step t1368 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule resolution :premises (t1343 t1367))
(step t1369 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule implies_neg2)
(step t1370 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule resolution :premises (t1368 t1369))
(step t1371 (cl (=> (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule contraction :premises (t1370))
(step t1372 (cl (not (and (not (= y$n0s32 y$v3_1517448506_30$next_op)) (= y$n0s32 y$v3_1517448506_30_op) (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op)) (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op)) (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32)) (= y$s$199_op (BitWiseNot_32_32 y$w$6_op)) (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range)) (= y$v_range (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range)) (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op)) (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op)) (= y$v_range$next y$v_range$next_rhs_op) (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next)) (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op)) (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32)) (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op)) (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule implies :premises (t1371))
(step t1373 (cl (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule resolution :premises (t1342 t1372))
(step t1374 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (not (= y$n0s32 y$v3_1517448506_30$next_op))))) :rule or_neg :args (0))
(step t1375 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$n0s32 y$v3_1517448506_30_op)))) :rule or_neg :args (1))
(step t1376 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule or_neg :args (2))
(step t1377 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))))) :rule or_neg :args (3))
(step t1378 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))))) :rule or_neg :args (4))
(step t1379 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))))) :rule or_neg :args (5))
(step t1380 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))))) :rule or_neg :args (6))
(step t1381 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))))) :rule or_neg :args (7))
(step t1382 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v_range (ite y$f09 y$132 y$v_range))))) :rule or_neg :args (8))
(step t1383 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))))) :rule or_neg :args (9))
(step t1384 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule or_neg :args (10))
(step t1385 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))))) :rule or_neg :args (11))
(step t1386 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule or_neg :args (12))
(step t1387 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))))) :rule or_neg :args (13))
(step t1388 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v_range$next y$v_range$next_rhs_op)))) :rule or_neg :args (14))
(step t1389 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))))) :rule or_neg :args (15))
(step t1390 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))))) :rule or_neg :args (16))
(step t1391 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))))) :rule or_neg :args (17))
(step t1392 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))))) :rule or_neg :args (18))
(step t1393 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule or_neg :args (19))
(step t1394 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (not (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule or_neg :args (20))
(step t1395 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule resolution :premises (t1373 t1374 t1375 t1376 t1377 t1378 t1379 t1380 t1381 t1382 t1383 t1384 t1385 t1386 t1387 t1388 t1389 t1390 t1391 t1392 t1393 t1394))
(step t1396 (cl (or (not (not (= y$n0s32 y$v3_1517448506_30$next_op))) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule contraction :premises (t1395))
(step t1397 (cl (or (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))))) :rule resolution :premises (t1332 t1341 t1396))
(step t1398 (cl (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$n0s32 y$v3_1517448506_30_op)) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule or :premises (t1397))
(step t1399 (cl (not (= y$n0s32 y$v3_1517448506_30_op)) (= y$n0s32 y$v3_1517448506_30$next_op) (not (= y$w$6_op (Concat_32_16_16 y$n0s16 y$v_range))) (not (= y$s$199_op (BitWiseNot_32_32 y$w$6_op))) (not (= y$v3_1517448506_32_op (ShiftR_32_32_32 y$s$199_op y$n16s32))) (not (= y$s$198_op (BitWiseNot_32_32 y$v3_1517448506_32_op))) (not (= y$v3_1517448506_30_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v3_1517448506_80_op (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_81_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v_range$next_rhs_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$v_range$next y$v_range$next_rhs_op)) (not (= y$w$6$next_op (Concat_32_16_16 y$n0s16 y$v_range$next))) (not (= y$s$199$next_op (BitWiseNot_32_32 y$w$6$next_op))) (not (= y$v3_1517448506_32$next_op (ShiftR_32_32_32 y$s$199$next_op y$n16s32))) (not (= y$s$198$next_op (BitWiseNot_32_32 y$v3_1517448506_32$next_op))) (not (= y$v3_1517448506_30$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op))) (not (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) (not (= y$v_range (ite y$f09 y$132 y$v_range))) (not (= y$v3_1517448506_80_op (ite y$f10 y$127 y$v3_1517448506_80_op))) (not (= y$v3_1517448506_81_op (ite y$f11 y$n3s16 y$v3_1517448506_81_op))) (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule reordering :premises (t1398))
(step t1400 (cl (= y$s$198_op (ite y$111 y$s$198_op y$v3_1517448506_31_op))) :rule resolution :premises (t955 t1328))
(step t1401 (cl (not (= y$s$198$next_op (ite y$1320 y$s$198$next_op y$v3_1517448506_31$next_op)))) :rule resolution :premises (t1399 t1312 t1318 a4 a5 a6 a7 t1326 t953 t950 t947 t943 a834 a835 a836 a837 t1322 t1400 t877 t907 t937))
(step t1402 (cl) :rule resolution :premises (t23 t1331 t1401))
)
